<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the number of axes of rotational symmetry a given solid has</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>How many axes of rotational symmetry does the solid have?</p>

[[jsxgraph]]

var board = JXG.JSXGraph.initBoard(divid,{
    boundingbox:[-7,6,8,-6],
    axis:false,
    showCopyright:false
});

var solid={#solid_id#};

/* =======================
   CUBE
======================= */

if(solid==1){

var A=board.create('point',[-2,2],{visible:false});
var B=board.create('point',[2,2],{visible:false});
var C=board.create('point',[2,-2],{visible:false});
var D=board.create('point',[-2,-2],{visible:false});

var dx=1.5,dy=1.3;

var E=board.create('point',[-2+dx,2+dy],{visible:false});
var F=board.create('point',[2+dx,2+dy],{visible:false});
var G=board.create('point',[2+dx,-2+dy],{visible:false});
var H=board.create('point',[-2+dx,-2+dy],{visible:false});

board.create('polygon',[A,B,C,D]);
board.create('polygon',[E,F,G,H]);

board.create('segment',[A,E]);
board.create('segment',[B,F]);
board.create('segment',[C,G]);
board.create('segment',[D,H]);

}

/* =======================
   CUBOID
======================= */

else if(solid==2){

var A=board.create('point',[-3,1.5],{visible:false});
var B=board.create('point',[3,1.5],{visible:false});
var C=board.create('point',[3,-1.5],{visible:false});
var D=board.create('point',[-3,-1.5],{visible:false});

var dx=2,dy=1.2;

var E=board.create('point',[-3+dx,1.5+dy],{visible:false});
var F=board.create('point',[3+dx,1.5+dy],{visible:false});
var G=board.create('point',[3+dx,-1.5+dy],{visible:false});
var H=board.create('point',[-3+dx,-1.5+dy],{visible:false});

board.create('polygon',[A,B,C,D]);
board.create('polygon',[E,F,G,H]);

board.create('segment',[A,E]);
board.create('segment',[B,F]);
board.create('segment',[C,G]);
board.create('segment',[D,H]);

}

/* =======================
   CONE
======================= */

else if(solid==3){

// Apex
var A = board.create('point',[0,3],{visible:false});

// Approximate circular base
var pts = [];

for(var i=0;i<24;i++){
    var t = 2*Math.PI*i/24;
    pts.push(board.create('point',[
        2*Math.cos(t),
        -2 + 0.45*Math.sin(t)
    ],{visible:false}));
}

board.create('polygon',pts,{
    fillColor:'#dbeafe',
    fillOpacity:0.5,
    borders:{strokeWidth:2}
});

// Join apex to left and right extremes
board.create('segment',[A,pts[12]],{strokeWidth:2});
board.create('segment',[A,pts[0]],{strokeWidth:2});

}

/* =======================
   CYLINDER
======================= */

else if(solid==4){

var top=[];
var bottom=[];

for(var i=0;i<24;i++){

    var t=2*Math.PI*i/24;

    top.push(
        board.create('point',[
            2*Math.cos(t),
            2+0.45*Math.sin(t)
        ],{visible:false})
    );

    bottom.push(
        board.create('point',[
            2*Math.cos(t),
            -2+0.45*Math.sin(t)
        ],{visible:false})
    );

}

board.create('polygon',top,{
    fillColor:'#dbeafe',
    fillOpacity:0.3,
    borders:{strokeWidth:2}
});

board.create('polygon',bottom,{
    fillColor:'#dbeafe',
    fillOpacity:0.3,
    borders:{strokeWidth:2}
});

// Vertical edges
board.create('segment',[top[0],bottom[0]],{strokeWidth:2});
board.create('segment',[top[12],bottom[12]],{strokeWidth:2});

}

/* =======================
   TETRAHEDRON
======================= */

else if(solid==5){

var A=board.create('point',[0,3],{visible:false});
var B=board.create('point',[-3,-2],{visible:false});
var C=board.create('point',[3,-2],{visible:false});
var D=board.create('point',[0,0],{visible:false});

board.create('polygon',[A,B,C]);

board.create('segment',[A,D]);
board.create('segment',[B,D]);
board.create('segment',[C,D]);

}

/* =======================
   TRIANGULAR PRISM
======================= */

else{

var A=board.create('point',[-3,2],{visible:false});
var B=board.create('point',[-4,-2],{visible:false});
var C=board.create('point',[-1,-2],{visible:false});

var D=board.create('point',[2,3],{visible:false});
var E=board.create('point',[1,-1],{visible:false});
var F=board.create('point',[4,-1],{visible:false});

board.create('polygon',[A,B,C]);
board.create('polygon',[D,E,F]);

board.create('segment',[A,D]);
board.create('segment',[B,E]);
board.create('segment',[C,F]);

}

[[/jsxgraph]]

<p>Number of axes of rotational symmetry:
[[input:ans1]][[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>The solid shown is a <strong>{@solid_name@}</strong>.</p>
<p>It has <strong>{@num_axes@}</strong> axis (or axes) of rotational symmetry.</p>
<p>An <strong>axis of rotational symmetry</strong> is a line about which a solid can be rotated so that it appears unchanged.</p>
<p> Therefore, the correct answer is
\[
\boxed{{@num_axes@}}
\] </p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1 = Cube
   2 = Cuboid
   3 = Cone
   4 = Cylinder
   5 = Regular tetrahedron
   6 = Regular triangular prism
*/

solid_id : rand(6)+1;

/* Number of axes of rotational symmetry */

num_axes :
if solid_id=1 then 13
elseif solid_id=2 then 3
elseif solid_id=3 then 1
elseif solid_id=4 then 1
elseif solid_id=5 then 4
else 4;

/* Name of solid */

solid_name :
if solid_id=1 then "cube"
elseif solid_id=2 then "cuboid"
elseif solid_id=3 then "cone"
elseif solid_id=4 then "cylinder"
elseif solid_id=5 then "regular tetrahedron"
else "regular triangular prism";

/* Teacher answer */

ta : num_axes;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta@}</p>
<p>{@solid_id@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct number of axes</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1283928561</deployedseed>
    <deployedseed>558342849</deployedseed>
    <deployedseed>1582995902</deployedseed>
    <deployedseed>1566718150</deployedseed>
    <deployedseed>220809435</deployedseed>
  </question>
</quiz>
