<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Mathematical computations using reciprocals</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<h3 style="text-align:center;">Solve the following using reciprocals</h3>


<ol type="a">
<li>\( {@a_num@} \div \frac{{@a_den@}}{{@a_num@}} \)</li>
<li>\( {@b_coeff@}x = {@b_rhs@} \)</li>
<li>\( \frac{{@c_num@}}{{@c_den@}} x = {@c_rhs@} \)</li>
<li>\( {@dec1@} \div {@dec2@} \)</li>
</ol>


<p>
(a) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]<br>
(b) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]<br>
(c) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]<br>
(d) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Dividing by a number is the same as multiplying by its reciprocal.</p>


<ol type="i">
<li>First, replace division by multiplication with the reciprocal.</li>
<li>Then simplify the expression.</li>
<li>Finding the value of\(x\) is like dividing the value after the equal sign by the reciprocal of the coefficient of \(x\) </li>
</ol>


<ol type="a">
<li>\( {@a_num@} \div \frac{{@a_den@}}{{@a_num@}} = {@a_num@} \times  \frac{{@a_num@}}{{@a_den@}}\) <br>
\(= {@ta_a@}\).<br>
</li>
<li>To find the value of \(x\) in \( {@b_coeff@}x = {@b_rhs@} \), first find the reciprocal of the \(x\) co-efficient then multiply it by the value after the equal sign. <br>
\(x= \frac{1}{ {@b_coeff@}} \times  {@b_rhs@}\)<br>
Therefore,  \(x = {@ta_b@}\).<br></li>
<li>Finding \(x\) in \( \frac{{@c_num@}}{{@c_den@}} x = {@c_rhs@} \), we have:<br>
 \(x = \frac{{@c_den@}}{{@c_num@}} \times  {@c_rhs@} \)<br>         
Hence, \(x = {@ta_c@}\).</li>
<li>\( {@dec1@} \div {@dec2@}  = \frac{{@d1@}}{{@10@}} \div  \frac{{@d2@}}{{@10@}}\)<br>
   \( = \frac{{@d1@}}{{@10@}} \times \frac{{@10@}}{{@d2@}}\)   
<br> \(= {@ta_d@}\)</li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* (a) Integer divided by a fraction */
a_num : rand(8)+2; /* 2–9 */
a_den : rand_with_prohib(2,9,[a_num]);
frac_a : a_den/a_num; /* gives form like 6 ÷ (2/3) */
ta_a : a_num*(a_num/a_den);
/*mistake*/
mistake_ta1: 1/a_den;


/* (b) Linear equation 4x = 12 type */
b_coeff : rand(8)+2; /* 2–9 */
b_rhs : b_coeff*(rand(6)+2);
ta_b : b_rhs/b_coeff;
mistake_ta2:b_rhs*b_coeff;


/* (c) Fraction times x equals integer */
c_num : rand(4)+1; /* 1–5 */
c_den : rand_with_prohib(2,6,[c_num]);
c_rhs : (rand(6)+2);
ta_c : c_rhs*(c_den/c_num);
mistake_ta3:c_rhs*(c_num/c_den);


/* (d) Decimal division */
d1 : rand(8)+1; /* 1–9 */
d2 : rand_with_prohib(2,11,[d1]); 
dec1 : float(d1/10);
dec2 : float(d2/10);
ta_d : dec1/dec2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta_a@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_a</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_b</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_c</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta_d</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct computation of the expression</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_a</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the incorrect value</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have simplified {@a_num@} with the denominator of \(\frac{{@a_den@}}{{@a_num@}}\). Remember division does not cancel out instead you multiply {@a_num@} by the reciprocal of \(\frac{{@a_den@}}{{@a_num@}}\). </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct computation</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_b</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect value of x</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mistake_ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You multiplied {@b_rhs@} by {@b_coeff@} instead of multiplying it by the reciprocal of  {@b_coeff@}</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct computation</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta_c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect value of x</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mistake_ta3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You failed to take the reciprocal of the coefficient fraction \(\frac{{@c_num@}}{{@c_den@}}\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct computation</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta_d</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
