<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Applying logarithms to a power</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>During a Chemistry laboratory experiment at <strong>Kiharu Comprehensive School</strong>, a chemical sample contained \({@base@}^{@exp@}\) units of <strong>substance X</strong>.</p>

<p>By applying logarithms, write \({@base@}^{@exp@}\) in terms of \(\log {@base@}\).</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
We use the power law of logarithms:
</p>

<p>
\[
\log(a^n) = n\log a
\]
</p>

<p>
Applying this to the expression:
</p>

<p>
\[
\log({@base@}^{@exp@}) = {@exp@}\log {@base@}
\]
</p>

<p>
So the correct answer is:
</p>

<p>
\[
\boxed{{@exp@}\log {@base@}}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random base and exponent */
base: rand([2,3,5,7,10]);
exp: rand([2,3,4,5]);

/* Correct answer using power law of logarithms */
correct: exp * log(base);

/* Common student mistakes */
m1: log(base) + log(exp);     /* adds logs incorrectly */
m2: log(base)^exp;            /* raises log to a power */
m3: log(base);                /* ignores the exponent */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Base: {@base@}</p>
<p>Exponent: {@exp@}</p>
<p>Correct form: {@correct@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"> This question assesses learners’ ability to apply the power law of logarithms by rewriting the logarithm of a power in terms of a single logarithm.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correct</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>adds logs incorrectly</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You added the logarithms. Kindly use the correct law to solve the equation</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>raises log to a power </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You raised the logarithm to a power.
The exponent should multiply the logarithm.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>ignores the exponent</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>You ignored the exponent.
The exponent becomes a coefficient in front of the logarithm.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1450243429</deployedseed>
    <deployedseed>867272340</deployedseed>
    <deployedseed>532483208</deployedseed>
    <deployedseed>443211014</deployedseed>
    <deployedseed>900463639</deployedseed>
    <deployedseed>1257927121</deployedseed>
    <deployedseed>1627459712</deployedseed>
    <deployedseed>1646372785</deployedseed>
    <deployedseed>712381604</deployedseed>
    <deployedseed>831400626</deployedseed>
    <deployedseed>1685445198</deployedseed>
    <deployedseed>2083318231</deployedseed>
    <deployedseed>532055375</deployedseed>
    <deployedseed>1008870816</deployedseed>
    <deployedseed>778190365</deployedseed>
    <deployedseed>1739805868</deployedseed>
    <deployedseed>665392991</deployedseed>
    <deployedseed>1888124032</deployedseed>
    <deployedseed>365199520</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>correct</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
