<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the value of unknown number given the sum of the number and its reciprocal</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> The sum of a number and its reciprocal is \( {@sum@} \). Find the number.</p>
<p>Number = [[input:ans1]] or  [[input:ans2]] [[validation:ans1]] [[feedback:prt1]] [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p> Let the number be \( x \) such that<br>
\( x + \frac{1}{x} = {@sum@} \)<br>
Next multiply both sides by \( x \): \( x^2 + 1 = {@sum@}x \)<br>
\( x^2 - {@sum@}x + 1 = 0 \)<br>
Solving the quadratic: \( x = {@ta1@} \) or \( x = {@ta1_alt@} \)<br>
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Problem 1: x + 1/x = sum */
/* Choose sum so solutions are nice: sum = n + 1/n for some integer n */
n: rand(84) + 2;
sum: n + 1/n;
ta1: n;
ta1_alt: 1/n;
/*mistake-student used the sum as the reciprocal of the number*/
mistake_ta1:denom(sum)/num(sum);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>first value- {@ta1@}</p>
<p>sum- {@sum@}</p>
<p>second value-{@ta1_alt@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta1_alt</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect value of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mistake_ta1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You treated the sum of a number and its reciprocal as if it were the reciprocal itself, which led to an incorrect value for the number. Review the worked solution to see the correct method.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct value of the unknown</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta1_alt</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect value of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have incorrectly calculated the value of the unknown number</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>621370435</deployedseed>
    <deployedseed>54393411</deployedseed>
    <deployedseed>230575001</deployedseed>
    <deployedseed>2090693206</deployedseed>
    <deployedseed>2101808052</deployedseed>
    <deployedseed>2062230551</deployedseed>
    <deployedseed>104164921</deployedseed>
    <deployedseed>519798144</deployedseed>
    <deployedseed>31086005</deployedseed>
    <deployedseed>420242876</deployedseed>
    <deployedseed>429616471</deployedseed>
    <deployedseed>1066912499</deployedseed>
    <deployedseed>603426625</deployedseed>
    <deployedseed>649808487</deployedseed>
    <deployedseed>1669490696</deployedseed>
    <deployedseed>907241125</deployedseed>
    <deployedseed>462387202</deployedseed>
    <deployedseed>416564543</deployedseed>
    <deployedseed>706381964</deployedseed>
    <deployedseed>258840972</deployedseed>
    <deployedseed>687283990</deployedseed>
    <deployedseed>2093575940</deployedseed>
    <deployedseed>67775817</deployedseed>
    <deployedseed>474702097</deployedseed>
    <deployedseed>987062261</deployedseed>
    <deployedseed>88449811</deployedseed>
    <deployedseed>1912893965</deployedseed>
    <deployedseed>1002272282</deployedseed>
    <deployedseed>1414982366</deployedseed>
    <deployedseed>525842160</deployedseed>
    <deployedseed>1520324563</deployedseed>
    <deployedseed>322074896</deployedseed>
    <deployedseed>704786658</deployedseed>
    <deployedseed>667993169</deployedseed>
    <deployedseed>1431793384</deployedseed>
    <deployedseed>1600217741</deployedseed>
    <deployedseed>484700349</deployedseed>
    <deployedseed>1465531994</deployedseed>
    <deployedseed>456550722</deployedseed>
    <deployedseed>875772933</deployedseed>
    <deployedseed>521036749</deployedseed>
    <deployedseed>617860328</deployedseed>
    <deployedseed>1917821901</deployedseed>
    <deployedseed>1879595915</deployedseed>
    <deployedseed>258349044</deployedseed>
    <deployedseed>570771862</deployedseed>
    <deployedseed>1870058577</deployedseed>
    <deployedseed>1016546567</deployedseed>
    <deployedseed>810583488</deployedseed>
    <deployedseed>1662220161</deployedseed>
    <deployedseed>576150696</deployedseed>
    <deployedseed>1624922665</deployedseed>
    <deployedseed>1913403583</deployedseed>
    <deployedseed>1413810922</deployedseed>
  </question>
</quiz>
