<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Using division to find reciprocal of the given numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Using division, find the reciprocal of each of the following numbers. </p>

<ol style="list-style-type: lower-alpha;">
<li>\({@a@}\) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li>
<li>\({@b@}\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>
<li>\({@c@}\) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]] </li>
<li>\({@d@}\) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]] </li>
<li>\({@e@}\) [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</li>
<li>\({@f@}\) [[input:ans6]] [[validation:ans6]]  [[feedback:prt6]]</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>Using division to find the reciprocal of a real number is dividing 1 by the number (1 รท number)</p>
<p>Therefore, the reciprocal of the given numbers are:</p>

<p><b>a)</b></p>
<p>  \(\begin{align}{@a@} =& 1 \div {@a@}\\ =& {@ta1@} \end{align}\)</p><br>

<p><b>b)</b></p>
<p> \(\begin{align}  {@b@} =& 1 \div {@b@}\\ =& {@ta2@}  \end{align}\) </p><br>

<p><b>c)</b></p> <p>\(\begin{align}{@c@} =& 1 \div {@c@}\\ =& {@ta3@}  \end{align}\) </p><br>

<p><b>d)</b></p> <p>\(\begin{align}{@d@} =& 1 \div {@d@}\\ =& {@ta4@}  \end{align}\) </p><br>
<p><b>e)</b></p> <p>\(\begin{align}{@e@} =& 1 \div {@e@}\\ =& {@ta5@}  \end{align}\)</p><br>
<p><b>f)</b></p><p> \(\begin{align}{@f@} =& 1 \div {@f@}\\ =& {@ta6@}  \end{align}\) </p><br>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Positive integer: 1 to 20 */
a: rand(20) + 1;

/* Negative integer: -1 to -20 */
b: -(rand(20) + 1);

/* Decimal between 0.1 and 5 with 1 decimal place */
c: (rand(50) + 1)/10;

/* Negative decimal between -0.1 and -5 with 1 decimal place */
d: -(rand(50) + 1)/10;

/* Positive fraction a/b where a < b */
num1: rand(8) + 2;
den1: num1 + rand(8) + 1;
e: num1/den1;

/* Negative fraction a/b where a < b */
num2: rand(8) + 2;
den2: num2 + rand(8) + 1;
f: -num2/den2;

/* Calculate reciprocals using division: 1 รท number */
ta1: 1/a;
ta2: 1/b;
ta3: 1/c;
ta4: 1/d;
ta5: 1/e;
ta6: 1/f;
/*mistakes*/
ta1_mistake:-a;
ta2_mistake:-1*(b);
ta3_mistake:-c;
ta4_mistake:-1*(d);
ta5_mistake:-e;
ta6_mistake:-1*(f);
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta1@}</p>
<p> ({@a@}) </p>
<p>{@ta2@}</p>
<p> ({@b@}) </p>

<p>{@ta3@}</p>
<p> ({@c@}) </p>

<p>{@ta4@}</p>
<p> ({@d@}) </p>
<p>{@ta5@}</p>
<p> ({@e@}) </p>
<p>{@ta6@}</p>
<p> ({@f@}) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct reciprocal using long division</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect reciprocals</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>a</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not use division correctly when finding the reciprocal of a number. Remember: the reciprocal of a number is <em>1 divided by the number</em>, not the number divided by 1. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1_mistake</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have confused the reciprocal of a number with its opposite (changing the sign). Remember, the reciprocal is <em>1 divided by the number</em>, and it keeps the same sign as the original number. Please review the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>b</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not use division correctly when finding the reciprocal of a number. Remember: the reciprocal of a number is <em>1 divided by the number</em>, not the number divided by 1. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2_mistake</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have confused the reciprocal of a number with its opposite (changing the sign). Remember, the reciprocal is <em>1 divided by the number</em>, and it keeps the same sign as the original number. Please review the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect reciprocal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>c</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not use division correctly when finding the reciprocal of a number. Remember: the reciprocal of a number is <em>1 divided by the number</em>, not the number divided by 1. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for incorrect reciprocal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3_mistake</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have confused the reciprocal of a number with its opposite (changing the sign). Remember, the reciprocal is <em>1 divided by the number</em>, and it keeps the same sign as the original number. Please review the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>d</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not use division correctly when finding the reciprocal of a number. Remember: the reciprocal of a number is <em>1 divided by the number</em>, not the number divided by 1. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4_mistake</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have confused the reciprocal of a number with its opposite (changing the sign). Remember, the reciprocal is <em>1 divided by the number</em>, and it keeps the same sign as the original number. Please review the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct reciprocal</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect reciprocal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>e</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not use division correctly when finding the reciprocal of a number. Remember: the reciprocal of a number is <em>1 divided by the number</em>, not the number divided by 1. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt5-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for incorrect reciprocal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5_mistake</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have confused the reciprocal of a number with its opposite (changing the sign). Remember, the reciprocal is <em>1 divided by the number</em>, and it keeps the same sign as the original number. Please review the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for incorrect reciprocal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>f</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not use division correctly when finding the reciprocal of a number. Remember: the reciprocal of a number is <em>1 divided by the number</em>, not the number divided by 1. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt6-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for incorrect reciprocal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>ta6_mistake</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have confused the reciprocal of a number with its opposite (changing the sign). Remember, the reciprocal is <em>1 divided by the number</em>, and it keeps the same sign as the original number. Please review the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>465518365</deployedseed>
    <deployedseed>815979139</deployedseed>
    <deployedseed>1990959436</deployedseed>
    <deployedseed>262689654</deployedseed>
    <deployedseed>1688501730</deployedseed>
    <deployedseed>827328484</deployedseed>
    <deployedseed>184111486</deployedseed>
    <deployedseed>947175450</deployedseed>
    <deployedseed>765698440</deployedseed>
    <deployedseed>748301951</deployedseed>
    <deployedseed>2013111401</deployedseed>
    <deployedseed>1699958499</deployedseed>
    <deployedseed>884856950</deployedseed>
    <deployedseed>512812876</deployedseed>
    <deployedseed>1429401560</deployedseed>
    <deployedseed>56799939</deployedseed>
    <deployedseed>637435605</deployedseed>
    <deployedseed>955694543</deployedseed>
    <deployedseed>1529054489</deployedseed>
    <deployedseed>757940225</deployedseed>
    <deployedseed>1646774901</deployedseed>
    <deployedseed>904150187</deployedseed>
    <deployedseed>714681107</deployedseed>
    <deployedseed>1115481382</deployedseed>
    <deployedseed>2128333971</deployedseed>
    <deployedseed>465371808</deployedseed>
    <deployedseed>1110083923</deployedseed>
    <deployedseed>1980873697</deployedseed>
    <deployedseed>142472710</deployedseed>
    <deployedseed>506193551</deployedseed>
    <deployedseed>1345694125</deployedseed>
    <deployedseed>1345180312</deployedseed>
    <deployedseed>1714313278</deployedseed>
    <deployedseed>359154510</deployedseed>
    <deployedseed>561801262</deployedseed>
    <deployedseed>1023666363</deployedseed>
    <deployedseed>186673072</deployedseed>
    <deployedseed>1309752417</deployedseed>
    <deployedseed>452963562</deployedseed>
    <deployedseed>329165458</deployedseed>
    <deployedseed>1752567953</deployedseed>
    <deployedseed>881041044</deployedseed>
    <deployedseed>1009764235</deployedseed>
    <deployedseed>596234351</deployedseed>
    <deployedseed>816148089</deployedseed>
    <deployedseed>543186199</deployedseed>
    <deployedseed>1524178602</deployedseed>
    <deployedseed>1846949772</deployedseed>
    <deployedseed>1875913515</deployedseed>
    <deployedseed>216692282</deployedseed>
    <deployedseed>1717039698</deployedseed>
    <deployedseed>12498809</deployedseed>
    <deployedseed>1210806795</deployedseed>
    <deployedseed>44086476</deployedseed>
    <deployedseed>1856417570</deployedseed>
    <deployedseed>272121326</deployedseed>
    <deployedseed>346054250</deployedseed>
    <deployedseed>1957425539</deployedseed>
    <deployedseed>2060565869</deployedseed>
    <deployedseed>1577728830</deployedseed>
    <deployedseed>87827585</deployedseed>
    <deployedseed>2047761627</deployedseed>
    <deployedseed>351844756</deployedseed>
  </question>
</quiz>
