<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Quadratic equation from theatre seating (non-monic)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Rows of seats are arranged in a theatre. There are \((2x - {@b@})\) rows, and each row contains \(x\) seats. Altogether, {@T@} seats are arranged.
</p>

<p>(a) Form a quadratic equation, in standard form, representing the situation.</p>
<p>[[input:eqn1]] [[validation:eqn1]] [[feedback:prt1]]</p>

<p>(b) Determine the value of \(x\).</p>
<p>\(x =\) [[input:ans1]] [[validation:ans1]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Let \(x\) be the number of seats in each row.
</p>

<p>
The number of rows is \(2x - {@b@}\), and each row has \(x\) seats. Multiplying the number of rows by the number of seats per row gives the total number of seats:
</p>

<p>
\[
x(2x - {@b@}) = {@T@}
\]
</p>

<p>
Step 1: Expand the left-hand side.
</p>

<p>
\[
2x^2 - {@b@}x = {@T@}
\]
</p>

<p>
Step 2: Rearrange into standard quadratic form (equal to zero).
</p>

<p>
\[
2x^2 - {@b@}x - {@T@} = 0
\]
</p>

<p>
Step 3: Factorise. Because the coefficient of \(x^2\) is \(2\), the factors take the form \((x - r)(2x + s)\):
</p>

<p>
\[
{@qeq_factored@}
\]
</p>

<p>
Step 4: Apply the zero product rule and solve each factor.
</p>

<p>
\[
x = {@sol@} \quad \text{or} \quad x = {@negroot@}
\]
</p>

<p>
Since the number of seats in a row cannot be negative, we reject the negative solution. Therefore:
</p>

<p>
\[
x = {@sol@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Design backwards from a nice positive root r (seats per row).
   Fixed leading coefficient 2 makes the quadratic genuinely NON-MONIC.
   b is ODD, which guarantees the 2 does not factor out (matches the (2x-1) style). */
b : rand([1,3,5,7]);
/* choose r so that the number of rows (2r - b) is at least 2 and total stays reasonable */
r : b + rand(4) + 2;          /* r >= b+2, so rows = 2r - b >= b+4 >= 5 */
rows : 2*r - b;               /* number of rows when x = r */
T : r*rows;                   /* total seats */

/* standard-form expression and equation */
expr : 2*x^2 - b*x - T;
qeq  : expr = 0;

/* factored form for the worked solution: (x - r)(2x + (2r - b)) */
sfac : 2*r - b;               /* the s in (2x + s) */
qeq_factored : (x - r)*(2*x + sfac) = 0;

/* solutions */
sol : r;                      /* the meaningful (positive) root */
negroot : -sfac/2;            /* the rejected negative root */

/* --- distractors for the EQUATION part (a) --- */
me1 : 2*x^2 - b*x = T;        /* not rearranged to = 0 */
me2 : x*(2*x - b) = T;        /* correct relationship but unexpanded / not standard form */
me3 : 2*x^2 - b*x + T = 0;    /* sign error on the total */
me4 : expr;                   /* wrote the expression only, forgot "= 0" */

/* --- distractors for the SOLUTION part (b) --- */
ms1 : negroot;                /* gave the negative root */
ms3 : rows;                   /* number of rows (2x - b); also the x=s slip, same value */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@qeq@}, x = {@sol@} (rows = {@rows@}, seats/row = {@sol@}, total {@T@})</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Quadratic equations in real-life situations (non-monic)</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The learner forms a non-monic quadratic equation from a real-life theatre-seating situation, rearranges it into standard form, factorises a quadratic with a leading coefficient of 2, and determines the physically meaningful solution.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>sol</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>eqn1</name>
      <type>algebraic</type>
      <tans>qeq</tans>
      <boxsize>18</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Is the student's equation the standard-form quadratic (allowing any rearrangement
   that is algebraically the same equation, e.g. terms on either side)? We test that
   the two sides differ by the target expression. lhs - rhs of the student's entry: */
sdiff : lhs(eqn1) - rhs(eqn1);
/* accept if sdiff is equal to expr or to -expr (they may have written it the other way round) */
eqn_ok : is(sdiff = expr) or is(sdiff = -expr);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct standard-form equation (any equivalent arrangement)</description>
        <answertest>AlgEquiv</answertest>
        <sans>eqn_ok</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Not rearranged to = 0</description>
        <answertest>AlgEquiv</answertest>
        <sans>eqn1</sans>
        <tans>me1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right relationship, but standard form needs every term on one side so the equation equals zero. Subtract \({@T@}\) from both sides.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Unexpanded product form</description>
        <answertest>AlgEquiv</answertest>
        <sans>eqn1</sans>
        <tans>me2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That correctly describes the situation, but it is not yet in standard form. Expand the brackets and move the total across so the equation reads \(2x^2 - {@b@}x - {@T@} = 0\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Sign error on the total</description>
        <answertest>AlgEquiv</answertest>
        <sans>eqn1</sans>
        <tans>me3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Check the sign of the constant term. Moving \({@T@}\) from the right-hand side to the left makes it negative, giving \(-{@T@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Wrote expression only, not an equation</description>
        <answertest>AlgEquiv</answertest>
        <sans>eqn1</sans>
        <tans>me4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the correct expression, but it needs to be written as an equation. Add \(= 0\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Start from \(x(2x - {@b@}) = {@T@}\), then expand and rearrange to standard form \(2x^2 - {@b@}x - {@T@} = 0\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct positive solution</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>sol</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Gave the negative root</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ms1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the other root of the quadratic, but it is negative. A number of seats cannot be negative, so choose the positive solution.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Answered number of rows (or x=s slip) instead of seats per row</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ms3</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That value is the number of rows, \(2x - {@b@}\) — also what you get if you solve \(2x + s = 0\) as \(x = s\) by mistake. The question asks for \(x\), the number of seats in each row. From \(2x + s = 0\), \(x = -\tfrac{s}{2}\); take the positive root instead.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Solve the standard-form equation \(2x^2 - {@b@}x - {@T@} = 0\) by factorising, then reject the negative root.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1017619104</deployedseed>
    <deployedseed>2025936836</deployedseed>
    <deployedseed>162636307</deployedseed>
    <deployedseed>1676320229</deployedseed>
    <deployedseed>1343340477</deployedseed>
    <deployedseed>2129744239</deployedseed>
    <deployedseed>1952439169</deployedseed>
    <deployedseed>827552270</deployedseed>
    <deployedseed>243366119</deployedseed>
    <deployedseed>1376876844</deployedseed>
    <deployedseed>875591216</deployedseed>
    <deployedseed>1808768083</deployedseed>
    <deployedseed>1931043450</deployedseed>
    <deployedseed>2049769629</deployedseed>
    <deployedseed>1196309378</deployedseed>
    <deployedseed>1510594026</deployedseed>
    <deployedseed>1848273161</deployedseed>
    <deployedseed>782246914</deployedseed>
    <deployedseed>710598860</deployedseed>
    <deployedseed>997427932</deployedseed>
    <deployedseed>317451354</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Teacher's answers get full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>sol</value>
      </testinput>
      <testinput>
        <name>eqn1</name>
        <value>qeq</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Equation entered the other way round (0 = expr) still correct; solution correct.</description>
      <testinput>
        <name>ans1</name>
        <value>sol</value>
      </testinput>
      <testinput>
        <name>eqn1</name>
        <value>0 = expr</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Equation not rearranged to zero (part a distractor).</description>
      <testinput>
        <name>ans1</name>
        <value>sol</value>
      </testinput>
      <testinput>
        <name>eqn1</name>
        <value>me1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Unexpanded product form (part a distractor).</description>
      <testinput>
        <name>ans1</name>
        <value>sol</value>
      </testinput>
      <testinput>
        <name>eqn1</name>
        <value>me2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Negative root given for x (part b distractor).</description>
      <testinput>
        <name>ans1</name>
        <value>ms1</value>
      </testinput>
      <testinput>
        <name>eqn1</name>
        <value>qeq</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Answered number of rows instead of x (part b distractor).</description>
      <testinput>
        <name>ans1</name>
        <value>ms3</value>
      </testinput>
      <testinput>
        <name>eqn1</name>
        <value>qeq</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
