<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the value of a number given a reciprocal</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>If the reciprocal of \( x \) is \( {@frac@} \), find the value of \( x \).</p>

<p>\( x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Worked Solution</p>
<p>Since the reciprocal of \( x \) is \( {@frac@} \) and by definition, the reciprocal of a number 2 is 1 divided by the number 2: \(\frac{1}{2}\)</p>
<p>Then: \( \frac{1}{x} = {@frac@} \)</p>
<p>Therefore, \( x = {@float(ta)@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>den: rand(110) + 20;
num:rand(10)+1;

/* Create the fraction 1/den */
frac: num/den;

/* The answer is the denominator */
ta: float(den/num);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@float(ta)@}</p>
<p> {@frac@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct value given a reciprocal</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have failed to give the correct value of x , remember the value x given its reciprocal is  1/x is just  x and x/2 given the reciprocal is 2/x </p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>247331284</deployedseed>
    <deployedseed>1018330710</deployedseed>
    <deployedseed>1979314737</deployedseed>
    <deployedseed>1561590089</deployedseed>
    <deployedseed>1504645617</deployedseed>
    <deployedseed>2086121249</deployedseed>
    <deployedseed>1998432701</deployedseed>
    <deployedseed>851850821</deployedseed>
    <deployedseed>1710616377</deployedseed>
    <deployedseed>695000824</deployedseed>
    <deployedseed>1288317707</deployedseed>
    <deployedseed>1806065718</deployedseed>
    <deployedseed>918565941</deployedseed>
    <deployedseed>1266223687</deployedseed>
    <deployedseed>1411070289</deployedseed>
    <deployedseed>86148267</deployedseed>
    <deployedseed>1509887159</deployedseed>
    <deployedseed>965487058</deployedseed>
    <deployedseed>1425558884</deployedseed>
    <deployedseed>298441853</deployedseed>
    <deployedseed>1864589879</deployedseed>
    <deployedseed>1041996882</deployedseed>
    <deployedseed>597231621</deployedseed>
    <deployedseed>1699433257</deployedseed>
    <deployedseed>390270289</deployedseed>
    <deployedseed>536793664</deployedseed>
    <deployedseed>1419673893</deployedseed>
    <deployedseed>1687113431</deployedseed>
    <deployedseed>1101662410</deployedseed>
    <deployedseed>353978204</deployedseed>
    <deployedseed>1103483333</deployedseed>
    <deployedseed>533681069</deployedseed>
    <deployedseed>701709707</deployedseed>
    <deployedseed>775515316</deployedseed>
    <deployedseed>1444156537</deployedseed>
    <deployedseed>1653725950</deployedseed>
    <deployedseed>756125507</deployedseed>
    <deployedseed>1584650953</deployedseed>
    <deployedseed>1915804331</deployedseed>
    <deployedseed>288700277</deployedseed>
    <deployedseed>154853860</deployedseed>
    <deployedseed>1653899817</deployedseed>
    <deployedseed>1743031049</deployedseed>
    <deployedseed>1695085745</deployedseed>
    <deployedseed>918329384</deployedseed>
    <deployedseed>1693832</deployedseed>
    <deployedseed>1670351460</deployedseed>
    <deployedseed>764201322</deployedseed>
    <deployedseed>1656503474</deployedseed>
    <deployedseed>1347096249</deployedseed>
    <deployedseed>1853576786</deployedseed>
    <deployedseed>2087810704</deployedseed>
    <deployedseed>783940074</deployedseed>
    <deployedseed>1229106209</deployedseed>
    <deployedseed>581544893</deployedseed>
    <deployedseed>532986555</deployedseed>
    <deployedseed>1375449449</deployedseed>
    <deployedseed>176377096</deployedseed>
    <deployedseed>1602944517</deployedseed>
    <deployedseed>409950249</deployedseed>
    <deployedseed>2051826460</deployedseed>
    <deployedseed>374827089</deployedseed>
    <deployedseed>2021898737</deployedseed>
    <deployedseed>1351387803</deployedseed>
    <deployedseed>758916880</deployedseed>
    <deployedseed>1285630005</deployedseed>
    <deployedseed>1497080634</deployedseed>
    <deployedseed>1360708738</deployedseed>
    <deployedseed>75551755</deployedseed>
    <deployedseed>1372286986</deployedseed>
    <deployedseed>1420356470</deployedseed>
    <deployedseed>1416214922</deployedseed>
    <deployedseed>464776260</deployedseed>
    <deployedseed>980355861</deployedseed>
    <deployedseed>632613016</deployedseed>
    <deployedseed>1690858092</deployedseed>
    <deployedseed>2093898830</deployedseed>
    <deployedseed>327598056</deployedseed>
    <deployedseed>125957033</deployedseed>
    <deployedseed>389259216</deployedseed>
    <deployedseed>922432485</deployedseed>
    <deployedseed>922572681</deployedseed>
    <deployedseed>903473218</deployedseed>
    <deployedseed>896034934</deployedseed>
    <deployedseed>1823047944</deployedseed>
    <deployedseed>1158314255</deployedseed>
    <deployedseed>776852510</deployedseed>
    <deployedseed>1630142001</deployedseed>
    <deployedseed>93512526</deployedseed>
    <deployedseed>72430844</deployedseed>
    <deployedseed>1851016067</deployedseed>
    <deployedseed>1660571326</deployedseed>
    <deployedseed>610660783</deployedseed>
    <deployedseed>537789370</deployedseed>
    <deployedseed>23328735</deployedseed>
    <deployedseed>863543430</deployedseed>
    <deployedseed>982903276</deployedseed>
    <deployedseed>346899323</deployedseed>
    <deployedseed>1130889013</deployedseed>
    <deployedseed>2050215739</deployedseed>
  </question>
</quiz>
