<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding Complementary Angles and Trigonometric Ratios</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Angles \({@letter_1@}\) and \({@letter_2@}\) are complementary. If \( {@letter_1@}\) = {@fraction@}\({@letter_2@}\) ,
</p>

<ol type="a">
    <li>
        <p>Find the values of the angles:</p>
        <ol type="i">
            <li>\({@letter_1@}=\)[[input:ans1]]\(^\circ\) [[validation:ans1]]</li>
            <li>\({@letter_2@}=\)[[input:ans2]]\(^\circ\)[[validation:ans2]]</li>
           <p>[[feedback:prt1]]<p>
        </ol>
    </li>
    <li>
        <p>
        Using a calculator, find the following correct to four decimal places:
        </p>
        <ol type="i">
            <li>\(\sin\, {@letter_1@}=\)[[input:ans3]][[validation:ans3]]</li>
            <li>\(\cos\,{@letter_2@}=\)[[input:ans4]] [[validation:ans4]]</li>
           <p>[[feedback:prt2]]<p>
        </ol>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>When we say {@letter_1@} and {@letter_2@} are complementary, it implies that \({@letter_1@}+{@letter_2@}=90^\circ\).<br>
Also, we have \({@letter_1@}={@fraction@}{@letter_2@}\).</p>

<p>First, write \({@letter_2@}\) in terms of \({@letter_1@}\), that is,<br>
\({@eq_1@}\)<br>
\({@eq_2@}\)</p>

<p>Since \({@letter_1@}+{@letter_2@}=90^\circ\), substitute \({@eq_2@}\):<br>
\({@frac_2*letter_1@}+{@letter_1@}=90\)<br>
\({@eq_5@}\)<br>
\(\quad {@letter_1@}=\frac{90}{{@coeficient@}}\)<br>
\(\qquad ={@Ta_1@}^\circ\)</p>

<p>Therefore,<br>
<ol type='i'>
<li>\(\angle \,{@letter_1@}={@Ta_1@}^\circ\)</li>
<li>
<p>\(\angle\,{@letter_2@}\)</p>
<p>
\(\begin{align}
\angle \,{@letter_2@}&=90^\circ-{@Ta_1@}^\circ\\
&={@Ta_2@}^\circ
\end{align}\)
</p>
</li>
</ol>
</p>

<p>Use your calculator to find both \(\sin\,{@Ta_1@}^\circ\) and \(\cos\,{@Ta_2@}^\circ\), and write the answers correct to 4 significant figures.<br>
Therefore,
<ol type='i'>
<li>\(\sin\,{@Ta_1@}^\circ={@an_1@}\)</li>
<li>\(\cos\,{@Ta_2@}^\circ={@an_2@}\)</li>
</ol>
</p>
<p><strong>Notice:</strong> Because the angles are complementary, \(\sin\,{@Ta_1@}^\circ\) is exactly equal to \(\cos\,{@Ta_2@}^\circ\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>letter_1: rand([C,D,E,F,G,H,I,J,K,M]);
letter_2: rand([N,O,Q,R,U,W,X,Y,Z]);
frac: 1;
/* Only pick denominators that yield a whole number when dividing 90 by (frac_2 + 1) */
frac_2: rand([2, 4, 8, 14]); 
fraction: frac/frac_2;
eq_1: letter_1=fraction*letter_2;
eq_2: letter_2=frac_2*letter_1;
simp: false;
eq_3: letter_1+letter_2=90;
eq_4: frac_2*letter_1+letter_1=90;
simp: true;
eq_5: frac_2*letter_1+letter_1=90;
coeficient: frac_2+1;
Ta_1: 90/coeficient; /* round() is no longer needed */
Ta_2: 90-Ta_1;
/*Part b of the question*/
ans_1: sin(Ta_1*pi/180);
an_1: dispdp(ans_1,4);
ans_2: cos(Ta_2*pi/180); /* Fixed: Was calculating cos(Ta_1) */
an_2: dispdp(ans_2,4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Part a of the question</p>
<p>\(\angle \,{@letter_1@}={@Ta_1@}^\circ\)</p>
<p>\(\angle \,{@letter_2@}={@Ta_2@}^\circ\)</p>
<p>Part b of the question</p>
<p>\(\sin\,{@letter_1@}={@an_1@}\)</p>
<p>\(\cos\,{@letter_1@}={@an_2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Ta_1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>Ta_2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>an_1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>an_2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct angle.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your value for \(\angle {@letter_1@}\) is incorrect.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the correct angle.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>Ta_2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your value for \(\angle {@letter_2@}\) is incorrect.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the corect value computed in the calculator.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ans_1</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your calculation for \(\sin {@letter_1@}^\circ\) is incorrect.</text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the corect value computed in the calculator.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans4</sans>
        <tans>ans_2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Your calculation for \(\cos {@letter_2@}^\circ\) is incorrect.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>373903067</deployedseed>
    <deployedseed>1658010805</deployedseed>
    <deployedseed>1294555872</deployedseed>
    <deployedseed>2043432042</deployedseed>
    <deployedseed>729373309</deployedseed>
    <deployedseed>1231602281</deployedseed>
    <deployedseed>1141069541</deployedseed>
    <deployedseed>644243532</deployedseed>
    <deployedseed>484132993</deployedseed>
    <deployedseed>496402449</deployedseed>
    <deployedseed>559698325</deployedseed>
    <deployedseed>1078371541</deployedseed>
    <deployedseed>1047302120</deployedseed>
    <deployedseed>1776651278</deployedseed>
    <deployedseed>1394542424</deployedseed>
    <deployedseed>1897513368</deployedseed>
    <deployedseed>1054937381</deployedseed>
    <deployedseed>17485283</deployedseed>
    <deployedseed>142470031</deployedseed>
    <deployedseed>1432156303</deployedseed>
    <deployedseed>598878401</deployedseed>
    <deployedseed>1225060345</deployedseed>
    <deployedseed>1157804439</deployedseed>
    <deployedseed>1179649772</deployedseed>
    <deployedseed>1396505402</deployedseed>
    <deployedseed>501081548</deployedseed>
    <deployedseed>1850365025</deployedseed>
    <deployedseed>1228590386</deployedseed>
    <deployedseed>1561762746</deployedseed>
    <deployedseed>876443295</deployedseed>
    <deployedseed>784458271</deployedseed>
    <deployedseed>1601832106</deployedseed>
    <deployedseed>1716284061</deployedseed>
    <deployedseed>201411463</deployedseed>
    <deployedseed>1332468508</deployedseed>
    <deployedseed>1746399857</deployedseed>
    <deployedseed>680870465</deployedseed>
    <deployedseed>1000134962</deployedseed>
    <deployedseed>1370922448</deployedseed>
    <deployedseed>929294030</deployedseed>
    <deployedseed>65106470</deployedseed>
    <deployedseed>1758242948</deployedseed>
    <deployedseed>2136582747</deployedseed>
    <deployedseed>1697541546</deployedseed>
    <deployedseed>1084626956</deployedseed>
    <deployedseed>1721729363</deployedseed>
    <deployedseed>1073472835</deployedseed>
    <deployedseed>2141616708</deployedseed>
    <deployedseed>259976915</deployedseed>
    <deployedseed>590643874</deployedseed>
    <deployedseed>347392091</deployedseed>
    <deployedseed>1841797955</deployedseed>
    <deployedseed>724273637</deployedseed>
    <deployedseed>2019312817</deployedseed>
    <deployedseed>1474612261</deployedseed>
    <deployedseed>2084729454</deployedseed>
    <deployedseed>918727741</deployedseed>
    <deployedseed>1776299407</deployedseed>
    <deployedseed>1092865837</deployedseed>
    <deployedseed>1658769141</deployedseed>
    <deployedseed>595797854</deployedseed>
    <deployedseed>966399627</deployedseed>
    <deployedseed>659113360</deployedseed>
    <deployedseed>880733069</deployedseed>
    <deployedseed>944251707</deployedseed>
    <deployedseed>1662430013</deployedseed>
    <deployedseed>639064594</deployedseed>
    <deployedseed>1466709014</deployedseed>
    <deployedseed>2091481601</deployedseed>
    <deployedseed>1632991310</deployedseed>
    <deployedseed>646920566</deployedseed>
    <deployedseed>900460925</deployedseed>
    <deployedseed>1700799043</deployedseed>
    <deployedseed>1268022297</deployedseed>
    <deployedseed>1879254100</deployedseed>
    <deployedseed>883225486</deployedseed>
    <deployedseed>1095992335</deployedseed>
    <deployedseed>1334831539</deployedseed>
    <deployedseed>1528944517</deployedseed>
    <deployedseed>977156498</deployedseed>
    <deployedseed>76771976</deployedseed>
    <deployedseed>1563999735</deployedseed>
    <deployedseed>590244404</deployedseed>
    <deployedseed>444420397</deployedseed>
    <deployedseed>1574248529</deployedseed>
    <deployedseed>2029356402</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>Ta_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Ta_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>an_1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>an_2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
