<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Selecting appropriate data collection methods</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Read the following research scenario:</p>
<p>{@q_text@}</p>
<p>Which method of data collection is most appropriate for this scenario?</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>Correct Answer: {@correct_text@}</strong></p>
<p>Each method serves a specific purpose:</p>
<ul>
    <li><strong>Observation:</strong> Studying behavior in natural settings.</li>
    <li><strong>Focus Group:</strong> Gathering deep qualitative insights from small groups.</li>
    <li><strong>Census:</strong> Collecting data from the entire population.</li>
    <li><strong>Experimentation:</strong> Establishing cause-and-effect (control vs treatment).</li>
    <li><strong>Questionnaire:</strong> Gathering standardized data from large groups efficiently.</li>
    <li><strong>Secondary Data:</strong> Using existing records or data collected by others.</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. SCENARIO & CORRECT ANSWER SETUP */
/* Define the scenarios and their matching correct method */
scenarios: [
  ["You stand in a corner and record how long customers spend in each aisle without interfering.", "Observation"],
  ["You gather detailed opinions and feelings from a small group of 10 people regarding a new product.", "Focus Group"],
  ["You collect data from every single household in the entire country.", "Census"],
  ["You treat one plot with fertilizer and leave another without to compare the growth results.", "Experimentation"],
  ["You survey 500 employees anonymously using a standard form to gather data quickly.", "Questionnaire"],
  ["You use existing rainfall records from the Meteorological Department for your research.", "Secondary Data Collection"]
];

/* Pick a random scenario */
rand_idx: rand(length(scenarios)) + 1;
chosen_scenario: scenarios[rand_idx];

q_text: chosen_scenario[1];
correct_text: chosen_scenario[2];

/* 2. DISTRACTOR SETUP */
/* Master list of all possible methods */
all_methods: [
    "Observation", 
    "Focus Group", 
    "Census", 
    "Experimentation", 
    "Questionnaire", 
    "Secondary Data Collection", 
    "Face-to-Face Interview", 
    "Telephone Survey", 
    "Case Study"
];

/* Filter out the correct answer to ensure distractors are actually wrong */
wrong_pool: sublist(all_methods, lambda([x], x # correct_text));

/* Shuffle the wrong pool so we get random distractors every time */
shuffled_wrongs: random_permutation(wrong_pool);

/* 3. BUILD THE ANSWER LIST (Matching your sample structure) */

/* Correct Answer Block: ID 1, Boolean TRUE, Text */
correct_answer: rand_selection([
   [1, true, correct_text]
], 1);

/* Incorrect Answers Block: ID 2-4, Boolean FALSE, Text (picked from shuffled pool) */
incorrect_answers: rand_selection([
   [2, false, shuffled_wrongs[1]],
   [3, false, shuffled_wrongs[2]],
   [4, false, shuffled_wrongs[3]]
], 3);

/* 4. FINAL PERMUTATION */
/* Combine and shuffle to get the final list 'ta' */
ta: random_permutation(append(correct_answer, incorrect_answers));

/* 5. INPUT HELPER */
/* Extract just the text strings for the STACK Radio Button list */
/* 'third' takes the 3rd element of each sublist (the text) */
final_options: map(third, ta);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}
{@@}
{@@}
{@@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">In this question, you are presented with a randomized research scenario describing a specific data gathering context. You are required to identify and select the most appropriate method of data collection (such as Observation, Census, Questionnaire, etc.) that best fits the situation described.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
