<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>A bag has 6 red and 4 blue balls.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A bag has {@ball1@} {@ball1c@} and {@ball2@} {@ball2c@} balls. One ball is drawn, replaced, and then another is drawn. Find the probability that:</p>
<p></p>
<p>
     <ol type="a">
             <li><p>Both balls are {@ball2c@} [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p></li>
             <li><p>One ball is {@ball1c@} and the other is {@ball2c@} [[input:ans2]] [[validation:ans2]][[feedback:prt3]]</p></li>
     </ol>

</p>

<p></p>
<p></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><h3> Worked Solution</h3></p>
<p>The probability that:</p>
<p>
     <ol type="a">
             <li><p>Both balls are {@ball2c@}</p>
<p> To find the probability that both balls are {@ball2c@}, we multiply the probability of drawing a {@ball2c@} ball on the first turn by the probability of drawing one on the second turn.</p>
<p>Because the first ball is replaced before the second draw, the probability stays exactly the same for both draws:</p>
<p>{@pball2@} \(\times\) {@pball2@} = {@ta@}</p></li>
             <li><p>One ball is {@ball1c@} and the other is {@ball2c@}</p>
<p> To find the probability that one ball is {@ball1c@} and the other is {@ball2c@}, we calculate the probability of both possible ways and add them together:</p>
<p><strong>First way</strong> ({@ball1c@} then {@ball2c@}): Multiply {@pball1@} \(\times\) {@pball2@}</p>
<p>Second way ({@ball2c@} then {@ball1c@}): Multiply {@pball2@} \(\times\) {@pball1@}</p>
<p>Since both ways have the same probability, the final calculation is:2 \(\times\) ({@pball1@} \(\times\) {@pball2@}) = {@ta1@}</p></li>
     </ol>

</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>ball1: rand(10)+2;
ball2: rand(10)+4;

pball1: ball1/(ball1+ball2);
pball2: ball2/(ball1+ball2);

pball11: (ball1-1)/(ball1+ball2-1);
pball12: ball1/(ball1+ball2-1);
pball21: ball2/(ball1+ball2-1);
pball22: (ball2-1)/(ball1+ball2-1);

ball1c: rand([red, yellow, green, purple, orange]);
ball2c: rand([blue, black, white, pink, brown]);

ta: pball2*pball2;
taw: pball2*pball22;
ta1: 2*pball1*pball2;
ta1w: pball2*pball12 + pball1*pball21;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>\({@blueballs@}\)
\({@redballs@}\)
\({@ball1@}\)
\({@ball2@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Joseph Karisa Baya</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"></p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>taw</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you calculated the probability <strong>without replacement</strong>. </p>
<p>Check the wording of the question carefully: it states the ball is <strong>replaced</strong>. Because the first ball is put back, the total number of balls in the bag does not change for the second draw. Your second probability fraction should have the exact same numerator and denominator as the first.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta1w</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you calculated the probability <strong>without replacement</strong>. </p>
<p>Because the first ball is put back, the total number of balls in the bag does not change for the second draw. Your second probability fraction should have the exact same numerator and denominator as the first.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>110886361</deployedseed>
    <deployedseed>698745439</deployedseed>
    <deployedseed>28313238</deployedseed>
    <deployedseed>484539779</deployedseed>
    <deployedseed>1868190754</deployedseed>
    <deployedseed>238772431</deployedseed>
    <deployedseed>417405764</deployedseed>
    <deployedseed>1828172926</deployedseed>
    <deployedseed>658067987</deployedseed>
    <deployedseed>541047831</deployedseed>
    <deployedseed>30282039</deployedseed>
    <deployedseed>1948197227</deployedseed>
    <deployedseed>410013966</deployedseed>
    <deployedseed>1817001954</deployedseed>
    <deployedseed>716416434</deployedseed>
    <deployedseed>1572625601</deployedseed>
    <deployedseed>972519857</deployedseed>
    <deployedseed>1124684449</deployedseed>
    <deployedseed>1913818216</deployedseed>
    <deployedseed>817589328</deployedseed>
  </question>
</quiz>
