<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the distance and displacement of a delivery drone</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A delivery drone is programmed to deliver two packages along a straight East-West flight path. The drone's motion is recorded as follows:</p>
<ul>
    <li>It flies \({@d1@}\) meters East to deliver the first package.</li>
    <li>It then turns around and flies \({@d2@}\) meters West to deliver the second package.</li>
</ul>

<p>
<ol>
    <li>What is the total distance travelled by the drone? [[input:ans1]] m [[validation:ans1]][[feedback:prt1]]</li>
    <li>What is the drone's displacement from its starting point? [[input:ans2]] m [[validation:ans2]][[feedback:prt2]]</li>
</ol>
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<b>1. Total Distance</b>
<p>Distance is the total length of the path travelled, regardless of direction:</p>
<p>
\(\begin{align}
\text{Distance} &= \text{Leg 1} + \text{Leg 2} \\
&= {@d1@} \text{ m} + {@d2@} \text{ m} \\
&= {@ta1@} \text{ m}
\end{align}\)
</p>

<b>2. Displacement</b>
<p>Displacement is the change in position (Final Position - Initial Position). Since the drone moved in opposite directions, we treat East as positive and West as negative:</p>
<p>
\(\begin{align}
\text{Displacement} &= \text{Final Position} - \text{Initial Position} \\
&= {@d1@} \text{ m} - {@d2@} \text{ m} \\
&= {@ta2@} \text{ m East}
\end{align}\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Randomized Directions and Distances */
/* We use East as positive, West as negative */
d1: 150 + rand(101);     /* First leg: 150 to 250 m East */
d2: 50 + rand(51);       /* Second leg: 50 to 100 m West */

/* 2. Calculations */
/* Total Distance = sum of absolute magnitudes */
total_dist: d1 + d2;

/* Displacement = Final Position - Initial Position */
/* Since West is negative, we subtract d2 from d1 */
displacement: d1 - d2;

/* 3. Teacher Answers */
ta1: total_dist;
ta2: displacement;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta1@}</p>
<p>{@ta2@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>449005281</deployedseed>
    <deployedseed>809821926</deployedseed>
    <deployedseed>2009376217</deployedseed>
    <deployedseed>431917911</deployedseed>
    <deployedseed>827062701</deployedseed>
    <deployedseed>516417536</deployedseed>
    <deployedseed>1345760999</deployedseed>
    <deployedseed>1268376684</deployedseed>
    <deployedseed>928602837</deployedseed>
    <deployedseed>1416205389</deployedseed>
    <deployedseed>2027328549</deployedseed>
    <deployedseed>1621113939</deployedseed>
    <deployedseed>354883830</deployedseed>
    <deployedseed>502686951</deployedseed>
    <deployedseed>1631810802</deployedseed>
    <deployedseed>875636354</deployedseed>
    <deployedseed>1061391715</deployedseed>
    <deployedseed>994966224</deployedseed>
    <deployedseed>1061018256</deployedseed>
    <deployedseed>676088274</deployedseed>
  </question>
</quiz>
