<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Application of scale factors to a cylindrical model</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>An engineering firm has designed a prototype for a chemical storage tank. The prototype consists of a cylindrical body and a hemispherical top. The dimensions of the prototype are: {@r_p@} m Radius and a Cylindrical Height of {@h_p@} m  </p>
    
<p>The firm decides to produce an industrial version of the tank by enlarging every dimension by a linear scale factor of {@k@}.</p>

[[jsxgraph width="350px" height="400px"]]

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10, 20, 10, -5],
    axis: false,
    showCopyright: false
});

var R = {#r_p#} * 1.5;
var H = {#h_p#} * 1.5;
var p = 0.3;

/* =====================
   HEMISPHERICAL TOP
===================== */

// top visible arc
var hemiTop = board.create('curve',
[
    function(t){return R*Math.cos(t)},
    function(t){return H + R*Math.sin(t)},
    0, Math.PI
],
{
    strokeColor: '#1f77b4',
    strokeWidth: 3
});

// bottom hidden arc
var hemiBottom = board.create('curve',
[
    function(t){return R*Math.cos(t)},
    function(t){return H + (R*p)*Math.sin(t)},
    Math.PI, 2*Math.PI
],
{
    strokeColor: '#1f77b4',
    dash: 2,
    strokeWidth: 2
});


/* =====================
   CYLINDER SIDES
===================== */

var leftSide = board.create('segment', [[-R,0],[-R,H]],
{
    strokeColor:'#2ca02c',
    strokeWidth:3
});

var rightSide = board.create('segment', [[R,0],[R,H]],
{
    strokeColor:'#2ca02c',
    strokeWidth:3
});


/* =====================
   CYLINDER BASE
===================== */

// visible base
board.create('curve',
[
    function(t){return R*Math.cos(t)},
    function(t){return (R*p)*Math.sin(t)},
    Math.PI,2*Math.PI
],
{
    strokeColor:'#2ca02c',
    strokeWidth:3
});

// hidden base
board.create('curve',
[
    function(t){return R*Math.cos(t)},
    function(t){return (R*p)*Math.sin(t)},
    0,Math.PI
],
{
    strokeColor:'#2ca02c',
    dash:2,
    strokeWidth:2
});


/* =====================
   HEIGHT LABEL
===================== */

var heightLine = board.create('segment',
[[R+1,0],[R+1,H]],
{strokeWidth:2, strokeColor:'black', dash:1});

board.create('text',
[R+1.5, H/2, "Height = {#h_p#} m"]);


/* =====================
   RADIUS LABEL
===================== */

var radiusLine = board.create('segment',
[[0,H],[R,H]],
{strokeWidth:2, strokeColor:'black'});

board.create('text',
[R/2, H+1, "Radius = {#r_p#} m"]);


/* =====================
   TITLE LABEL
===================== */

board.create('text',[-4, H+R+2, "Prototype Fountain"]);

[[/jsxgraph]]

<p>Complete the following regarding the enlargement:</p>
<ol>
    <li>State the Linear Scale Factor \(k\): [[input:ans1]]  [[validation:ans1]] [[feedback:prt1]]</li>
    <li>Calculate the Area Scale Factor \(k^2\): [[input:ans2]][[validation:ans2]] [[feedback:prt2]]</li>
    <li>Calculate the Volume Scale Factor \(k^3\): [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</li>
    <li>If the prototype has a volume of \({@vol_p@}\) m³, calculate the volume of the industrial tank: [[input:ans4]] [[validation:ans4]] [[feedback:prt4]] m³</li>
</ol>

<p><i>( Round your final volume to 2 decimal places.)</i></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To solve problems involving similarity and enlargement, we apply the scale factor relationship across dimensions:</p>

<p>1. Linear Scale Factor \(k\)</p>
\(\begin{align}
k &= \frac{\text{Image length}}{\text{Object length}} \\
k &= {@ta1@}
\end{align}\)

<p>2. Area Scale Factor \(k^2\)</p>
<p>The area of the enlarged image increases by the square of the linear scale factor:</p>
\(\begin{align}
\text{Area Scale Factor} &= k^2 \\
&= {@ta1@}^2 \\
&= {@ta2@}
\end{align}\)

<p>3. Volume Scale Factor \(k^3\)</p>
<p>The volume of the enlarged image increases by the cube of the linear scale factor:</p>
\(\begin{align}
\text{Volume Scale Factor} &= k^3 \\
&= {@ta1@}^3 \\
&= {@ta3@}
\end{align}\)

<p>4. Capacity of the Industrial Tank</p>
\(\begin{align}
V_{\text{image}} &= V_{\text{object}} \times k^3 \\
&= {@float(vol_p)@} \times {@ta3@} \\
&= {@ta4@} \text{ m}^3
\end{align}\)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Linear Scale Factor (k) */
/* Enlarging from a prototype to a real industrial tank */
k: 5 + rand(6);           /* Scale factor k between 5 and 10 */

/* 2. Prototype (Object) Dimensions */
r_p: rand(2)+2;                   /* Radius in meters */
h_p: (2)+6;                   /* Cylindrical height in meters */

/* 3. Calculations for Prototype (The Object) */
area_base_p: %pi * r_p^2;
area_curved_p: (2 * %pi * r_p * h_p) + (2 * %pi * r_p^2); /* Cyl side + Hemi surface */
vol_p: (%pi * r_p^2 * h_p) + (2/3 * %pi * r_p^3);

/* 4. Calculations for Industrial Tank (The Image) */
/* Linear Dimensions */
r_i: r_p * k;
h_i: h_p * k;

/* Area and Volume Scale Factors */
asf: k^2;
vsf: k^3;

/* Values for the Industrial Tank */
total_area_i: float(area_curved_p * asf);
total_vol_i: float(vol_p * vsf);

/* 5. Teacher Answers */
ta1: k;               /* Linear Scale Factor */
ta2: asf;             /* Area Scale Factor */
ta3: vsf;             /* Volume Scale Factor */
ta4: decimalplaces(total_vol_i,2);     /* Final enlarged volume */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta1@}</p>
<p>{@ta2@}</p>
<p>{@ta3@}</p>
<p>{@ta4@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>779843015</deployedseed>
    <deployedseed>1391937414</deployedseed>
    <deployedseed>372733272</deployedseed>
    <deployedseed>395729003</deployedseed>
    <deployedseed>1918503176</deployedseed>
    <deployedseed>1259174155</deployedseed>
    <deployedseed>1582442312</deployedseed>
    <deployedseed>1125294073</deployedseed>
    <deployedseed>1246850767</deployedseed>
    <deployedseed>1105629775</deployedseed>
  </question>
</quiz>
