<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area scale factor and the volume of a final enlarged sculpture.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A designer creates a prototype of a cylindrical fountain pedestal with a radius of \({@r_p@}\) m and a height of \({@h_p@}\) m.</p>
<p><ul><li>The designer draws the pedestal on a blueprint using a scale factor of <b>1:50</b>. 
What is the linear scale factor \(k\) for this reduction? [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </li></ul></p>
<p>A sculptor is commissioned to create the final fountain by enlarging the prototype using a linear scale factor of \({@k_enlarge@}\).</p>

[[jsxgraph width="500px" height="550px"]]

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-6, 15, 16, -2],
    axis: false,
    showCopyright: false
});

var p = 0.3;

/* STACK variables */
var r = {#r_p#};
var h = {#h_p#};
var k = {#k_enlarge#};

/* Prototype dimensions */
var Rp = r;
var Hp = h;

/* Enlarged dimensions (automatic scaling) */
var Re = k * r;
var He = k * h;

/* Positions */
var x1 = -3;
var x2 = 3;


/* =====================
   PROTOTYPE CYLINDER
===================== */

// sides
board.create('segment', [[x1-Rp,0],[x1-Rp,Hp]], {strokeColor:'#1f77b4', strokeWidth:3});
board.create('segment', [[x1+Rp,0],[x1+Rp,Hp]], {strokeColor:'#1f77b4', strokeWidth:3});

// base
board.create('curve',
[
 function(t){return x1 + Rp*Math.cos(t)},
 function(t){return (Rp*p)*Math.sin(t)},
 Math.PI,2*Math.PI
],
{strokeColor:'#1f77b4', strokeWidth:3});

// top
board.create('curve',
[
 function(t){return x1 + Rp*Math.cos(t)},
 function(t){return Hp + (Rp*p)*Math.sin(t)},
 0,2*Math.PI
],
{strokeColor:'#1f77b4', strokeWidth:3});

board.create('text',[x1-1.5,Hp+1,"Prototype"]);



/* =====================
   ENLARGED CYLINDER
===================== */

// sides
board.create('segment', [[x2-Re,0],[x2-Re,He]], {strokeColor:'#d62728', strokeWidth:3});
board.create('segment', [[x2+Re,0],[x2+Re,He]], {strokeColor:'#d62728', strokeWidth:3});

// base
board.create('curve',
[
 function(t){return x2 + Re*Math.cos(t)},
 function(t){return (Re*p)*Math.sin(t)},
 Math.PI,2*Math.PI
],
{strokeColor:'#d62728', strokeWidth:3});

// top
board.create('curve',
[
 function(t){return x2 + Re*Math.cos(t)},
 function(t){return He + (Re*p)*Math.sin(t)},
 0,2*Math.PI
],
{strokeColor:'#d62728', strokeWidth:3});

board.create('text',[x2-2,He+1,"Final Sculpture (k = {#k_enlarge#})"]);

[[/jsxgraph]]
<ol type="a">
    <li>What is the Area Scale Factor for the final sculpture? [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>
    <li>If the prototype has a volume of {@vol_proto@} m³, calculate the volume of the final enlarged sculpture. [[input:ans3]] [[validation:ans3]] [[feedback:prt3]] m³</li>
</ol>

<p><i>( Round your final volume to 2 decimal places.)</i></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</p>
<p>\(\begin{align}
\text{Linear Scale Factor (Reduction)} &= \frac{1}{50} = 0.02 \\
\text{Area Scale Factor} &= k^2 = {@k_enlarge@}^2 = {@asf@} \\
\text{Volume of Image} &= \text{Volume of Object} \times k^3 \
&= {@vol_proto@} \times {@k_enlarge@}^3 \
&= {@vol_proto@} \times {@vsf@} \
&= {@vol_final@} \text{ m}^3
\end{align}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* k1: Reduction for the blueprint (Negative Enlargement/Scale Factor < 1) */
k_red: 1/50; 

/* k2: Positive Enlargement for the final fountain sculpture */
/* The sculptor enlarges the prototype by a factor of 4 */
k_enlarge: rand(2)+2;


/* 2. Original Prototype Dimensions (The Object) */
/* A cylindrical pedestal */
r_p: 1;   /*radius*/
h_p: rand(2)+2; /*height*/  
/* 3. Calculations */
/* Blueprint dimensions (Linear) */
r_blue: float(r_p * k_red * 100); /* converted to cm */

/* Final Sculpture dimensions (Area and Volume) */
asf: k_enlarge^2;
vsf: k_enlarge^3;

vol_proto: %pi * r_p^2 * h_p;
vol_final = vol_proto * k_enlarge^3;

/* 4. Teacher Answers */
ta1: k_red;
ta2: k_enlarge;
ta3: asf;
ta4: vol_final;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta1@}</p>
<p>{@ta2@}</p>
<p>{@ta3@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1524381166</deployedseed>
    <deployedseed>1113255033</deployedseed>
  </question>
</quiz>
