<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Area of a Regular Pentagon</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the area of the figure below, centered at \(O\).</p>

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-40, 40, 40, -40],
    axis: false,
    showCopyright: false
});

/* Centre point O */
var O = board.create('point', [0, 0], {
    name: 'O',
    fixed: true,
    visible: false
});

/* Pentagon vertices (numerical, fixed shape) */
var A = board.create('point', [0, 25.5], {visible: false});
var B = board.create('point', [24.2, 7.9], {visible: false});
var C = board.create('point', [14.8, -20.6], {visible: false});
var D = board.create('point', [-14.8, -20.6], {visible: false});
var E = board.create('point', [-24.2, 7.9], {visible: false});

/* Pentagon */
board.create('polygon', [A, B, C, D, E], {
    borders: {strokeColor: 'black'},
    fillColor: 'none'
});

/* Radius lines */
board.create('segment', [O, C], {
    strokeColor: 'black',
    straightFirst: false,
    straightLast: false
});

board.create('segment', [O, B], {
    strokeColor: 'black',
    straightFirst: false,
    straightLast: false
});

/* Labels */
board.create('text', [-1, -10, "{#slant#} cm"], {
    fontSize: 12,
    fixed: true
});

board.create('text', [-0.2, 2, "O"], {
    fontSize: 12,
    fixed: true
});

/* Base BC */
board.create('segment', [B, C], {
    strokeColor: 'black'
});

board.create('text', [
    function(){ return B.X(); },
    function(){ return B.Y() - 15; },
    "{#base_length#} cm"
], {
    fontSize: 12,
    fixed: true
});

/* Perpendicular height from O to BC */
var midpoint = board.create('midpoint', [B, C], {visible: false});

board.create('segment', [O, midpoint], {
    strokeColor: 'black',
    dash: 1
});

/* Right angle marker */
board.create('nonreflexangle', [B, midpoint, O], {
    radius: 3,
    type: 'square',
    name: '',
    fixed: true
});
[[/jsxgraph]]

<p>
\(\text{Area} =\) [[input:ans1]] \(\text{cm}^2\)
[[feedback:prt1]]
</p>

<p style="display:none;">
[[validation:ans1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p> To calculate the area of the pentagon, divide the shape into five identical isosceles triangles by joining each vertex to
the centre \(O\).
<p>
Each triangle has:
<ul>
  <li>base length \( {@base_length@} \) cm</li>
  <li>two equal sides of length \( {@slant@} \) cm</li>
</ul>
</p>

<p>
To find the perpendicular height of one triangle, we split it into two right-angled triangles.
The hypotenuse is the slanted side, and half the base is
\[
\frac{{@base_length@}}{2}.
\]
</p>

<p>
Using Pythagoras’ theorem:
\[
h = \sqrt{({@slant@})^2 - \left(\frac{{@base_length@}}{2}\right)^2}.
\]
</p>

<p>
The area of one triangle is therefore
\[
\frac{1}{2} \times {@base_length@} \times h.
\]
</p>

<p>
Since the pentagon consists of five identical triangles, its total area is
\[
5 \times \frac{1}{2} \times {@base_length@} \times h = {@ta@}.
\]
</p>

<p>
\[
\boxed{\text{Area} = {@ta@}\text{ cm}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Title: Random Pentagon Area Calculator with Expanded Range (NO conditionals) */

/* Base length: divisible by 10, between 20 and 260 */
base_length : 10 * (rand(25) + 2);

/* Half base */
half_base : base_length / 2;

/* Maximum slant height: always safely greater than half_base */
max_slant : max(base_length - 5, half_base + 5);

/* Slant height: divisible by 5, strictly between half_base and max_slant */
slant : 5 * (
  floor(half_base/5)
  + rand(floor((max_slant - half_base)/5))
  + 1
);

/* Height using Pythagoras */
height : sqrt(slant^2 - half_base^2);
height_d : decimalplaces(float(height), 1);

/* Area of one triangle */
area_triangle : (1/2) * base_length * height_d;
area_triangle_d : decimalplaces(area_triangle, 1);

/* Area of pentagon */
area_pentagon : 5 * area_triangle;
ta_exact : decimalplaces(area_pentagon, 1);
ta1: round(area_pentagon);

/* ---------- Randomised mistake generation (NO conditionals) ---------- */

/* Mistake selector */
mistake_number : rand(4) + 1;

/* List of mistake answers */
mistake_answers : [
  round((1/2) * base_length * slant),      /* using slant as height */
  round((1/2) * base_length * half_base),  /* using half base as height */
  round(area_triangle),                    /* forgetting ×5 */
  round((base_length * slant) / 2)         /* base × slant / 2 */
];

/* Corresponding feedback messages */
mistake_feedback : [
  "Great effort! Remember, the slant height is not the perpendicular height. Use Pythagoras to find the correct height.",
  "Nice try. Half the base helps form the right triangle, but it is not the perpendicular height.",
  "Almost there! You found the area of one triangle. Remember to multiply by 5.",
  "Good attempt! The slant height is not perpendicular to the base. Use the perpendicular height in the area formula."
];

/* Selected mistake answer and feedback */
mistake_answer : mistake_answers[mistake_number];
feedback : mistake_feedback[mistake_number];
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Monica Sherily</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">This question is on finding the area of a regular pentagon. The learner is required to observe the diagram and identify the various dimensions then calculate the area of the pentagon. Guide them to first find the area of one triangle formed by drawing lines from the center to each vertex, and then multiply it by five to get the total area of the pentagon.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>404107258</deployedseed>
    <deployedseed>1208046095</deployedseed>
    <deployedseed>1852468895</deployedseed>
    <deployedseed>936331563</deployedseed>
    <deployedseed>346189755</deployedseed>
    <deployedseed>1835566648</deployedseed>
    <deployedseed>764251822</deployedseed>
    <deployedseed>664016470</deployedseed>
    <deployedseed>1480651678</deployedseed>
    <deployedseed>411788114</deployedseed>
    <deployedseed>403306236</deployedseed>
    <deployedseed>869316976</deployedseed>
    <deployedseed>584605846</deployedseed>
    <deployedseed>2017442244</deployedseed>
    <deployedseed>495128563</deployedseed>
    <deployedseed>1619283471</deployedseed>
    <deployedseed>1522795602</deployedseed>
    <deployedseed>1825102366</deployedseed>
    <deployedseed>847770146</deployedseed>
    <deployedseed>2033574599</deployedseed>
    <deployedseed>156356757</deployedseed>
    <deployedseed>80755900</deployedseed>
    <deployedseed>1727800016</deployedseed>
    <deployedseed>1784772340</deployedseed>
    <deployedseed>1728945908</deployedseed>
    <deployedseed>996935779</deployedseed>
    <deployedseed>2035715110</deployedseed>
    <deployedseed>1996445353</deployedseed>
    <deployedseed>1361983137</deployedseed>
    <deployedseed>492746565</deployedseed>
    <deployedseed>833061516</deployedseed>
    <deployedseed>1027625165</deployedseed>
    <deployedseed>1302030072</deployedseed>
    <deployedseed>523245577</deployedseed>
    <deployedseed>1453999703</deployedseed>
    <deployedseed>1832362477</deployedseed>
    <deployedseed>1994607854</deployedseed>
    <deployedseed>1344438466</deployedseed>
    <deployedseed>1252482131</deployedseed>
    <deployedseed>1863638755</deployedseed>
    <deployedseed>2118616108</deployedseed>
    <deployedseed>1633665723</deployedseed>
    <deployedseed>2100004262</deployedseed>
    <deployedseed>1323253132</deployedseed>
    <deployedseed>924467589</deployedseed>
    <deployedseed>241372917</deployedseed>
    <deployedseed>1607815998</deployedseed>
    <deployedseed>1531409858</deployedseed>
    <deployedseed>146774169</deployedseed>
    <deployedseed>799075935</deployedseed>
    <deployedseed>1915410360</deployedseed>
    <deployedseed>1522736089</deployedseed>
    <deployedseed>1764458293</deployedseed>
    <deployedseed>2089082964</deployedseed>
    <deployedseed>1287238881</deployedseed>
    <deployedseed>1186616084</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
