<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Area of a Regular Pentagon</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the area of the figure below, centered at \(O\).</p>

[[jsxgraph width="500px" height="300px"]]
    var b = {#base_length#};
    var s = {#slant#};
    
    // Abstract geometric scaling for a clean visual representation
    var R = 20; 
    var angle = Math.PI / 5; // 36 degrees
    
    // Generous bounding box to prevent clipping the top vertex or base label
    // 500x300 canvas ratio requires a wider X span to keep Aspect Ratio true
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-40, 25, 40, -25], 
        maxboundingbox: [-60, 40, 60, -40],
        axis: false, 
        showCopyright: false,
        keepAspectRatio: true,
        showNavigation: false 
    });

    // Pentagon Vertices
    var p0 = board.create('point', [0, R], {visible: false});
    var p1 = board.create('point', [-R * Math.sin(2*angle), R * Math.cos(2*angle)], {visible: false});
    var p2 = board.create('point', [-R * Math.sin(angle), -R * Math.cos(angle)], {visible: false});
    var p3 = board.create('point', [R * Math.sin(angle), -R * Math.cos(angle)], {visible: false});
    var p4 = board.create('point', [R * Math.sin(2*angle), R * Math.cos(2*angle)], {visible: false});
    var center = board.create('point', [0, 0], {name: 'O', visible: true, label:{offset:[-5,15]}});

    board.create('polygon', [p0, p1, p2, p3, p4], {
        fillColor: '#fafafa', 
        strokeColor: '#212121', 
        strokeWidth: 2
    });

    // Slanted lines (radii)
    board.create('segment', [center, p3], {strokeColor: '#1976d2'});
    board.create('segment', [center, p2], {strokeColor: '#1976d2'});

    // Perpendicular height (apothem)
    var mid = board.create('midpoint', [p2, p3], {visible: false});
    board.create('segment', [center, mid], {dash: 2, strokeColor: '#d32f2f'});

    // Right angle marker
    var rad = R * 0.15;
    board.create('angle', [p3, mid, center], {type: 'square', radius: rad, strokeColor: '#d32f2f', name: ''});

    // Dynamic Labels securely anchored
    board.create('text', [0, -R * Math.cos(angle) - 2.5, b + " cm"], {anchorX: 'middle', anchorY: 'top', fontSize: 14});
    board.create('text', [R * Math.sin(angle)/2 + 2, -R * Math.cos(angle)/2, s + " cm"], {anchorX: 'left', anchorY: 'middle', fontSize: 14, strokeColor: '#1976d2'});
[[/jsxgraph]]

<p>\(\text{Area} =\) [[input:ans1]] <b>\(\text{cm}^2\)</b> [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To calculate the area of the pentagon, divide the shape into five identical isosceles triangles by joining each vertex to the centre \(O\).</p>
<p>Each triangle has:</p>
<ul>
  <li>base length \(= {@base_length@}\text{ cm}\)</li>
  <li>two equal slanted sides of \(= {@slant@}\text{ cm}\)</li>
</ul>

<p>To find the perpendicular height (\(h\)) of one triangle, we split it into two right-angled triangles. The hypotenuse is the slanted side, and the base of this smaller triangle is exactly half of the total base:</p>
<p>
\[ \text{Half base} = \frac{{@base_length@}}{2} = {@half_base@}\text{ cm} \]
</p>

<p>Using Pythagoras’ theorem (\(a^2 + b^2 = c^2\)):</p>
<p>
\[ \begin{aligned}
h^2 + ({@half_base@})^2 &= ({@slant@})^2 \\
h^2 &= {@slant^2@} - {@half_base^2@} \\
h &= \sqrt{{@slant^2 - half_base^2@}} \\
h &= {@height@}\text{ cm}
\end{aligned} \]
</p>

<p>The area of one full isosceles triangle is therefore:</p>
<p>
\[ \begin{aligned}
\text{Area of one triangle} &= \frac{1}{2} \times \text{base} \times h \\
&= \frac{1}{2} \times {@base_length@} \times {@height@} \\
&= {@area_triangle@}\text{ cm}^2
\end{aligned} \]
</p>

<p>Since the pentagon consists of five identical triangles, its total area is:</p>
<p>
\[ \begin{aligned}
\text{Total Area} &= 5 \times {@area_triangle@} \\
\text{Total Area} &= {@ta@}\text{ cm}^2
\end{aligned} \]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Generate Pythagorean triples to ensure height and area are perfectly clean integers */
m : rand(4) + 2; 
n : rand(m-1) + 1;
k : rand(4) + 2; 

/* Assign dimensions */
half_base : k * (m^2 - n^2);
height : k * (2*m*n);

/* Randomly swap half_base and height for variety */
if rand(2) = 1 then (
    temp : half_base,
    half_base : height,
    height : temp
);

slant : k * (m^2 + n^2);
base_length : 2 * half_base;

/* Calculations */
area_triangle : (1/2) * base_length * height;
ta : 5 * area_triangle;

/* Specific Mistakes for PRT */
m1 : 5 * (1/2) * base_length * slant;     /* Used slant as height */
m2 : 5 * (1/2) * base_length * half_base; /* Used half-base as height */
m3 : area_triangle;                       /* Forgot to multiply by 5 */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Base: {@base_length@}, Slant: {@slant@}, Height: {@height@}, Area: {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct total area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Mistake: Used slant as height</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Great effort! However, the slant is not the perpendicular height. You must use Pythagoras' theorem to find the correct height first.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Mistake: Used half base as height</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Nice try. Half the base helps form the right triangle, but it is not the perpendicular height. Use Pythagoras' theorem.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Mistake: Forgot to multiply by 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Almost there! You successfully found the area of <b>one</b> triangle. Remember a pentagon is made of 5 identical triangles, so you must multiply your answer by 5.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1839088632</deployedseed>
    <deployedseed>1740798787</deployedseed>
    <deployedseed>1723749193</deployedseed>
    <deployedseed>958903117</deployedseed>
    <deployedseed>425017788</deployedseed>
    <deployedseed>815537197</deployedseed>
    <deployedseed>1564918769</deployedseed>
    <deployedseed>1319437888</deployedseed>
    <deployedseed>1297143349</deployedseed>
    <deployedseed>216100797</deployedseed>
    <deployedseed>870939523</deployedseed>
    <deployedseed>1307675360</deployedseed>
    <deployedseed>1138086982</deployedseed>
    <deployedseed>1947039655</deployedseed>
    <deployedseed>1465110896</deployedseed>
    <deployedseed>24966155</deployedseed>
    <deployedseed>1873621106</deployedseed>
    <deployedseed>716650257</deployedseed>
    <deployedseed>1476547840</deployedseed>
    <deployedseed>2090645631</deployedseed>
    <deployedseed>526863890</deployedseed>
    <deployedseed>951334633</deployedseed>
    <deployedseed>518248963</deployedseed>
    <deployedseed>380369547</deployedseed>
    <deployedseed>602955583</deployedseed>
    <deployedseed>1928377984</deployedseed>
    <deployedseed>1193396904</deployedseed>
    <deployedseed>372847929</deployedseed>
    <deployedseed>2016151772</deployedseed>
    <deployedseed>272741209</deployedseed>
    <deployedseed>488592621</deployedseed>
    <deployedseed>1397530667</deployedseed>
    <deployedseed>1864258694</deployedseed>
    <deployedseed>2101835481</deployedseed>
    <deployedseed>1503702049</deployedseed>
    <deployedseed>1215849670</deployedseed>
    <deployedseed>1747370019</deployedseed>
    <deployedseed>1110688177</deployedseed>
    <deployedseed>2017240293</deployedseed>
    <deployedseed>1724625416</deployedseed>
    <deployedseed>1056689518</deployedseed>
    <deployedseed>2016029064</deployedseed>
    <deployedseed>414261965</deployedseed>
    <deployedseed>816732193</deployedseed>
    <deployedseed>299996436</deployedseed>
    <deployedseed>184759994</deployedseed>
    <deployedseed>2000189506</deployedseed>
    <deployedseed>1975698117</deployedseed>
    <deployedseed>193427669</deployedseed>
    <deployedseed>1857207083</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
