<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Definitions of sources of data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Select the correct definition for each of the following statistical terms from the dropdown menus:</p>
<div class="card mb-3">
  <div class="card-body">
    <p><strong>(a) Data Source</strong></p>
    [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(b) Raw Data</strong></p>
    [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(c) Tertiary Sources</strong></p>
    [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(d) Primary Sources</strong></p>
    [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
  </div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<ul>
    <li><strong>Data Source:</strong> The origin or specific location from where statistical facts are obtained.</li>
    <li><strong>Raw Data:</strong> Data recorded in its original form as collected (ungrouped), without any organization or processing.</li>
    <li><strong>Tertiary Sources:</strong> These are sources that facilitate the identification and location of primary and secondary sources. Examples include bibliographies, indexes, and encyclopedias.</li>
    <li><strong>Primary Sources:</strong> Original materials or first-hand information collected directly by the researcher for a specific purpose (e.g., surveys, interviews).</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* --- 1. SETUP FOR DATA SOURCE (Input 1) --- */
correct_text_ds: "The specific origin, location, or medium from where statistical facts are obtained.";

/* Pool of wrong answers specific to Data Source */
wrongs_ds: [
    "The software tools used to analyze numerical information.", 
    "The final report generated after data analysis is complete.", 
    "A small group of people selected to represent a population.",
    "The mathematical formula used to calculate the mean.",
    "A graphical representation of frequency distribution."
];

/* Build the list: 1 Correct + 3 Random Wrongs */
/* Note: Boolean set to false for distractors to maintain logical consistency */
correct_ans_ds: rand_selection([[1, true, correct_text_ds]], 1);
wrong_ans_ds: rand_selection([
    [2, false, wrongs_ds[1]],
    [3, false, wrongs_ds[2]],
    [4, false, wrongs_ds[3]],
    [5, false, wrongs_ds[4]],
    [6, false, wrongs_ds[5]]
], 3);

/* Combine and Shuffle */
ta_ds: random_permutation(append(correct_ans_ds, wrong_ans_ds));
/* Extract just the text for the dropdown list */
final_options_ds: map(third, ta_ds);


/* --- 2. SETUP FOR RAW DATA (Input 2) --- */
correct_text_rd: "Data recorded in its original form as collected, without any organization or processing.";

wrongs_rd: [
    "Data that has been summarized into frequency tables and charts.", 
    "Conclusions and inferences drawn from a statistical study.", 
    "The average value calculated from a set of numbers.",
    "Data that has been rounded off to the nearest decimal.",
    "Information that has been sorted alphabetically."
];

correct_ans_rd: rand_selection([[1, true, correct_text_rd]], 1);
wrong_ans_rd: rand_selection([
    [2, false, wrongs_rd[1]],
    [3, false, wrongs_rd[2]],
    [4, false, wrongs_rd[3]],
    [5, false, wrongs_rd[4]],
    [6, false, wrongs_rd[5]]
], 3);

ta_rd: random_permutation(append(correct_ans_rd, wrong_ans_rd));
final_options_rd: map(third, ta_rd);


/* --- 3. SETUP FOR TERTIARY SOURCES (Input 3) --- */
correct_text_ts: "Sources that compile, index, or summarize primary and secondary sources (e.g., encyclopedias).";

wrongs_ts: [
    "First-hand accounts and original data collected directly by a researcher.", 
    "Detailed analysis and interpretation of data found in academic journals.", 
    "The raw responses collected immediately from a survey.",
    "Oral interviews conducted with participants.",
    "Government records of daily transactions."
];

correct_ans_ts: rand_selection([[1, true, correct_text_ts]], 1);
wrong_ans_ts: rand_selection([
    [2, false, wrongs_ts[1]],
    [3, false, wrongs_ts[2]],
    [4, false, wrongs_ts[3]],
    [5, false, wrongs_ts[4]],
    [6, false, wrongs_ts[5]]
], 3);

ta_ts: random_permutation(append(correct_ans_ts, wrong_ans_ts));
final_options_ts: map(third, ta_ts);


/* --- 4. SETUP FOR PRIMARY SOURCES (Input 4) --- */
correct_text_ps: "Original data or first-hand information collected directly by the researcher for a specific purpose.";

wrongs_ps: [
    "Information obtained from existing records like government censuses or books.", 
    "Abstracts and indexes used to locate other research materials.", 
    "Data that has already been manipulated and published by another author.",
    "News articles commenting on a recent event.",
    "Textbooks that summarize a specific field of study."
];

correct_ans_ps: rand_selection([[1, true, correct_text_ps]], 1);
wrong_ans_ps: rand_selection([
    [2, false, wrongs_ps[1]],
    [3, false, wrongs_ps[2]],
    [4, false, wrongs_ps[3]],
    [5, false, wrongs_ps[4]],
    [6, false, wrongs_ps[5]]
], 3);

ta_ps: random_permutation(append(correct_ans_ps, wrong_ans_ps));
final_options_ps: map(third, ta_ps);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}
{@@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">Matching definitions for Data Source, Raw Data, Tertiary Sources, and Primary Sources.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_ds</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_rd</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_ts</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta_ps</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_ds)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_rd)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_ts)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mcq_correct(ta_ps)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
