<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating area of a rectangle</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Calculate the area of the rectangle below. [[input:ans1]] \(\text{m}^2\) [[validation:ans1]][[feedback:prt1]]</p>
[[jsxgraph]]
        var board = JXG.JSXGraph.initBoard('jxgbox', {
            boundingbox: [-5, 30, 30, -10], 
            axis: false, 
            showCopyright: false,
            keepAspectRatio:true
        });

       
        var A = board.create('point', [0, 0], {name: 'A', visible: false});
        var B = board.create('point', [30, 0], {name: 'B', visible: false});
        var C = board.create('point', [30, 15], {name: 'C', visible: false});
        var D = board.create('point', [0,15], {name: 'D', visible: false});

        
        board.create('polygon', [A, B,C,D],{fillColor: 'none',strokeColor:'red'});

        
        board.create('text', [15, -2, "{#length#} m"], {fontSize: 12});  
        board.create('text', [31, 7.5, "{#width#} m"], {fontSize: 12}); 

       
        board.create('angle', [B, A, D], {
            type: 'square',
            radius: 1,
            strokeColor:'skyblue',
            fillColor: 'none',  
            label: {visible: false}
        });
        board.create('angle', [A, B, C], {
            type: 'square',
            radius: 1,
            strokeColor:'skyblue',
            fillColor: 'none',  
            label: {visible: false} 
        });
        board.create('angle', [B, C, D], {
            type: 'square',
            radius: 1,
            strokeColor:'skyblue',
            fillColor: 'none',  
             label: {visible: false} 
        });
        board.create('angle', [C, D, A], {
            type: 'square',
            radius: 1,
            strokeColor:'skyblue',
            fillColor: 'none', 
            label: {visible: false}
        });
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>The area of a rectangle is calculated by multiplying its length and width. </p>
<p>Given a rectangle with a length of \({@length@}\) meters and a width of \({@width@}\) meters, the area is found by multiplying \({@length@}\)  meters by \({@width@}\)  meters, resulting in an area of \({@area@}\)  square meters.
</p>Therefore area = \({@length@}\)  \(\times\) ({@width@}\) = \({@area@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>units:rand([m]);/*units*/
square_unit:(units)^2;/*square units*/
length:rand_with_step(55,100,2);/*length of the rectangle*/
width:rand_with_step(6,49,2);/*width of the rectangle*/
area:length*width;/*area of the rectangle*/
ta:area;/*teacher's answer*/</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>ta</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
