<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Expressing cube power results in index form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A school has a computer storage server that saves student work in blocks. Each block holds repeated copies of the same unit file. One unit file has a size of \({@base@}\) MB. After processing, the total saved size is \({@number@}\) MB.</p>

<p>Write \({@number@}\) in index form using base \({@base@}\).</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>The total saved size is \({@number@}\) MB.</p>

<p>We check how many times \({@base@}\) appears as a repeated factor:</p>

<p>\({@number@} = {@base@} \times {@base@} \times \dots \times {@base@}\) (repeated {@exp@} times)</p>

<p>Thus, the index form is:</p>

<p>\[
{@base@}^{@exp@}.
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Random base from common perfect powers */
base: rand([2,3,5,7]);

/* Random exponent (bigger than Question 1) */
exp:  rand([3,4,5]);

/* Actual number */
number: base^exp;

/* Common mistakes */
m1: base * exp;      /* multiplication instead of exponent */
m2: exp ^ base;      /* reversed exponent */
m3: base + exp;      /* addition */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Base: {@base@}</p>
<p>Exponent: {@exp@}</p>
<p>Number: {@number@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"> The question requires the learner to rewrite a large number in index form within a real-life context by recognising how many times a given base is multiplied to produce the total value.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>base^exp</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>base^exp</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>113844411</deployedseed>
    <deployedseed>1374825739</deployedseed>
    <deployedseed>1696156580</deployedseed>
    <deployedseed>1388243201</deployedseed>
    <deployedseed>611240069</deployedseed>
    <deployedseed>1357438451</deployedseed>
    <deployedseed>2006968171</deployedseed>
    <deployedseed>1478302477</deployedseed>
    <deployedseed>678440851</deployedseed>
    <deployedseed>1219367446</deployedseed>
    <deployedseed>1557771426</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>base^exp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
