<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Lines of Symmetry of a Regular polygon.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A <b>line of symmetry</b> divides a shape into two halves that are mirror images of each other. Work out the number of lines of symmetry for each regular polygon shown below.</p>
<p><b>Tip:</b> drag the red point to rotate the dashed axis about the centre, then press <b>Reflect</b>. If the orange copy lands exactly on the polygon, the dashed line is a line of symmetry. Press <b>Reset</b> to start again.</p>
<p>(a)</p>
[[jsxgraph height='380px' width='460px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-2, 12, 16, -2],
    keepaspectratio: true,
    axis: false,
    showNavigation: false,
    showCopyright: false,
    pan: { enabled: false },
    zoom: { enabled: false }
});

var n  = {#n1#};
var cx = 7, cy = 6, r = 4;
var rot = Math.random() * 360;
function deg(d){ return d * Math.PI / 180; }

// Polygon (vertices hidden, only the filled shape shows)
var verts = [];
for (var i = 0; i < n; i++) {
    var th = 90 + i * 360 / n + rot;
    verts.push(board.create('point',
        [cx + r * Math.cos(deg(th)), cy + r * Math.sin(deg(th))],
        { visible: false, fixed: true }));
}
board.create('polygon', verts, {
    fillColor: '#1f77b4', fillOpacity: 0.30,
    borders: { strokeColor: '#1f77b4', strokeWidth: 3 },
    fixed: true, vertices: { visible: false }
});

// Draggable axis constrained through the centre
var C = board.create('point', [cx, cy], { name: '', size: 2, color: '#555', fixed: true });
var guide = board.create('circle', [C, r + 1.6], { visible: false });
var D = board.create('glider', [cx + (r + 1.6), cy, guide],
    { name: 'drag', size: 4, color: '#d62728', label: { offset: [10, 8] }, withLabel: true });
var mirror = board.create('line', [C, D],
    { strokeColor: '#d62728', dash: 2, strokeWidth: 2, fixed: true });

// Reflected copy (hidden until the button is pressed)
var refl = board.create('transform', [mirror], { type: 'reflect' });
var rverts = verts.map(function (P) {
    return board.create('point', [P, refl], { visible: false });
});
var rpoly = board.create('polygon', rverts, {
    fillColor: '#ff7f0e', fillOpacity: 0.35,
    borders: { strokeColor: '#ff7f0e', strokeWidth: 3, dash: 2 },
    visible: false, fixed: true, vertices: { visible: false }
});

function setReflVisible(v) {
    rpoly.setAttribute({ visible: v });
    for (var i = 0; i < rpoly.borders.length; i++) {
        rpoly.borders[i].setAttribute({ visible: v });
    }
    board.update();
}

var shown = false;
var reflectBtn = board.create('button', [-1.3, 0.4, 'Reflect', function () {
    shown = !shown;
    setReflVisible(shown);
    reflectBtn.rendNodeButton.innerHTML = shown ? 'Hide reflection' : 'Reflect';
}], { fixed: true });

board.create('button', [5.2, 0.4, 'Reset', function () {
    D.moveTo([cx + (r + 1.6), cy]);
    shown = false;
    setReflVisible(false);
    reflectBtn.rendNodeButton.innerHTML = 'Reflect';
}], { fixed: true });

board.create('text', [cx, 11.3, 'Drag the red point, then press Reflect'],
    { anchorX: 'middle', fontSize: 12, fixed: true, color: '#333' });
[[/jsxgraph]]
<div>Number of lines of symmetry: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<p>(b)</p>
[[jsxgraph height='380px' width='460px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-2, 12, 16, -2],
    keepaspectratio: true,
    axis: false,
    showNavigation: false,
    showCopyright: false,
    pan: { enabled: false },
    zoom: { enabled: false }
});

var n  = {#n2#};
var cx = 7, cy = 6, r = 4;
var rot = Math.random() * 360;
function deg(d){ return d * Math.PI / 180; }

var verts = [];
for (var i = 0; i < n; i++) {
    var th = 90 + i * 360 / n + rot;
    verts.push(board.create('point',
        [cx + r * Math.cos(deg(th)), cy + r * Math.sin(deg(th))],
        { visible: false, fixed: true }));
}
board.create('polygon', verts, {
    fillColor: '#2ca02c', fillOpacity: 0.30,
    borders: { strokeColor: '#2ca02c', strokeWidth: 3 },
    fixed: true, vertices: { visible: false }
});

var C = board.create('point', [cx, cy], { name: '', size: 2, color: '#555', fixed: true });
var guide = board.create('circle', [C, r + 1.6], { visible: false });
var D = board.create('glider', [cx + (r + 1.6), cy, guide],
    { name: 'drag', size: 4, color: '#d62728', label: { offset: [10, 8] }, withLabel: true });
var mirror = board.create('line', [C, D],
    { strokeColor: '#d62728', dash: 2, strokeWidth: 2, fixed: true });

var refl = board.create('transform', [mirror], { type: 'reflect' });
var rverts = verts.map(function (P) {
    return board.create('point', [P, refl], { visible: false });
});
var rpoly = board.create('polygon', rverts, {
    fillColor: '#ff7f0e', fillOpacity: 0.35,
    borders: { strokeColor: '#ff7f0e', strokeWidth: 3, dash: 2 },
    visible: false, fixed: true, vertices: { visible: false }
});

function setReflVisible(v) {
    rpoly.setAttribute({ visible: v });
    for (var i = 0; i < rpoly.borders.length; i++) {
        rpoly.borders[i].setAttribute({ visible: v });
    }
    board.update();
}

var shown = false;
var reflectBtn = board.create('button', [-1.3, 0.4, 'Reflect', function () {
    shown = !shown;
    setReflVisible(shown);
    reflectBtn.rendNodeButton.innerHTML = shown ? 'Hide reflection' : 'Reflect';
}], { fixed: true });

board.create('button', [5.2, 0.4, 'Reset', function () {
    D.moveTo([cx + (r + 1.6), cy]);
    shown = false;
    setReflVisible(false);
    reflectBtn.rendNodeButton.innerHTML = 'Reflect';
}], { fixed: true });

board.create('text', [cx, 11.3, 'Drag the red point, then press Reflect'],
    { anchorX: 'middle', fontSize: 12, fixed: true, color: '#333' });
[[/jsxgraph]]
<div>Number of lines of symmetry: [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>A regular polygon with \(n\) sides has exactly \(n\) lines of symmetry &mdash; one for every vertex. (When \(n\) is odd each axis runs from a vertex to the midpoint of the opposite side; when \(n\) is even, half the axes join opposite vertices and half join the midpoints of opposite sides.)</p>
<p>(a) The regular \({@n1@}\)-gon has \({@ta1@}\) lines of symmetry, shown below.</p>
[[jsxgraph height='360px' width='440px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-2, 12, 16, -2],
    keepaspectratio: true, axis: false,
    showNavigation: false, showCopyright: false,
    pan: { enabled: false }, zoom: { enabled: false }
});
var n = {#n1#};
var cx = 7, cy = 6, r = 4;
var rot = Math.random() * 360;
function deg(d){ return d * Math.PI / 180; }

var verts = [];
for (var i = 0; i < n; i++) {
    var th = 90 + i * 360 / n + rot;
    verts.push(board.create('point',
        [cx + r * Math.cos(deg(th)), cy + r * Math.sin(deg(th))],
        { visible: false, fixed: true }));
}
board.create('polygon', verts, {
    fillColor: '#cfe2f3', fillOpacity: 0.5,
    borders: { strokeColor: '#000', strokeWidth: 3 },
    fixed: true, vertices: { visible: false }
});

var palette = ['#e41a1c','#377eb8','#4daf4a','#984ea3','#ff7f00','#a65628','#f781bf','#999999'];
for (var k = 0; k < n; k++) {
    var a = deg(90 + k * 180 / n + rot);
    var A = board.create('point', [cx + r*Math.cos(a), cy + r*Math.sin(a)], { visible: false, fixed: true });
    var B = board.create('point', [cx - r*Math.cos(a), cy - r*Math.sin(a)], { visible: false, fixed: true });
    board.create('segment', [A, B], { strokeColor: palette[k % palette.length], strokeWidth: 2, fixed: true });
}
[[/jsxgraph]]
<p>(b) The regular \({@n2@}\)-gon has \({@ta2@}\) lines of symmetry, shown below.</p>
[[jsxgraph height='360px' width='440px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-2, 12, 16, -2],
    keepaspectratio: true, axis: false,
    showNavigation: false, showCopyright: false,
    pan: { enabled: false }, zoom: { enabled: false }
});
var n = {#n2#};
var cx = 7, cy = 6, r = 4;
var rot = Math.random() * 360;
function deg(d){ return d * Math.PI / 180; }

var verts = [];
for (var i = 0; i < n; i++) {
    var th = 90 + i * 360 / n + rot;
    verts.push(board.create('point',
        [cx + r * Math.cos(deg(th)), cy + r * Math.sin(deg(th))],
        { visible: false, fixed: true }));
}
board.create('polygon', verts, {
    fillColor: '#d9ead3', fillOpacity: 0.5,
    borders: { strokeColor: '#000', strokeWidth: 3 },
    fixed: true, vertices: { visible: false }
});

var palette = ['#e41a1c','#377eb8','#4daf4a','#984ea3','#ff7f00','#a65628','#f781bf','#999999'];
for (var k = 0; k < n; k++) {
    var a = deg(90 + k * 180 / n + rot);
    var A = board.create('point', [cx + r*Math.cos(a), cy + r*Math.sin(a)], { visible: false, fixed: true });
    var B = board.create('point', [cx - r*Math.cos(a), cy - r*Math.sin(a)], { visible: false, fixed: true });
    board.create('segment', [A, B], { strokeColor: palette[k % palette.length], strokeWidth: 2, fixed: true });
}
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>perm : random_permutation([3, 4, 5, 6]);
n1 : perm[1];
n2 : perm[2];
ta1 : n1;
ta2 : n2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>(a) regular {@n1@}-gon &rarr; {@ta1@} lines; (b) regular {@n2@}-gon &rarr; {@ta2@} lines.</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1/2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have counted only one family of axes &mdash; for example just the diagonals, or just the lines through the midpoints of the sides. A regular polygon also has the <em>other</em> family. Use the <b>Reflect</b> tool to rotate the axis right around and count <em>every</em> position where the copy lands exactly on the shape.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Every line of symmetry passes through the centre. Use the <b>Reflect</b> tool: rotate the dashed axis and count each position where the reflected copy sits exactly on top of the polygon.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2/2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have counted only one family of axes &mdash; for example just the diagonals, or just the lines through the midpoints of the sides. A regular polygon also has the <em>other</em> family. Use the <b>Reflect</b> tool to rotate the axis right around and count <em>every</em> position where the copy lands exactly on the shape.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Every line of symmetry passes through the centre. Use the <b>Reflect</b> tool: rotate the dashed axis and count each position where the reflected copy sits exactly on top of the polygon.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2049988493</deployedseed>
    <deployedseed>1029871755</deployedseed>
    <deployedseed>935762523</deployedseed>
    <deployedseed>369071232</deployedseed>
    <deployedseed>570600636</deployedseed>
    <deployedseed>1305793817</deployedseed>
    <deployedseed>1627827484</deployedseed>
    <deployedseed>351637518</deployedseed>
    <deployedseed>956469304</deployedseed>
    <deployedseed>271497181</deployedseed>
    <deployedseed>1237799226</deployedseed>
    <deployedseed>1477996121</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both parts correct: n1 and n2 lines of symmetry.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Generic wrong answer (n+1, always an integer) routes to the generic hint branch on every variant.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
