<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Volume of a Cube</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the volume of the figure below.</p>
<!DOCTYPE html>
<html>
<head>
  <title>Sector OAB</title>
  <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraphcore.js"></script>
  <link rel="stylesheet" href="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraph.css" />
</head>
<body>

    <div id="box" class="jxgbox" style="width: 500px; height: 500px;"></div>
    <script type="text/javascript">
        const board = JXG.JSXGraph.initBoard('box', {
            boundingbox: [-2, 5, 5, -2], showCopyright: false,
            
            axis: false, showNavigation: false
        });

        const side = 2.5;

        // Define points
        let A = board.create('point', [0, 0], {visible: false});
        let B = board.create('point', [side, 0], {visible: false});
        let C = board.create('point', [side, side], {visible: false});
        let D = board.create('point', [0, side], {visible: false});

        let A1 = board.create('point', [0.8, 1], {visible: false,'fixed':true});
        let B1 = board.create('point', [side + 0.8, 1], {visible: false,'fixed':true});
        let C1 = board.create('point', [side + 0.8, side + 1], {visible: false,'fixed':true});
        let D1 = board.create('point', [0.8, side + 1], {visible: false,'fixed':true});

        // Create edges
        board.create('polygon', [A, B, C, D], {borders: {strokeColor: 'blue','fixed':true}});
        board.create('polygon', [A1, B1, C1, D1], {borders: {strokeColor: 'blue','fixed':true}});
        board.create('line', [A, A1], {straightFirst: false, straightLast: false, strokeColor: 'blue','fixed':true});
        board.create('line', [B, B1], {straightFirst: false, straightLast: false, strokeColor: 'blue','fixed':true});
        board.create('line', [C, C1], {straightFirst: false, straightLast: false, strokeColor: 'blue','fixed':true});
        board.create('line', [D, D1], {straightFirst: false, straightLast: false, strokeColor: 'blue','fixed':true});

        // Label the 3 visible edges with 13 cm
        board.create('text', [1.2, -0.2, "{#Side#} cm"],{'fixed':true});
        board.create('text', [3.5, 2, "{#Side#} cm"],{'fixed':true});
        board.create('text', [3, 0.4, "{#Side#} cm"],{'fixed':true});
    </script>


</body>
</html>

<p>\(\text{Volume}=\)[[input:ans1]]\(\text{cm}^2\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The volume of the cube is given by:</p>
<p>\(\text{Volume}=\text{Side}^3\)</p>
<p>Therefore,</p>
<p>\(\begin{align}
\text{Volume} &=\text{Side}^3 \\
&={@Side@}^3 \\
&= {@Side@} \times {@Side@} \times {@Side@} \\
&={@Volume@}\,\text{cm}^3

\end{align}\)</p>
<p></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Random integer side between 5 and 60, then compute volume */
Side: rand_with_step(5,65,5);   
V: Side^3;
simp:true;
Volume:V;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The Side of the cube is }={@Side@}\,\text{cm}\)</p>
<p>\(\text{The volume of the cube is }={@Volume@}\,\text{cm}^2\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Volume</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct volume of the cube.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Volume</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>649816401</deployedseed>
    <deployedseed>1002337883</deployedseed>
    <deployedseed>1868479943</deployedseed>
    <deployedseed>1972681787</deployedseed>
    <deployedseed>1189055306</deployedseed>
    <deployedseed>538643759</deployedseed>
    <deployedseed>47965473</deployedseed>
    <deployedseed>1096163115</deployedseed>
    <deployedseed>725543799</deployedseed>
    <deployedseed>892683946</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>Volume</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
