<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interpreting a given displacement time graph</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Use the displacement–time graph below to answer the following questions.
</p>

[[jsxgraph width="700px" height="450px"]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-1, 65, 11, -5],
    axis: true,
    grid: true,
    showNavigation: false,
    showCopyright: false,
   defaultAxes: {
    x: { 
        name: 'Time (s)', 
        withLabel: true,
        label: { position: 'cb',
offset: [0, 5] }
    },
    y: { 
        name: 'Displacement (m)', 
        withLabel: true,
        label: { 
             position: 'lm',
            offset: [-30, 0],
            rotate: 90 
        }
    }
}
});

// Line: constant velocity
var A = board.create('point', [0, 0], {fixed:true, size:2, name:''});
var B = board.create('point', [{# t_end #}, {# s_end #}], {fixed:true, size:2, name:''});

board.create('segment', [A, B], {strokeWidth:2});

// Highlighted point
var P = board.create('point', [{# t_mark #}, {# s_mark #}], {
    fixed:true,
    size:3,
    color:'black',
    name:''
});

// Guide lines
board.create('dashline', [[{# t_mark #}, 0], P]);
board.create('dashline', [[0, {# s_mark #}], P]);

// Numeric labels
board.create('text', [{# t_mark #}, -4, '{# t_mark #} s']);
board.create('text', [-0.8, {# s_mark #}, '{# s_mark #} m'], {anchorX:'right'});
[[/jsxgraph]]


<hr>

<p><strong>(a)</strong> What type of motion does the graph represent? Explain your answer.</p>
<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>

<p><strong>(b)</strong> What is the displacement of the car at \(t = 8\) seconds?</p>
<p>
[[input:ans2]] m [[validation:ans2]] [[feedback:prt2]]
</p>

<p><strong>(c)</strong> What is the total displacement at \(t = 10\) seconds?</p>
<p>
[[input:ans3]] m [[validation:ans3]] [[feedback:prt3]]
</p>

<p><strong>(d)</strong> Determine the velocity of the car from the graph.</p>
<p>
[[input:ans4]] m/s [[validation:ans4]] [[feedback:prt4]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<ul>
<li>The graph shows a straight line with a constant slope, indicating uniform motion. This means the car is moving at a constant velocity with no acceleration.</li>
<li>At \(t = 8\) s, the displacement is \(48\) m.</li>
<li>At \(t = 10\) s, the displacement is \(60\) m.</li>
<li>The velocity is given by the slope:
\[
v = \frac{48}{8} = 6\ \text{m/s}.
\]
</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fixed parameters matching the graph */
t_mark : 8;        /* seconds */
s_mark : 48;       /* metres */
t_end  : 10;       /* seconds */

/* Constant velocity */
v : s_mark / t_mark;

/* Total displacement at t = 10 s */
s_end : v * t_end;

/* Teacher answers */
ta_s8     : s_mark;
ta_s10    : s_end;
ta_v      : v;



/* MCQ options for part (a) */
mc_motion : [
  [1, true,  castext("Constant velocity")],
  [2, false, castext("Constant acceleration")],
  [3, false, castext("Increasing acceleration")],
  [4, false, castext("Object at rest")]
];

]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta_s8@}</p>
<p>{@ta_v @}</p>
<p>{@t_s10 @}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>mc_motion</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_s8</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_s10 </tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta_v </tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have failed to identify the type of motion from the graph, remember a constant acceleration has a changing velocity while Increasing acceleration would show as a curve with increasing slope and an object at rest would show as a horizontal line indicating velocity is zero. Kindly check the worked solution below.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_s8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have failed to correctly read the displacement at the time \(t = 8\). Kindly check the worked solution below.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta_s10 </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You've entered the velocity, not the displacement. Check the vertical axis label: it's displacement (m), not velocity.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta_v </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have failed give the correct velocity from the graph. Remember velocity is the slope of displacement-time graph</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
