<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding Trigonometric Ratios in a Right-Angled Triangle</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The right-angled triangle \(ABC\) below has a right angle at \(B\). The side \(AB\) has length \({@AB@}\) and the side \(BC\) has length \({@BC@}\).</p>
<p>Find the exact values of \(\sin \theta\), \(\cos \theta\), and \(\tan \theta\). Enter exact values, using e.g. <tt>sqrt(2)</tt> for \(\sqrt{2}\).</p>
[[jsxgraph]]
/* Triangle drawn to the true proportions of this variant. */
var ab = {#AB#};                    /* vertical leg AB (opposite theta) */
var bc = {#BC#};                    /* horizontal leg BC (adjacent), always the longer leg */
var s  = 5 / bc;                    /* scale the longer leg to 5 board units */
var w  = bc * s;
var h  = ab * s;

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-1.4, h + 0.9, w + 1.3, -1.0],
    axis: false,
    keepaspectratio: true,
    showNavigation: false,
    showCopyright: false
});

var B = board.create('point', [0, 0], {name: 'B', fixed: true, visible: false, label: {offset: [-15, -12], visible: true}});
var C = board.create('point', [w, 0], {name: 'C', fixed: true, visible: false, label: {offset: [10, -12], visible: true}});
var A = board.create('point', [0, h], {name: 'A', fixed: true, visible: false, label: {offset: [-15, 10], visible: true}});

board.create('segment', [B, C]);
board.create('segment', [B, A]);
board.create('segment', [A, C]);

/* Right-angle mark at B */
board.create('angle', [C, B, A], {radius: 0.35, type: 'square', name: ''});

/* Angle theta at C */
board.create('angle', [A, C, B], {radius: 0.6, name: '&theta;'});

/* Side-length labels, anchored away from the segments */
board.create('text', [-0.25, h / 2, '{#AB#}'], {anchorX: 'right', anchorY: 'middle', fixed: true});
board.create('text', [w / 2, -0.35, '{#BC#}'], {anchorX: 'middle', anchorY: 'top', fixed: true});
[[/jsxgraph]]
<div>\(\sin \theta = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<div>\(\cos \theta = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>
<div>\(\tan \theta = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In a right-angled triangle the three trigonometric ratios are defined relative to the angle (SOH-CAH-TOA): sine is the ratio of the side <em>opposite</em> the angle to the <em>hypotenuse</em>, cosine is the <em>adjacent</em> side over the hypotenuse, and tangent is the opposite side over the adjacent side. Here \(AB\) is opposite \(\theta\), \(BC\) is adjacent to \(\theta\), and the hypotenuse is \(AC\).</p>
<p>Only the hypotenuse needs to be calculated. By Pythagoras' theorem, \[AC = \sqrt{AB^2 + BC^2} = \sqrt{{@AB@}^2 + {@BC@}^2} = {@AC@}.\] The trigonometric ratios are then \begin{align*}\sin(\theta) &amp;= \frac{AB}{AC} = \frac{{@AB@}}{{@AC@}} = {@ta1@},\\ \cos(\theta) &amp;= \frac{BC}{AC} = \frac{{@BC@}}{{@AC@}} = {@ta2@}, \ \text{and}\\ \tan(\theta) &amp;= \frac{AB}{BC} = \frac{{@AB@}}{{@BC@}} = {@ta3@}.\end{align*}</p>]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Design backwards: choose the two legs; everything else follows.        */
/* AB is the vertical leg (opposite theta at C); BC is horizontal (adjacent). */
AB : 2 + rand(5);              /* 2 .. 6 */
BC : AB + 1 + rand(4);         /* AB+1 .. AB+4 : strictly longer than AB, so   */
                               /* sin(theta) # cos(theta) (the swap diagnostic */
                               /* is safe) and the diagram is always wide.     */
AC : sqrt(AB^2 + BC^2);        /* hypotenuse, usually a surd */

/* Model answers */
ta1 : AB/AC;                   /* sin(theta) */
ta2 : BC/AC;                   /* cos(theta) */
ta3 : AB/BC;                   /* tan(theta) */

/* Anticipated wrong answer: reciprocal of the tangent ratio */
wrong_tan_recip : BC/AB;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(AB={@AB@}\), \(BC={@BC@}\), \(AC={@AC@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the sine ratio correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Diagnostic: did the student use the adjacent side (i.e. give the cosine)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You appear to have used the <em>adjacent</em> side in your ratio. For sine, use the side <em>opposite</em> the angle: \(\sin\theta = \dfrac{\text{opposite}}{\text{hypotenuse}}\). Look at the diagram again: which of the two labelled sides does <em>not</em> touch the angle \(\theta\)?</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>First find the hypotenuse \(AC\) using Pythagoras' theorem, then use \(\sin\theta = \dfrac{\text{opposite}}{\text{hypotenuse}}\) (the SOH in SOH-CAH-TOA). The side opposite \(\theta\) is the one that does not touch the angle.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the cosine ratio correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The cosine ratio is \(\cos\theta = \dfrac{\text{adjacent}}{\text{hypotenuse}}\) (the CAH in SOH-CAH-TOA). The adjacent side runs from \(\theta\) to the right angle, and the hypotenuse — found by Pythagoras' theorem — is the side opposite the right angle.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the tangent ratio correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Diagnostic: reciprocal of the tangent ratio (adjacent over opposite)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>wrong_tan_recip</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is the reciprocal of the required ratio: you have divided the adjacent side by the opposite side. The tangent ratio is \(\tan\theta = \dfrac{\text{opposite}}{\text{adjacent}}\) — check which way up your fraction should be.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Use \(\tan\theta = \dfrac{\text{opposite}}{\text{adjacent}}\) (the TOA in SOH-CAH-TOA). Note that the hypotenuse is not involved: both sides in the ratio are given directly in the diagram.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1027259654</deployedseed>
    <deployedseed>666866064</deployedseed>
    <deployedseed>457220572</deployedseed>
    <deployedseed>636240078</deployedseed>
    <deployedseed>2070352536</deployedseed>
    <deployedseed>1568804558</deployedseed>
    <deployedseed>194538831</deployedseed>
    <deployedseed>1135656153</deployedseed>
    <deployedseed>1727439596</deployedseed>
    <deployedseed>412072013</deployedseed>
    <deployedseed>1868679708</deployedseed>
    <deployedseed>303595769</deployedseed>
    <deployedseed>1876310260</deployedseed>
    <deployedseed>793276716</deployedseed>
    <deployedseed>1650476901</deployedseed>
    <deployedseed>22881685</deployedseed>
    <deployedseed>1536292791</deployedseed>
    <deployedseed>1604668073</deployedseed>
    <deployedseed>1017374395</deployedseed>
    <deployedseed>2146328873</deployedseed>
    <deployedseed>34033514</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All three ratios correct: full marks on every part.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Targeted wrong answers: sin/cos swapped in part 1; cosine given as sine in part 2 (generic wrong); reciprocal tangent in part 3.</description>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>wrong_tan_recip</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Generic wrong answers that match no diagnostic: catch-all branches fire.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
