<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interpreting of histograms of data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A meteorological station recorded the rainfall pattern for a specific region over a period of 12 months. The data was grouped into unequal time intervals based on seasonal shifts.</p>
<p>The histogram below represents this data. The \(x\)-axis represents Time in Months (0 = Start of Jan, 12 = End of Dec), and the \(y\)-axis represents the Frequency Density (Rainfall Intensity in mm/month).</p>

<p>
[[jsxgraph width="750px" height="450px"]]
    JXG.Options.text.useMathJax = true;
    
    // 1. IMPORT DATA
    var x = [0, {#x1#}, {#x2#}, {#x3#}, {#x4#}, {#x5#}];
    var h = [{#h1#}, {#h2#}, {#h3#}, {#h4#}, {#h5#}];
    
    // Determine bounds
    var max_h = Math.max(...h);
    var max_x = 12;

    // 2. INIT BOARD
    // Grid:true gives the background lines. 
    // boundingBox padded to make axes clear.
    var board = JXG.JSXGraph.initBoard(divid, {
        boundingBox: [-1, max_h + 2, max_x + 1, -1],
        grid: true,
        showNavigation: false,
        showCopyright: false,
        zoom: {enabled: false},
        pan: {enabled: false}
    });

    // 3. AXES
    // Ticks set to distance 1 to match the integer heights
    var yAxis = board.create("axis", [[0, 0], [0, max_h + 2]], {
        name: "\\(Density\\)",
        withLabel: true,
        label: { offset: [-25, 0] },
        ticks: { distance: 1, minorTicks: 0 } 
    });

    var xAxis = board.create("axis", [[0, 0], [max_x + 1, 0]], {
        name: "\\(Months\\)",
        withLabel: true,
        label: { offset: [-10, -20] },
        ticks: { distance: 1 }
    });

    // 4. DRAW BARS (Histogram)
    for(let i=0; i<5; i++){
        board.create('polygon', [
            [x[i], 0], [x[i+1], 0], [x[i+1], h[i]], [x[i], h[i]]
        ], {
            fixed: true,          // Prevent dragging
            fillColor: '#2E8B57', // SeaGreen
            fillOpacity: 0.5,
            borders: { 
                strokeColor: 'black', 
                strokeWidth: 2,
                fixed: true 
            },
            vertices: { 
                visible: false,   // No corner dots
                fixed: true 
            }
        });
    }
    // No Frequency Polygon drawn.

[[/jsxgraph]]
</p>

<p><b>(a)</b> The "Monsoon" period occurs in the interval \({@x2@} < t \leq {@x3@}\). Using the frequency density shown on the graph, calculate the total rainfall (in mm) that fell during this specific month.</p>
<p>[[input:ans1]] [[validation:ans1]]</p>

<p><b>(b)</b> Estimate the cumulative rainfall (total amount of rain) recorded from the start of the year up to the middle of the 8th month (i.e., \(t = 7.5\)).</p>
<p>[[input:ans2]] [[validation:ans2]]</p>

<p><b>(c)</b> Calculate what percentage of the total annual rainfall occurred during the "Monsoon" interval (\({@x2@} < t \leq {@x3@}\)). (Give your answer as a percentage rounded to 2 decimal places).</p>
<p>[[input:ans3]]% [[validation:ans3]]</p>

<p><b>(d)</b> Calculate the <b>mean monthly rainfall</b> for the entire year.</p>
<p>[[input:ans4]] [[validation:ans4]]</p>

<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p><b>Understanding the Graph</b></p>
<p>In this histogram, the Area of each bar represents the total quantity of rainfall (Frequency) for that period.</p>
\( \text{Total Rainfall (Area)} = \text{Width (Months)} \times \text{Frequency Density (Height)} \)

<p><b>(a) Monsoon Rainfall (Interval \({@x2@} < t \leq {@x3@}\))</b></p>
<ul>
    <li><b>Width:</b> The interval is from {@x2@} to {@x3@}, so width = \({@w3@}\) month.</li>
    <li><b>Height:</b> Reading from the graph, the Frequency Density is {@h3@} mm/mo.</li>
    <li><b>Calculation:</b> \( \text{Rainfall} = {@w3@} \times {@h3@} = \mathbf{@freq3@} \text{ mm} \).</li>
</ul>

<p><b>(b) Estimation up to \(t = 7.5\)</b></p>
<p>We need the area of the graph from \(t=0\) to \(t=7.5\).</p>
<ol>
    <li><b>Complete Intervals:</b>
    <br>Months 0–{@x1@}: \( {@w1@} \times {@h1@} = {@freq1@} \)
    <br>Months {@x1@}–{@x2@}: \( {@w2@} \times {@h2@} = {@freq2@} \)
    <br>Months {@x2@}–{@x3@}: \( {@w3@} \times {@h3@} = {@freq3@} \)
    <br><i>Subtotal = {@rain_prev@} mm</i>
    </li>
    <li><b>Partial Interval:</b>
    <br>The cut-off \(t=7.5\) is inside the interval \({@x3@} < t \leq {@x4@}\).
    <br>Width of partial section = \(7.5 - {@x3@} = {@p_width@}\).
    <br>Height of this bar = {@h4@}.
    <br>Partial Rain = \({@p_width@} \times {@h4@} = {@rain_partial@}\) mm.
    </li>
</ol>
<p><b>Total Estimate:</b> \({@rain_prev@} + {@rain_partial@} = \mathbf{@ta2@} \text{ mm}\).</p>

<p><b>(c) Percentage Calculation</b></p>
<p>First, calculate the Total Annual Rainfall by summing the area of all 5 bars:</p>
\( \text{Total} = {@freq1@} + {@freq2@} + {@freq3@} + {@freq4@} + {@freq5@} = {@total_rain@} \text{ mm} \)
<p>The percentage contribution of the Monsoon interval is:</p>
\( \frac{@freq3@}{@total_rain@} \times 100 = \mathbf{@ta3_disp@}\% \)

<p><b>(d) Mean Monthly Rainfall</b></p>
<p>The mean monthly rainfall is the Total Rainfall divided by the total number of months (12).</p>
\( \text{Mean} = \frac{@total_rain@}{12} = \mathbf{@ta4_disp@} \text{ mm/month} \)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. SETUP INTERVALS (Months 0-12) */
/* Unequal widths to force use of Frequency Density */
x0: 0;
x1: 2;   /* Jan-Feb (2 months) */
x2: 5;   /* Mar-May (3 months) */
x3: 6;   /* Jun (1 month - wet season start) */
x4: 9;   /* Jul-Sep (3 months) */
x5: 12;  /* Oct-Dec (3 months) */

w1: x1 - x0;
w2: x2 - x1;
w3: x3 - x2;
w4: x4 - x3;
w5: x5 - x4;

/* 2. RANDOMIZE HEIGHTS (Frequency Density) */
/* Using integers 1-9 for clear y-axis readability */
h1: rand([2,3]);
h2: rand([1,2]);
h3: rand([7,8,9]); /* The "Monsoon" peak */
h4: rand([4,5,6]);
h5: rand([2,3]);

/* 3. CALCULATE FREQUENCIES (Total Rainfall per interval) */
/* Area = Width * Height */
freq1: w1 * h1;
freq2: w2 * h2;
freq3: w3 * h3;
freq4: w4 * h4;
freq5: w5 * h5;

total_rain: freq1 + freq2 + freq3 + freq4 + freq5;

/* 4. DEFINE TEACHER ANSWERS (ta) */

/* Q1: Rain in Monsoon month (Interval 3: x2 to x3) */
ta1: freq3;

/* Q2: Estimate cumulative rain up to T = 7.5 months */
T_cut: 7.5;
/* Sum of full bars before interval 4 */
rain_prev: freq1 + freq2 + freq3;
/* Partial bar calculation for interval 4 */
p_width: T_cut - x3;
rain_partial: p_width * h4;
ta2: rain_prev + rain_partial;

/* Q3: Percentage of total rain in the Monsoon month */
ta3: (freq3 / total_rain) * 100;

/* Q4: Mean Monthly Rainfall */
ta4: total_rain / 12;

/* Formatting for display (rounded) */
ta3_disp: float(round(ta3*100)/100);
ta4_disp: float(round(ta4*100)/100);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@1@}</p>
<p>{@1@}</p>
<p>{@1@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="reviewer"></p>
<p id="reviewer"> </p>
<p id="description">
Given a histogram representing annual rainfall grouped into unequal seasonal intervals, the student must calculate specific interval totals (using Area = Width \(\times\) Density), estimate cumulative rainfall up to a non-integer time point (linear interpolation within a bin), calculate the percentage contribution of a specific season, and determine the mean monthly rainfall.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>node_test : if ta1=ans1 and ans2=ta2  and ans3=ta3 and ans4=ta4  then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer in all the inputs</description>
        <answertest>AlgEquiv</answertest>
        <sans>node_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span style="color:green; font-weight:bold;">Correct.</span> You correctly calculated the area (Width \(\times\) Height) for the monsoon month.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><span style="color:red; font-weight:bold;">Incorrect.</span> Remember that Total Rainfall is the <b>Area</b> of the bar. Multiply the width of the interval (1 month) by the density shown on the y-axis.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span style="color:green; font-weight:bold;">Correct.</span> You accurately estimated the cumulative rainfall by summing the full intervals and the partial section.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><span style="color:red; font-weight:bold;">Incorrect.</span> Calculate the area of all bars to the left of \(t=6\), then add the area of the partial rectangle from \(t=6\) to \(t=7.5\).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span style="color:green; font-weight:bold;">Correct.</span></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><span style="color:red; font-weight:bold;">Incorrect.</span> Calculate (Monsoon Rain / Total Annual Rain) \(\times\) 100.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.25</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><span style="color:green; font-weight:bold;">Correct.</span></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><span style="color:red; font-weight:bold;">Incorrect.</span> Divide the Total Annual Rainfall by 12 months.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2144083686</deployedseed>
  </question>
</quiz>
