<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of an heptagon</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A school is designing a flower garden in the shape of a regular heptagon. Each side has length <b>{#s#} m</b>, and the apothem is <b>{#a#} m</b>.</p>

[[jsxgraph width="400px" height="400px"]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-8, 8, 8, -8],
    axis: false,
    showNavigation: false,
    showCopyright: false
});

var n = 7;
var R = 5;

/* 1. Create Vertices */
var points = [];
for (var i = 0; i < n; i++) {
    /* Rotating slightly to ensure a side is clearly angled for the label */
    var angle = (2 * Math.PI * i / n) + (Math.PI / 2); 
    points.push(board.create('point', [R * Math.cos(angle), R * Math.sin(angle)], {visible: false}));
}

/* 2. Create Polygon */
var poly = board.create('polygon', points, {
    withLines: true,
    fillColor: '#dfefff',
    borders: {strokeWidth: 2, strokeColor: '#0072b2'}
});

/* 3. Center Point */
var O = board.create('point', [0, 0], {visible: false});

/* 4. Apothem Label (Rotated along the dashed line) */
var M_apo = board.create('midpoint', [points[0], points[1]], {visible: false});
board.create('segment', [O, M_apo], {dash: 2, strokeWidth: 2, strokeColor: '#0072b2'});

board.create('text', [
    (O.X() + M_apo.X()) / 2, 
    (O.Y() + M_apo.Y()) / 2 - 0.4, 
    "{#a#} m"
], {
    fontSize: 14,
    anchorX: 'middle',
    rotate: function() {
        /* Calculate angle of the apothem segment */
        return Math.atan2(M_apo.Y() - O.Y(), M_apo.X() - O.X()) * 180 / Math.PI;
    }
});

/* 5. Side Length Label (Rotated along the boundary line) */
var pA = points[1];
var pB = points[2];
var M_side = board.create('midpoint', [pA, pB], {visible: false});

board.create('text', [
    M_side.X() - 0.5, 
    M_side.Y() + 0.5, 
    "{#s#} m"
], {
    fontSize: 14,
    anchorX: 'middle',
    rotate: function() {
        /* Calculate angle of the side segment */
        var ang = Math.atan2(pB.Y() - pA.Y(), pB.X() - pA.X()) * 180 / Math.PI;
        /* Flip text if it's upside down (between 90 and 270 degrees) */
        return (ang > 90 || ang < -90) ? ang + 180 : ang;
    }
});
[[/jsxgraph]]
<p>Calculate the area of the heptagon.</p>

<p style="text-align: center;">
Area = [[input:ans1]] m<sup>2</sup> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p><strong>Step 1: Perimeter</strong></p>
<p>
A regular heptagon has 7 equal sides:
</p>
<p>
The area of a regular polygon is:
</p>
<p>
\( A = \frac{7}{2} \times \text{side length} \times \text{apothem} \)
</p>

<p>
\( A = \frac{7}{2} \times {@s@} \times {@a@} \)
</p>

<p>
\( A = {@ta1@} \text{ m}^2 \)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Side length (m) */
s : rand(4)+6;   /* 6–9 */

/* Apothem (m) */
a : rand(3)+4;   /* 4–6 */

/* Area */
ta1 : (7/2)*s*a;

</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1926318901</deployedseed>
    <deployedseed>58158914</deployedseed>
    <deployedseed>2085164398</deployedseed>
    <deployedseed>1364358825</deployedseed>
    <deployedseed>76627291</deployedseed>
    <deployedseed>999995501</deployedseed>
    <deployedseed>1374004260</deployedseed>
    <deployedseed>1439211178</deployedseed>
    <deployedseed>121764709</deployedseed>
    <deployedseed>167314452</deployedseed>
    <deployedseed>1393084378</deployedseed>
  </question>
</quiz>
