<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Using logarithms with standard form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Write the number \({@value@}\) in standard form using logarithms.</p>

<p>Express your answer in the form \(A \times 10^B\).</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>We want to express the number in standard form \(A \times 10^B\).</p>

<p>\({@value@} = {@coef@} \times 10^{@exp@}\).</p>

<p>Using logarithms:</p>

<p>\[
\log_{10}({@value@}) = \log_{10}({@coef@}) + \log_{10}(10^{@exp@})
\]</p>

<p>\[
= \log_{10}({@coef@}) + {@exp@}.
\]</p>

<p>Therefore the standard form is:</p>

<p>\[
{@coef@} \times 10^{@exp@}.
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Choose a power of 10 */
exp: rand([3,4,5,6]);

/* Number formed using power of 10 */
number: 10^exp;

/* Coefficient between 1 and 9 */
coef: rand_with_step(1,9,1);

/* Final number in expanded form */
value: coef * number;

/* Standard form components */
A: coef;
B: exp;

/* Wrong answers */
m1: coef + exp;      /* adds instead of SF */
m2: coef * exp;      /* multiplies instead */
m3: number;          /* gives full number instead of SF */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@value@} = {@coef@} × 10^{@exp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"> The question requires the learner to express a large number in standard form by identifying the power of ten using the meaning of base ten logarithms.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A*10^B</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>A*10^B</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1010947386</deployedseed>
    <deployedseed>1831155997</deployedseed>
    <deployedseed>833733291</deployedseed>
    <deployedseed>738802563</deployedseed>
    <deployedseed>468244295</deployedseed>
    <deployedseed>2022690933</deployedseed>
    <deployedseed>1504850441</deployedseed>
    <deployedseed>1553148887</deployedseed>
    <deployedseed>734238376</deployedseed>
    <deployedseed>60733463</deployedseed>
    <deployedseed>836059070</deployedseed>
    <deployedseed>702981637</deployedseed>
    <deployedseed>301329819</deployedseed>
    <deployedseed>1505628615</deployedseed>
    <deployedseed>545307751</deployedseed>
    <deployedseed>123506974</deployedseed>
    <deployedseed>2004358962</deployedseed>
    <deployedseed>1782120530</deployedseed>
    <deployedseed>540219338</deployedseed>
    <deployedseed>1539382092</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A*10^B</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
