<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Applying the product law of indices</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The Agriculture Club is doing mango processing and is making mango juice. Each jar of mango juice needs \({@base@}^{@exp1@}\) mangoes. The group makes \({@base@}^{@exp2@}\) jars of mango juice.</p>

<p>How many mangoes are needed in total?
Write your answer as a <strong>single power</strong>, in index form.
</p>

<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
The total number of mangoes is found by multiplying the number of mangoes per jar by the number of jars.
</p>

<p>
We use the product law of indices:
</p>

<p>
\[
a^m \times a^n = a^{m+n}
\]
</p>

<p>
So,
</p>

<p>
\[
{@base@}^{@exp1@} \times {@base@}^{@exp2@}
= {@base@}^{@exp1 + exp2@}
\]
</p>

<p>
Therefore, the total number of mangoes needed is:
</p>

<p>
\[
\boxed{{@base@}^{@exp_total@}}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random base and exponents */
base: rand([2,3,5,7]);
exp1: rand([1,2,3,4]);
exp2: rand([1,2,3,4]);

/* Correct simplified exponent */
exp_total: exp1 + exp2;

/* Original expression */
expr: base^exp1 * base^exp2;

/* Common student mistakes */
m1: base^(exp1*exp2);        /* multiplies exponents */
m2: base^exp1 + base^exp2;  /* adds instead of multiplying */
m3: base^(exp1 - exp2);     /* subtracts exponents */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Original expression: {@expr@}</p>
<p>Total exponent: {@exp_total@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"> This question helps learners apply the product law of indices in a real-life situation by finding the total number of items expressed using powers.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>base^exp_total</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>base^exp_total</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Mistake1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You multiplied the exponents. When bases are the same, we add exponents.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Mistake2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You added the powers instead of multiplying them</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Mistake3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>You used subtraction to arrive to your answer. Subtraction is used when dividing, not multiplying.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1065589520</deployedseed>
    <deployedseed>829002031</deployedseed>
    <deployedseed>1280944213</deployedseed>
    <deployedseed>187214546</deployedseed>
    <deployedseed>834538727</deployedseed>
    <deployedseed>384847905</deployedseed>
    <deployedseed>1729088992</deployedseed>
    <deployedseed>125889534</deployedseed>
    <deployedseed>584534641</deployedseed>
    <deployedseed>1704799368</deployedseed>
    <deployedseed>1202983379</deployedseed>
    <deployedseed>1433112642</deployedseed>
    <deployedseed>684498647</deployedseed>
    <deployedseed>672702521</deployedseed>
    <deployedseed>809444564</deployedseed>
    <deployedseed>1264279383</deployedseed>
    <deployedseed>2073465178</deployedseed>
    <deployedseed>1579859510</deployedseed>
    <deployedseed>307554636</deployedseed>
    <deployedseed>638799585</deployedseed>
    <deployedseed>2141287681</deployedseed>
    <deployedseed>6227094</deployedseed>
    <deployedseed>570461657</deployedseed>
    <deployedseed>1354425370</deployedseed>
    <deployedseed>68949546</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>base^exp_total</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
