<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Classifying primary and secondary data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Classify each of the following data collection activities as either <strong>Primary Data</strong> or <strong>Secondary Data</strong>:</p>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(a)</strong> {@scen_1@}</p>
    [[input:ans1]] [[validation:ans1]][[feedback:prt1]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(b)</strong> {@scen_2@}</p>
    [[input:ans2]] [[validation:ans2]][[feedback:prt2]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(c)</strong> {@scen_3@}</p>
    [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(d)</strong> {@scen_4@}</p>
    [[input:ans4]] [[validation:ans4]][[feedback:prt4]]
  </div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<ul>
    <li><strong>Primary Data:</strong> This is data that you collect <em>yourself</em> for a specific purpose. It is "fresh" and original. (Examples: Counting cars, interviewing people, measuring height).</li>
    <li><strong>Secondary Data:</strong> This is data that <em>someone else</em> has already collected. You are just using it. (Examples: Census reports, newspaper articles, textbooks, internet searches).</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* --- SCENARIO POOLS --- */

/* List of Primary Data Scenarios (Collection) */
pool_primary: [
    "You stand at the school gate and count the number of cars passing by.",
    "You interview 20 students to ask them about their favorite lunch meal.",
    "You measure the height of every seedling in your agriculture plot yourself.",
    "You conduct a science experiment to see how fast ice melts.",
    "You hand out a questionnaire to your neighbors about water usage."
];

/* List of Secondary Data Scenarios (Existing Sources) */
pool_secondary: [
    "You look up the population of Kenya in the 2019 Census report.",
    "You read a newspaper article to find out yesterday's football scores.",
    "You use a graph from a textbook to analyze rainfall patterns from 1990.",
    "You search the internet for the current exchange rate of the Dollar.",
    "You analyze patient records from a hospital archive for a history project."
];

/* --- SELECTION & RANDOMIZATION --- */

/* Select 4 unique scenarios: 2 Primary and 2 Secondary to ensure balance */
/* We shuffle the pools first to get random ones */
shuffled_p: random_permutation(pool_primary);
shuffled_s: random_permutation(pool_secondary);

/* Assign scenarios to the 4 parts (a, b, c, d) */
/* We mix the order: P, S, S, P or similar. Let's do P, S, P, S for simplicity in coding, but we can randomize the display order if needed. */
/* Here I map them explicitly to variables for clarity */

scen_1: shuffled_p[1]; /* Primary */
scen_2: shuffled_s[1]; /* Secondary */
scen_3: shuffled_p[2]; /* Primary */
scen_4: shuffled_s[2]; /* Secondary */


/* --- 1. SETUP FOR PART A (Primary) --- */
correct_text_1: "Primary Data";
wrong_text_1: "Secondary Data";

/* Build the list using your format */
correct_ans_1: rand_selection([[1, true, correct_text_1]], 1);
wrong_ans_1: rand_selection([[2, false, wrong_text_1]], 1);

ta_1: random_permutation(append(correct_ans_1, wrong_ans_1));
final_options_1: map(third, ta_1);


/* --- 2. SETUP FOR PART B (Secondary) --- */
correct_text_2: "Secondary Data";
wrong_text_2: "Primary Data";

correct_ans_2: rand_selection([[1, true, correct_text_2]], 1);
wrong_ans_2: rand_selection([[2, false, wrong_text_2]], 1);

ta_2: random_permutation(append(correct_ans_2, wrong_ans_2));
final_options_2: map(third, ta_2);


/* --- 3. SETUP FOR PART C (Primary) --- */
correct_text_3: "Primary Data";
wrong_text_3: "Secondary Data";

correct_ans_3: rand_selection([[1, true, correct_text_3]], 1);
wrong_ans_3: rand_selection([[2, false, wrong_text_3]], 1);

ta_3: random_permutation(append(correct_ans_3, wrong_ans_3));
final_options_3: map(third, ta_3);


/* --- 4. SETUP FOR PART D (Secondary) --- */
correct_text_4: "Secondary Data";
wrong_text_4: "Primary Data";

correct_ans_4: rand_selection([[1, true, correct_text_4]], 1);
wrong_ans_4: rand_selection([[2, false, wrong_text_4]], 1);

ta_4: random_permutation(append(correct_ans_4, wrong_ans_4));
final_options_4: map(third, ta_4);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">In this classification question, students are presented with four randomized real-world research scenarios (e.g., conducting an experiment vs. reading a government report). They must identify whether each activity represents <strong>Primary Data</strong> collection or <strong>Secondary Data</strong> collection.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mcq_correct(ta_4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>660633</deployedseed>
  </question>
</quiz>
