<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Classifying primary and secondary data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Classify each of the following data collection activities as either <strong>Primary Data</strong> or <strong>Secondary Data</strong>:</p>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(a)</strong> {@scen_1@}</p>
    [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(b)</strong> {@scen_2@}</p>
    [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(c)</strong> {@scen_3@}</p>
    [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(d)</strong> {@scen_4@}</p>
    [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]
  </div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Ask, for each activity: <em>who</em> gathered the data?</p>
<ul>
    <li><strong>Primary Data:</strong> data you collect <em>yourself</em>, first-hand, for this specific purpose. (Examples: counting cars, interviewing people, measuring plants, running an experiment.)</li>
    <li><strong>Secondary Data:</strong> data that <em>someone else</em> has already collected and published, which you are simply re-using. (Examples: a census report, a newspaper article, a textbook graph, an internet search.)</li>
</ul>
<p>For this variant:</p>
<ul>
    <li><strong>(a)</strong> {@scen_1@} &rarr; <strong>{@type_1@}</strong></li>
    <li><strong>(b)</strong> {@scen_2@} &rarr; <strong>{@type_2@}</strong></li>
    <li><strong>(c)</strong> {@scen_3@} &rarr; <strong>{@type_3@}</strong></li>
    <li><strong>(d)</strong> {@scen_4@} &rarr; <strong>{@type_4@}</strong></li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* --- SCENARIO POOLS --- */

/* Primary Data scenarios (data collected first-hand) */
pool_primary: [
    "You stand at the school gate and count the number of cars passing by.",
    "You interview 20 students to ask them about their favourite lunch meal.",
    "You measure the height of every seedling in your agriculture plot yourself.",
    "You conduct a science experiment to see how fast ice melts.",
    "You hand out a questionnaire to your neighbours about water usage."
];

/* Secondary Data scenarios (data already collected by someone else) */
pool_secondary: [
    "You look up the population of Kenya in the 2019 Census report.",
    "You read a newspaper article to find out yesterday's football scores.",
    "You use a graph from a textbook to analyse rainfall patterns from 1990.",
    "You search the internet for the current exchange rate of the Dollar.",
    "You analyse patient records from a hospital archive for a history project."
];

/* --- SELECTION --- */
/* Pick 2 distinct primary and 2 distinct secondary scenarios */
sp: random_permutation(pool_primary);
ss: random_permutation(pool_secondary);

/* Tag each chosen scenario with its correct classification */
tagged: [
    [sp[1], "Primary Data"],
    [sp[2], "Primary Data"],
    [ss[1], "Secondary Data"],
    [ss[2], "Secondary Data"]
];

/* --- RANDOMISE POSITION --- */
/* Shuffle so the correct-answer sequence varies between variants          */
/* (fixes the old fixed P,S,P,S pattern that could be answered by rote).   */
items: random_permutation(tagged);

scen_1: items[1][1];  type_1: items[1][2];
scen_2: items[2][1];  type_2: items[2][2];
scen_3: items[3][1];  type_3: items[3][2];
scen_4: items[4][1];  type_4: items[4][2];

/* --- DROPDOWN OPTIONS --- */
/* Value 1 = "Primary Data", value 2 = "Secondary Data".                   */
/* Correctness flag depends on this part's type; display order shuffled.   */
opts_for(t) := random_permutation([
    [1, is(t = "Primary Data"),   "Primary Data"],
    [2, is(t = "Secondary Data"), "Secondary Data"]
]);

ta_1: opts_for(type_1);
ta_2: opts_for(type_2);
ta_3: opts_for(type_3);
ta_4: opts_for(type_4);

/* Short summary of this variant for the question note */
variantnote: sconcat("a=", type_1, ", b=", type_2, ", c=", type_3, ", d=", type_4);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@variantnote@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">Sources of data: classifying activities as primary or secondary data collection.</p>
<p id="reviewer">Ben (STACK audit)</p>
<p id="description">In this classification question, students are presented with four randomised real-world research scenarios (e.g., conducting an experiment vs. reading a government report). They must identify whether each activity represents <strong>Primary Data</strong> collection or <strong>Secondary Data</strong> collection. The four scenarios (two primary, two secondary) are shuffled across the parts so the correct sequence differs between variants, and each wrong answer receives a diagnostic hint.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>ta_4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Ask yourself: did <em>you</em> gather this data first-hand for this task (that would make it <strong>primary</strong>), or was it already collected and published by someone else, so you are only re-using it (<strong>secondary</strong>)?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_2)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Ask yourself: did <em>you</em> gather this data first-hand for this task (that would make it <strong>primary</strong>), or was it already collected and published by someone else, so you are only re-using it (<strong>secondary</strong>)?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_3)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Ask yourself: did <em>you</em> gather this data first-hand for this task (that would make it <strong>primary</strong>), or was it already collected and published by someone else, so you are only re-using it (<strong>secondary</strong>)?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>mcq_correct(ta_4)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Ask yourself: did <em>you</em> gather this data first-hand for this task (that would make it <strong>primary</strong>), or was it already collected and published by someone else, so you are only re-using it (<strong>secondary</strong>)?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>660633</deployedseed>
    <deployedseed>121212</deployedseed>
    <deployedseed>247613</deployedseed>
    <deployedseed>883051</deployedseed>
    <deployedseed>519274</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All four parts correct: full marks on every PRT.</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta_1)[1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>mcq_correct(ta_2)[1]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>mcq_correct(ta_3)[1]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>mcq_correct(ta_4)[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>All four parts wrong (opposite classification): fires each F branch with the diagnostic.</description>
      <testinput>
        <name>ans1</name>
        <value>3-mcq_correct(ta_1)[1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>3-mcq_correct(ta_2)[1]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>3-mcq_correct(ta_3)[1]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>3-mcq_correct(ta_4)[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
