<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Definitions of sources of data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Select the correct definition for each of the following statistical terms from the dropdown menus:</p>
<div class="card mb-3">
  <div class="card-body">
    <p><strong>(a) Data</strong></p>
    [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(b) Primary Data</strong></p>
    [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
  </div>
</div>

<div class="card mb-3">
  <div class="card-body">
    <p><strong>(c) Secondary Data</strong></p>
    [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
  </div>
</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<ul>
    <li><strong>Data:</strong> a collection of facts, figures, numbers, observations, or measurements that can be recorded and used to study or understand something.</li>
    <li><strong>Primary Data:</strong> data collected first-hand by the researcher for the specific study at hand (e.g., your own surveys, experiments, or observations).</li>
    <li><strong>Secondary Data:</strong> data that has already been collected by someone else for a different purpose, which the researcher now re-uses (e.g., a census report, a textbook, a newspaper).</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Helper: position (1-3) of the correct option (value 1) within a shuffled dropdown list */
correctpos(ta) := first(sublist_indices(map(first, ta), lambda([v], is(v = 1))));
/* A wrong option's value, for question testing */
wrongval(ta)   := first(delete(1, map(first, ta)));

/* --- 1. DATA (Input 1) --- */
correct_text_da: "A collection of facts, figures, numbers, observations, or measurements that can be recorded and used to study something.";
wrongs_da: [
    "The final conclusions and recommendations drawn after a study is complete.",
    "A graph or chart used to display information visually.",
    "The average value calculated from a set of numbers.",
    "A small group of people selected to represent a population.",
    "The software tools used to analyse numerical information."
];
correct_ans_da: rand_selection([[1, true, correct_text_da]], 1);
wrong_ans_da: rand_selection([
    [2, false, wrongs_da[1]],
    [3, false, wrongs_da[2]],
    [4, false, wrongs_da[3]],
    [5, false, wrongs_da[4]],
    [6, false, wrongs_da[5]]
], 3);
ta_da: random_permutation(append(correct_ans_da, wrong_ans_da));

/* --- 2. PRIMARY DATA (Input 2) --- */
correct_text_pd: "Data collected first-hand by the researcher for the specific study at hand.";
wrongs_pd: [
    "Information obtained from existing records such as government censuses or textbooks.",
    "Data that has already been published and analysed by another author.",
    "A summary of research compiled from many other studies.",
    "News articles reporting on a past event.",
    "Figures copied from a published statistical report."
];
correct_ans_pd: rand_selection([[1, true, correct_text_pd]], 1);
wrong_ans_pd: rand_selection([
    [2, false, wrongs_pd[1]],
    [3, false, wrongs_pd[2]],
    [4, false, wrongs_pd[3]],
    [5, false, wrongs_pd[4]],
    [6, false, wrongs_pd[5]]
], 3);
ta_pd: random_permutation(append(correct_ans_pd, wrong_ans_pd));

/* --- 3. SECONDARY DATA (Input 3) --- */
correct_text_sd: "Data already collected by someone else for a different purpose, which the researcher now re-uses.";
wrongs_sd: [
    "Information a researcher collects first-hand through their own surveys or experiments.",
    "Measurements you record yourself during a science experiment.",
    "Responses gathered directly from interviewing participants.",
    "Original observations made by the investigator for this specific study.",
    "First-hand counts taken by a student during fieldwork."
];
correct_ans_sd: rand_selection([[1, true, correct_text_sd]], 1);
wrong_ans_sd: rand_selection([
    [2, false, wrongs_sd[1]],
    [3, false, wrongs_sd[2]],
    [4, false, wrongs_sd[3]],
    [5, false, wrongs_sd[4]],
    [6, false, wrongs_sd[5]]
], 3);
ta_sd: random_permutation(append(correct_ans_sd, wrong_ans_sd));

/* Short summary of this variant for the question note */
variantnote: sconcat("correct positions a/b/c = ",
    correctpos(ta_da), "/", correctpos(ta_pd), "/", correctpos(ta_sd));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@variantnote@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">Sources of data: matching the terms Data, Primary Data and Secondary Data to their definitions.</p>
<p id="reviewer">Ben (STACK audit)</p>
<p id="description">Students match three statistical terms to their correct definitions using dropdown menus. Each dropdown shows the correct definition plus three distractors sampled at random from a pool of five. Terms restricted to those taught in the lesson (Data, Primary Data, Secondary Data); each wrong answer gives a diagnostic hint.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_da</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_pd</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>ta_sd</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_da)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. <em>Data</em> is the most general term: the raw facts, figures, or observations you record &mdash; not a chart, a conclusion, or a calculation done with them.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_pd)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. <em>Primary data</em> is first-hand: collected directly by the researcher for this specific study. If someone else gathered it earlier, it is not primary.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>mcq_correct(ta_sd)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. <em>Secondary data</em> was collected by someone else, for a different purpose, and is being re-used. If you gathered it yourself, it would be primary.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1310784277</deployedseed>
    <deployedseed>445120</deployedseed>
    <deployedseed>739218</deployedseed>
    <deployedseed>982004</deployedseed>
    <deployedseed>167355</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All three parts correct: full marks on every PRT.</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta_da)[1]</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>mcq_correct(ta_pd)[1]</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>mcq_correct(ta_sd)[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>All three wrong (a distractor selected): fires each F branch.</description>
      <testinput>
        <name>ans1</name>
        <value>wrongval(ta_da)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>wrongval(ta_pd)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>wrongval(ta_sd)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
