<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Expressing an exponential number in logarithmic form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A computing device stores data in powers of a specific base. In one task, it processes a total of \({@number@}\) units using base \({@base@}\). </p>
<p>Write the logarithmic form of the statement:</p>

<p>\[
{@base@}^{@exp@} = {@number@}
\]</p>

<p>Express your answer as: \(\log_{a}(b)\).</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>The exponential form is:</p>

<p>\[
{@base@}^{@exp@} = {@number@}.
\]</p>

<p>To convert from exponential to logarithmic form, we use:</p>

<p>\[
a^b = c \quad \longrightarrow \quad \log_a(c) = b.
\]</p>

<p>Therefore:</p>

<p>\[
\log_{ {@base@} }({@number@}) = {@exp@}.
\]</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Random base for logarithm */
base: rand([2,3,5,10]);

/* Random exponent */
exp: rand([2,3,4,5]);

/* Resulting number */
number: base^exp;

/* Correct logarithmic expression in STACK form */
log_expr: log(number)/log(base);

/* Common mistakes */
m1: log(exp)/log(base);        /* wrong input */
m2: log(base)/log(number);     /* reversed */
m3: exp^base;                  /* nonsense algebraic error */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@base@}^{@exp@} = {@number@}</p>
<p>Logarithmic: log_base(number) = exponent</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"> The question requires the learner to convert an exponential statement into its equivalent logarithmic form by identifying the base, the result, and the corresponding exponent.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>log(number)/log(base)</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>log(number)/log(base)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2065750636</deployedseed>
    <deployedseed>916985200</deployedseed>
    <deployedseed>1370623816</deployedseed>
    <deployedseed>1665816775</deployedseed>
    <deployedseed>851641330</deployedseed>
    <deployedseed>463721273</deployedseed>
    <deployedseed>1975449274</deployedseed>
    <deployedseed>1207162177</deployedseed>
    <deployedseed>2036983640</deployedseed>
    <deployedseed>79543480</deployedseed>
    <deployedseed>2013661981</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>log(number)/log(base)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
