<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Factorising a monic quadratic with negative terms</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Factorise the following quadratic expression completely:
\[
{@expr@}
\]
</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
We want to factorise:
</p>

<p>
\[
{@expr@}
\]
</p>

<p>
Step 1: Identify the coefficients.
</p>

<p>
Coefficient of \(x\) = {@mid@}
</p>

<p>
Constant term = {@const@}
</p>

<p>
Step 2: Find two numbers whose product is {@const@} and whose sum is {@mid@}.
</p>

<p>
The numbers are {@a@} and {@b@}, since:
</p>

<p>
\[
{@a@} \times {@b@} = {@const@}
\]
</p>

<p>
\[
-{@a@} + (-{@b@}) = {@mid@}
\]
</p>

<p>
Step 3: Write the factorisation.
</p>

<p>
\[
(x - {@a@})(x - {@b@})
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random factors */
a: rand(8) + 1;
b: rand(8) + 1;

/* Quadratic expression */
expr: expand((x - a)*(x - b));

/* Correct factorisation */
correct: (x - a)*(x - b);

/* Coefficients for explanation */
mid: -(a + b);
const: a*b;

/* Common mistakes */
m1: (x + a)*(x + b);
m2: x^2 - a*x - b;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@expr@}, {@correct@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Factorising monic quadratic expressions</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to factorise a monic quadratic expression with negative terms by identifying two numbers whose product equals the constant term and whose sum equals the coefficient of the linear term.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer in factored form</description>
        <answertest>FacForm</answertest>
        <sans>ans1</sans>
        <tans>correct</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>m1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Check your positive and negative signs inside the brackets.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>m2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>Once you have found the two numbers, remember to factorise.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>Find two numbers whose product is {@const@} and whose sum is {@mid@}.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1925935328</deployedseed>
    <deployedseed>2003747367</deployedseed>
    <deployedseed>875986891</deployedseed>
    <deployedseed>1848601307</deployedseed>
    <deployedseed>474406003</deployedseed>
    <deployedseed>574864037</deployedseed>
    <deployedseed>1307068457</deployedseed>
    <deployedseed>120425781</deployedseed>
    <deployedseed>913077904</deployedseed>
    <deployedseed>813327751</deployedseed>
    <deployedseed>443479003</deployedseed>
    <deployedseed>379651381</deployedseed>
    <deployedseed>637766757</deployedseed>
    <deployedseed>835572277</deployedseed>
    <deployedseed>959179065</deployedseed>
    <deployedseed>326322674</deployedseed>
    <deployedseed>1650548530</deployedseed>
    <deployedseed>437485428</deployedseed>
    <deployedseed>475639725</deployedseed>
    <deployedseed>1580069314</deployedseed>
    <deployedseed>4845273</deployedseed>
    <deployedseed>1651744419</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>correct</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
