<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Question 2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<h3 style="text-align:center;">
Expand and simplify the following expression:
</h3>

<p style="text-align:center;">
\[
({@a@}x + {@b@})({@c@}x + {@d@})
\]
</p>

<p style="text-align:center;">
Give your answer in the form \(Ax^2 + Bx + C\).
</p>

<p style="text-align:center;">
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>To expand the product of two binomials, we use the distributive property 
(multiply each term in the first bracket by each term in the second bracket).</p>

<p>
\(
({@a@}x + {@b@})({@c@}x + {@d@})
\)
</p>

<p>Multiply each term:</p>

<p>
\(
= ({@a@}x)({@c@}x) + ({@a@}x)({@d@}) + ({@b@})({@c@}x) + ({@b@})({@d@})
\)
</p>

<p>Simplify each product:</p>

<p>
\(
= {@a*c@}x^2 + {@a*d@}x + {@b*c@}x + {@b*d@}
\)
</p>

<p>Now combine like terms:</p>

<p>
\(
= {@a*c@}x^2 + {@a*d + b*c@}x + {@b*d@}
\)
</p>

<p><strong>Final Answer:</strong></p>

<p>
\(
{@ta@}
\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random non-zero coefficients */
a : rand_with_prohib(-2,2,[0]);
b : rand_with_prohib(-3,3,[0]);
c : rand_with_prohib(-6,6,[0]);
d : rand_with_prohib(-4,4,[0]);

/* Define expression */
expr : (a*x + b)*(c*x + d);

/* Teacher answer */
ta : expand(expr);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>199681474</deployedseed>
    <deployedseed>2003328829</deployedseed>
    <deployedseed>830492636</deployedseed>
    <deployedseed>1109137395</deployedseed>
    <deployedseed>1370764639</deployedseed>
    <deployedseed>1785391495</deployedseed>
    <deployedseed>733665031</deployedseed>
    <deployedseed>680861079</deployedseed>
    <deployedseed>951859339</deployedseed>
    <deployedseed>738210911</deployedseed>
    <deployedseed>376288497</deployedseed>
    <deployedseed>1243045465</deployedseed>
    <deployedseed>1591983009</deployedseed>
    <deployedseed>1421555500</deployedseed>
    <deployedseed>1613275231</deployedseed>
    <deployedseed>1183189647</deployedseed>
    <deployedseed>638595777</deployedseed>
    <deployedseed>717384430</deployedseed>
    <deployedseed>1241262619</deployedseed>
    <deployedseed>972476929</deployedseed>
    <deployedseed>977516090</deployedseed>
    <deployedseed>1510564716</deployedseed>
    <deployedseed>1843132596</deployedseed>
    <deployedseed>33321392</deployedseed>
    <deployedseed>1123224494</deployedseed>
    <deployedseed>1895596180</deployedseed>
    <deployedseed>1086702291</deployedseed>
    <deployedseed>378251228</deployedseed>
    <deployedseed>1814111948</deployedseed>
    <deployedseed>369682981</deployedseed>
    <deployedseed>1260798296</deployedseed>
    <deployedseed>1768042527</deployedseed>
    <deployedseed>615306389</deployedseed>
    <deployedseed>878023180</deployedseed>
    <deployedseed>1122606638</deployedseed>
    <deployedseed>2065312115</deployedseed>
    <deployedseed>1878648934</deployedseed>
    <deployedseed>933384639</deployedseed>
    <deployedseed>1279077315</deployedseed>
    <deployedseed>1401129383</deployedseed>
    <deployedseed>199973521</deployedseed>
    <deployedseed>915259598</deployedseed>
    <deployedseed>1628425535</deployedseed>
    <deployedseed>11266180</deployedseed>
    <deployedseed>927835485</deployedseed>
    <deployedseed>472536277</deployedseed>
    <deployedseed>2020103643</deployedseed>
    <deployedseed>1128088252</deployedseed>
    <deployedseed>1125313826</deployedseed>
    <deployedseed>1083566452</deployedseed>
    <deployedseed>697038379</deployedseed>
    <deployedseed>1687815591</deployedseed>
    <deployedseed>1396526384</deployedseed>
    <deployedseed>1543929594</deployedseed>
    <deployedseed>1690916758</deployedseed>
    <deployedseed>821535196</deployedseed>
    <deployedseed>836390334</deployedseed>
    <deployedseed>672153539</deployedseed>
    <deployedseed>1601510307</deployedseed>
    <deployedseed>1871179407</deployedseed>
    <deployedseed>694213150</deployedseed>
    <deployedseed>400843331</deployedseed>
    <deployedseed>1291738641</deployedseed>
    <deployedseed>673967450</deployedseed>
    <deployedseed>2112039303</deployedseed>
    <deployedseed>307391830</deployedseed>
    <deployedseed>1363272518</deployedseed>
    <deployedseed>1327647609</deployedseed>
    <deployedseed>811305842</deployedseed>
    <deployedseed>1297719486</deployedseed>
    <deployedseed>556538800</deployedseed>
    <deployedseed>1319578239</deployedseed>
    <deployedseed>243361292</deployedseed>
    <deployedseed>1072490710</deployedseed>
    <deployedseed>1024709856</deployedseed>
    <deployedseed>79641710</deployedseed>
    <deployedseed>1112853497</deployedseed>
    <deployedseed>178639213</deployedseed>
    <deployedseed>257891328</deployedseed>
    <deployedseed>2087259120</deployedseed>
    <deployedseed>757248884</deployedseed>
    <deployedseed>808348164</deployedseed>
    <deployedseed>44140270</deployedseed>
    <deployedseed>1532918280</deployedseed>
    <deployedseed>1983658409</deployedseed>
    <deployedseed>2032414431</deployedseed>
    <deployedseed>906997047</deployedseed>
    <deployedseed>546965263</deployedseed>
    <deployedseed>1909055865</deployedseed>
    <deployedseed>780532647</deployedseed>
    <deployedseed>1018040394</deployedseed>
    <deployedseed>821698512</deployedseed>
    <deployedseed>587338845</deployedseed>
    <deployedseed>119381442</deployedseed>
    <deployedseed>857547175</deployedseed>
    <deployedseed>1273489256</deployedseed>
    <deployedseed>581641319</deployedseed>
    <deployedseed>1595977719</deployedseed>
    <deployedseed>1473935572</deployedseed>
    <deployedseed>1308639633</deployedseed>
    <deployedseed>899741419</deployedseed>
    <deployedseed>622669994</deployedseed>
    <deployedseed>1306583970</deployedseed>
    <deployedseed>1511752619</deployedseed>
    <deployedseed>226007670</deployedseed>
    <deployedseed>151673649</deployedseed>
    <deployedseed>261812263</deployedseed>
    <deployedseed>1886360103</deployedseed>
    <deployedseed>1421829295</deployedseed>
    <deployedseed>220942558</deployedseed>
    <deployedseed>1578383398</deployedseed>
    <deployedseed>479793117</deployedseed>
    <deployedseed>1172141120</deployedseed>
    <deployedseed>1555105360</deployedseed>
    <deployedseed>772118275</deployedseed>
    <deployedseed>1806670874</deployedseed>
    <deployedseed>899064993</deployedseed>
    <deployedseed>981309893</deployedseed>
    <deployedseed>630903451</deployedseed>
    <deployedseed>1829059107</deployedseed>
    <deployedseed>573698001</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
