<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the midpoint</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>In the figure below, \(ABC\) is a triangle where \(M\) is the midpoint of \(AC\) and \(N\) is the midpoint of \(AB\).</p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    /* Increased bounding box to 22 to fit coordinates up to 20 comfortably */
    boundingbox: [-2, 22, 22, -2],
    axis: true,
    showCopyright: false
});

var A = board.create('point', [{#ax#}, {#ay#}], {name:'A', fixed:true, size:2});
var B = board.create('point', [{#bx#}, {#by#}], {name:'B', fixed:true, size:2});
var C = board.create('point', [{#cx#}, {#cy#}], {name:'C', fixed:true, size:2});

var poly = board.create('polygon', [A,B,C], {borders:{strokeColor:'black'}});

/* Visual Helpers */
var M = board.create('midpoint', [A, C], {name:'M', color:'red', size:3});
var N = board.create('midpoint', [A, B], {name:'N', color:'red', size:3});

board.create('segment', [B, M], {dash:2, strokeColor:'gray'});
board.create('segment', [C, N], {dash:2, strokeColor:'gray'});
[[/jsxgraph]]

<p>Determine the coordinates of \(M\) and \(N\).</p>
<p><i>(Enter your answers as coordinates, e.g., <code>[x,y]</code>)</i></p>

<p>\( M = \) [[input:ans1]] [[validation:ans1]] </p>
<p>\( N = \) [[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Method:</b></p>
<p>The midpoint \( (x_m, y_m) \) between two points \( (x_1, y_1) \) and \( (x_2, y_2) \) is calculated by taking the average of their coordinates:</p>
\[ M = \left( \frac{x_1 + x_2}{2}, \frac{y_1 + y_2}{2} \right) \]

<p><b>For point M (Midpoint of AC):</b></p>
\[ x_M = \frac{{@ax@} + {@cx@}}{2} = {@mx@} \]
\[ y_M = \frac{{@ay@} + {@cy@}}{2} = {@my@} \]
<p>So, \( M = {@ta_m@} \).</p>

<p><b>For point N (Midpoint of AB):</b></p>
\[ x_N = \frac{{@ax@} + {@bx@}}{2} = {@nx@} \]
\[ y_N = \frac{{@ay@} + {@by@}}{2} = {@ny@} \]
<p>So, \( N = {@ta_n@} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. Randomize Start Point A (Keep small to allow room for triangle) */
/* Range: 2, 4, 6, 8 */
ax: 2 * rand_with_step(1, 4, 1);
ay: 2 * rand_with_step(1, 4, 1);

/* 2. Define B (Right of A) */
/* We add between 4 and 10 units. Max x = 8 + 10 = 18. Safe < 20 */
bx: ax + 2 * rand_with_step(2, 5, 1);
by: ay; 

/* 3. Define C (Above A, shifted horizontally slightly) */
/* cx shift: 2 to 6. Max cx = 8 + 6 = 14. Safe. */
cx: ax + 2 * rand_with_step(1, 3, 1);
/* cy shift: 4 to 10. Max cy = 8 + 10 = 18. Safe < 20 */
cy: ay + 2 * rand_with_step(2, 5, 1);

/* 4. Calculate Midpoints */

/* M is midpoint of AC */
mx: (ax + cx) / 2;
my: (ay + cy) / 2;
ta_m: [mx, my];

/* N is midpoint of AB (Bottom edge) */
nx: (ax + bx) / 2;
ny: (ay + by) / 2;
ta_n: [nx, ny];

/* 5. Wrong Answers for Feedback */
wa_m_sub: [(ax - cx) / 2, (ay - cy) / 2];
wa_n_sub: [(ax - bx) / 2, (ay - by) / 2];

wa_m_sum: [ax + cx, ay + cy];
wa_n_sum: [ax + bx, ay + by];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@ta_m@}\)</p>
<p>\({@ta_n@}\)</p>
<p>\({@ax@}\) and \({@ay@}\) </p>
<p>\({@bx@}\) and \({@by@}\) </p>
<p>\({@cx@}\) and \({@cy@}\) </p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_m</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_n</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>correct_midpoints_pos : if (ans1=ta_m) and (ans2=ta_n) then true else false;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct midpoints for line AB and line AC</description>
        <answertest>AlgEquiv</answertest>
        <sans>correct_midpoints_pos</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_m</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct answer. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_n</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct answer. </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wa_n_sub</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Incorrect. It looks like you subtracted the coordinates.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Incorrect answer.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1522579780</deployedseed>
    <deployedseed>1261513782</deployedseed>
    <deployedseed>49043099</deployedseed>
    <deployedseed>1115177532</deployedseed>
    <deployedseed>519696522</deployedseed>
    <deployedseed>1264480472</deployedseed>
    <deployedseed>1758150157</deployedseed>
    <deployedseed>351041153</deployedseed>
    <deployedseed>1780218462</deployedseed>
    <deployedseed>2114374321</deployedseed>
    <deployedseed>412949843</deployedseed>
    <deployedseed>1582219536</deployedseed>
    <deployedseed>1430630457</deployedseed>
    <deployedseed>1856326426</deployedseed>
    <deployedseed>1110795934</deployedseed>
    <deployedseed>1469068536</deployedseed>
    <deployedseed>2005952897</deployedseed>
    <deployedseed>1349101578</deployedseed>
    <deployedseed>377927767</deployedseed>
    <deployedseed>1713833957</deployedseed>
    <deployedseed>1783636671</deployedseed>
    <deployedseed>1347397865</deployedseed>
    <deployedseed>1176838739</deployedseed>
    <deployedseed>1820069962</deployedseed>
    <deployedseed>999910353</deployedseed>
    <deployedseed>1590852895</deployedseed>
    <deployedseed>2067877058</deployedseed>
    <deployedseed>1721052316</deployedseed>
    <deployedseed>118323211</deployedseed>
    <deployedseed>1483002656</deployedseed>
    <deployedseed>55072114</deployedseed>
    <deployedseed>161873527</deployedseed>
    <deployedseed>562432595</deployedseed>
    <deployedseed>1603448873</deployedseed>
    <deployedseed>50702860</deployedseed>
    <deployedseed>1515856902</deployedseed>
    <deployedseed>184907084</deployedseed>
    <deployedseed>1321357576</deployedseed>
    <deployedseed>1459811238</deployedseed>
    <deployedseed>1345381427</deployedseed>
    <deployedseed>1065910977</deployedseed>
    <deployedseed>993141001</deployedseed>
    <deployedseed>976485105</deployedseed>
    <deployedseed>126586190</deployedseed>
    <deployedseed>1620921517</deployedseed>
    <deployedseed>1475992450</deployedseed>
    <deployedseed>838886118</deployedseed>
    <deployedseed>39699909</deployedseed>
    <deployedseed>2112405004</deployedseed>
    <deployedseed>1318491937</deployedseed>
  </question>
</quiz>
