<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of a sector</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Ekadeli opened a door of width \({@r@}\,\text{m}\), fixed at a hinge point \(O\). As the door swung open, it swept through an angle of \({@theta_deg@}^\circ\), tracing out the shaded region of floor shown below.</p>

<p>Calculate the area of the floor swept by the door. Use \(\pi=\dfrac{22}{7}\), and give your answer correct to \(2\) decimal places.</p>

[[jsxgraph]]
      var board = JXG.JSXGraph.initBoard(divid, {
        boundingbox: [-2, 8, 10, -2],
        axis: false,
        showCopyright: false,
        showNavigation: false,
        keepAspectRatio: true
      });

      // Centre point O (the door's hinge)
      var O = board.create('point', [0, 0], {
        name: 'O',
        fixed: true,
        size: 2,
        showInfobox: false,
        label: { offset: [-14, -12], fontSize: 16 }
      });

      var angleDeg = {#theta_deg#};
      var angleRad = angleDeg * Math.PI / 180;

      // Fixed DISPLAY radius (independent of the true randomized r)
      // This guarantees the sector always fills the board nicely,
      // regardless of whether the real r is 1.0 m or 1.9 m.
      var displayRadius = 6;

      // Point A: door's closed position (along the horizontal)
      var A = board.create('point', [displayRadius, 0], {
        name: 'A',
        fixed: true,
        size: 0,
        visible: false,
        label: { visible: false }
      });

      // Point B: door's open position, at the true swept angle
      var B = board.create('point', [
        displayRadius * Math.cos(angleRad),
        displayRadius * Math.sin(angleRad)
      ], {
        name: 'B',
        fixed: true,
        size: 0,
        visible: false,
        label: { visible: false }
      });

      // Shaded sector representing the swept floor area
      board.create('sector', [O, A, B], {
        fillColor: '#e8e8e8',
        fillOpacity: 1,
        strokeColor: 'blue',
        strokeWidth: 2,
        withLabel: false
      });

      // Small arc near O to mark the angle itself
      board.create('sector', [O, A, B], {
        radius: 1.2,
        fillColor: 'none',
        strokeColor: 'blue',
        strokeWidth: 1,
        withLabel: false
      });

      // Angle label, positioned along the bisector of the swept angle
      board.create('text', [
        function() { return 1.8 * Math.cos(angleRad / 2); },
        function() { return 1.8 * Math.sin(angleRad / 2); },
        function() { return angleDeg + '°'; }
      ], { fontSize: 14 });

      // Label for the door width (radius), placed along OB, the open edge
      board.create('text', [
        function() { return displayRadius * 0.55 * Math.cos(angleRad); },
        function() { return displayRadius * 0.55 * Math.sin(angleRad) + 0.5; },
        function() { return {#r#} + ' m'; }
      ], { fontSize: 14 });
[[/jsxgraph]]

<p>\(\text{Area through which the door swept}=\)[[input:ans1]]\(\,\text{m}^2\)[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>We are given:</p>
<ul>
  <li>Radius (door width): \(r={@r@}\,\text{m}\)</li>
  <li>Angle swept: \(\theta={@theta_deg@}^\circ\)</li>
  <li>\(\pi={@pi_val@}\)</li>
</ul>

<p>The swept floor area is a sector of a circle. The area of a sector is:</p>
<p>\(\text{Area}=\dfrac{\theta}{360}\times \pi r^2\)</p>

<p>Substituting the values:</p>
<p>
\(
\begin{align}
\text{Area} &=\dfrac{{@theta_deg@}}{360} \times {@pi_val@} \times ({@r@})^2\\
&=\dfrac{{@theta_deg@}}{360} \times {@pi_val@} \times {@r_square@} \\
&= {@area@}\,\text{m}^2  \\
&= {@Area@}\,\text{m}^2 
\end{align}
\)
</p>

<p>Rounded to \(2\) decimal places:</p>
<p>\(\text{Area through which the door swept}\approx {@Area@}\,\text{m}^2\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomize radius r in metres, range [1.0, 2.0), to 1 decimal place */
r: (rand(10) + 10) / 10;

/* Randomize angle theta in degrees, range [25, 85] */
theta_deg: rand(60) + 25;

/* Fraction of full turn swept */
theta_frac: theta_deg / 360;

/* Use pi = 22/7 as specified in the question */
pi_val: 22/7;

r_square: r^2;

/* Area of the sector: A = (theta/360) * pi * r^2 */
area: theta_frac * pi_val * r_square;
Area: float(decimalplaces(area, 2));

/* Distractor: mistake of using arc length formula instead of sector area */
area_arclength_mistake: theta_frac * 2 * pi_val * r;
Area_arclength_mistake: float(decimalplaces(area_arclength_mistake, 2));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{Radius of the sector}=\){@r@}</p>
<p>\(\text{The angle }\, \theta =\){@theta_deg @}</p>
<p>\(\text{Area}=\){@Area@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Henry Onyango</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question provided requires the student to find the area of a sector.
To solve it correctly, the student must carefully read the word problem and extract the relevant information to apply the correct formula.

The formula for the area of a sector is:  \(\text{Area}=\frac{\theta}{360} \times \pi r^2\)</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Area</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct area calculated.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Area</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for wrong use of the formular.</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Area_arclength_mistake</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> It looks like you used the formula for arc length instead of the formula for the area of the sector. You may have done something like this:</p>
<p>\(\textbf{Arc length}=\frac{\theta}{360^\circ} \times 2 \pi r\)</p>
<p>This gives the distance along the curved edge of the door's swing, not the area of the floor that the door actually swept through.</p>
<p>\(\textbf{What You Should Do Instead:}\)</p>
<p>Use the area of a sector formula:</p>
<p>\(\text{Area}= \frac{\theta}{360^\circ} \times \pi r^2\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1676065475</deployedseed>
    <deployedseed>755370642</deployedseed>
    <deployedseed>227248341</deployedseed>
    <deployedseed>118401119</deployedseed>
    <deployedseed>1677184340</deployedseed>
    <deployedseed>1231560853</deployedseed>
    <deployedseed>1451232413</deployedseed>
    <deployedseed>2006218653</deployedseed>
    <deployedseed>506884214</deployedseed>
    <deployedseed>1594147002</deployedseed>
    <deployedseed>792974397</deployedseed>
    <deployedseed>720618612</deployedseed>
    <deployedseed>1939106053</deployedseed>
    <deployedseed>1483859556</deployedseed>
    <deployedseed>82687235</deployedseed>
    <deployedseed>675935412</deployedseed>
    <deployedseed>736801933</deployedseed>
    <deployedseed>916685649</deployedseed>
    <deployedseed>1280095253</deployedseed>
    <deployedseed>55120571</deployedseed>
    <deployedseed>1600868853</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>Area</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Checks for wrong use of the formular.</description>
      <testinput>
        <name>ans1</name>
        <value>Area1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.1000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
