<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Lines of Symmetry of a Regular polygon.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Calculate the number of lines of symmetry for each regular polygon shown below.</p>
<p>(a).</p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [-2, 12, 16, -2],
    axis: true,
    pan: { enabled: false },
    zoom: { enabled: false },
    showNavigation: false
});

/* === Rotating Regular Pentagon === */

// Fixed centre and radius
var cx = 7;
var cy = 6;
var r  = 4;

// Random rotation angle (±)
var angle = (Math.random() * 360) * (Math.random() < 0.5 ? -1 : 1);

// Random fill colour
var colours = ['#1f77b4', '#d62728', '#2ca02c', '#9467bd', '#ff7f0e'];
var fillColour = colours[Math.floor(Math.random() * colours.length)];

// Degrees → radians
function deg(d){ return d * Math.PI / 180; }

// Pentagon vertices with rotation
function v(theta){
    return [
        cx + r * Math.cos(deg(theta + angle)),
        cy + r * Math.sin(deg(theta + angle))
    ];
}

// Invisible vertices
var P = board.create('point', v(90),  {visible:false, fixed:true});
var Q = board.create('point', v(18),  {visible:false, fixed:true});
var R = board.create('point', v(-54), {visible:false, fixed:true});
var S = board.create('point', v(-126),{visible:false, fixed:true});
var T = board.create('point', v(162), {visible:false, fixed:true});

// Filled polygon
var pentagon = board.create('polygon', [P,Q,R,S,T], {
    fillColor: fillColour,
    fillOpacity: 0.35,
    borders: {
        strokeWidth: 3,
        strokeColor: fillColour
    },
    fixed: true
});

// Figure label
board.create('text', [
    cx - 1.7,
    cy - r - 1.2,
    // "Regular Pentagon"
], {
    fontSize: 16,
    fontWeight: 'bold',
    fixed: true
});

/* Prevent interaction */
board.renderer.svgRoot.style.pointerEvents = "none";

[[/jsxgraph]]

<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>(b).</p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [-2, 12, 16, -2],
    axis: true,
    pan: { enabled: false },
    zoom: { enabled: false },
    showNavigation: false
});

/* === Rotating Square (No Symmetry Lines) === */

// Fixed centre and half-diagonal
var cx = 7;
var cy = 6;
var r  = 4;

// Random rotation (CW or CCW)
var angle = (Math.random() * 360) * (Math.random() < 0.5 ? -1 : 1);

// Random fill colour
var fillColours = ['#cfe2f3', '#f4cccc', '#d9ead3', '#ead1dc', '#fff2cc'];
var fillColour = fillColours[Math.floor(Math.random() * fillColours.length)];

// Degrees → radians
function deg(d){ return d * Math.PI / 180; }

// Vertex generator (90° apart)
function v(theta){
    return [
        cx + r * Math.cos(deg(theta + angle)),
        cy + r * Math.sin(deg(theta + angle))
    ];
}

// Invisible vertices
var A = board.create('point', v(45),  {visible:false, fixed:true});
var B = board.create('point', v(135), {visible:false, fixed:true});
var C = board.create('point', v(225), {visible:false, fixed:true});
var D = board.create('point', v(315), {visible:false, fixed:true});

// Square (filled)
board.create('polygon', [A,B,C,D], {
    fillColor: fillColour,
    fillOpacity: 0.5,
    borders: {
        strokeWidth: 3,
        strokeColor: '#000'
    },
    fixed: true
});

// Label
board.create('text', [
    cx - 1.9,
    cy - r - 1.3,
    // "Square"
], {
    fontSize: 16,
    fontWeight: 'bold',
    fixed: true
});

/* Prevent interaction */
board.renderer.svgRoot.style.pointerEvents = "none";

[[/jsxgraph]]
<p>[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>(a). A regular pentagon has \(5\) lines of symmetry as shown below.</p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [-2, 12, 16, -2],
    axis: true,
    pan: { enabled: false },
    zoom: { enabled: false },
    showNavigation: false
});

/* === Rotating Regular Pentagon with Coloured Symmetry Lines === */

// Fixed centre and radius
var cx = 7;
var cy = 6;
var r  = 4;

// Random rotation (CW or CCW)
var angle = (Math.random() * 360) * (Math.random() < 0.5 ? -1 : 1);

// Random fill colour (avoid clash with symmetry lines)
var fillColours = ['#cfe2f3', '#f4cccc', '#d9ead3', '#ead1dc', '#fff2cc'];
var fillColour = fillColours[Math.floor(Math.random() * fillColours.length)];

// Degrees → radians
function deg(d){ return d * Math.PI / 180; }

// Vertex generator
function v(theta){
    return [
        cx + r * Math.cos(deg(theta + angle)),
        cy + r * Math.sin(deg(theta + angle))
    ];
}

// Invisible vertices
var P = board.create('point', v(90),  {visible:false, fixed:true});
var Q = board.create('point', v(18),  {visible:false, fixed:true});
var R = board.create('point', v(-54), {visible:false, fixed:true});
var S = board.create('point', v(-126),{visible:false, fixed:true});
var T = board.create('point', v(162), {visible:false, fixed:true});

// Pentagon (filled)
board.create('polygon', [P,Q,R,S,T], {
    fillColor: fillColour,
    fillOpacity: 0.5,
    borders: {
        strokeWidth: 3,
        strokeColor: '#000'
    },
    fixed: true
});

// === Coloured symmetry lines (bounded at vertices) ===
var symmetryData = [
    {angle: 90,   colour: '#e41a1c'},
    {angle: 18,   colour: '#377eb8'},
    {angle: -54,  colour: '#4daf4a'},
    {angle: -126, colour: '#984ea3'},
    {angle: 162,  colour: '#ff7f00'}
];

symmetryData.forEach(function(s){

    var A = board.create('point', [
        cx + r*Math.cos(deg(s.angle + angle)),
        cy + r*Math.sin(deg(s.angle + angle))
    ], {visible:false, fixed:true});

    var B = board.create('point', [
        cx - r*Math.cos(deg(s.angle + angle)),
        cy - r*Math.sin(deg(s.angle + angle))
    ], {visible:false, fixed:true});

    board.create('segment', [A, B], {
        strokeColor: s.colour,
        strokeWidth: 2,
        fixed: true
    });
});

// Label
board.create('text', [
    cx - 2.8,
    cy - r - 1.3,
    // "Lines of symmetry of a regular pentagon"
], {
    fontSize: 15,
    fontWeight: 'bold',
    fixed: true
});

/* Prevent interaction */
board.renderer.svgRoot.style.pointerEvents = "none";
[[/jsxgraph]]


<p>(b). The square has \(4\) lines of symmetry.</p>
[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [-2, 12, 16, -2],
    axis: true,
    pan: { enabled: false },
    zoom: { enabled: false },
    showNavigation: false
});

/* === Rotating Square with Lines of Symmetry === */

// Fixed centre and half-diagonal
var cx = 7;
var cy = 6;
var r  = 4;   // distance from centre to vertex

// Random rotation (CW or CCW)
var angle = (Math.random() * 360) * (Math.random() < 0.5 ? -1 : 1);

// Random fill colour
var fillColours = ['#cfe2f3', '#f4cccc', '#d9ead3', '#ead1dc', '#fff2cc'];
var fillColour = fillColours[Math.floor(Math.random() * fillColours.length)];

// Degrees → radians
function deg(d){ return d * Math.PI / 180; }

// Vertex generator (90° apart)
function v(theta){
    return [
        cx + r * Math.cos(deg(theta + angle)),
        cy + r * Math.sin(deg(theta + angle))
    ];
}

// Invisible vertices
var A = board.create('point', v(45),   {visible:false, fixed:true});
var B = board.create('point', v(135),  {visible:false, fixed:true});
var C = board.create('point', v(225),  {visible:false, fixed:true});
var D = board.create('point', v(315),  {visible:false, fixed:true});

// Square (filled)
board.create('polygon', [A,B,C,D], {
    fillColor: fillColour,
    fillOpacity: 0.5,
    borders: {
        strokeWidth: 3,
        strokeColor: '#000'
    },
    fixed: true
});

// === Lines of symmetry ===
// 4 lines: 2 diagonals, 2 through midpoints
var symmetryData = [
    {angle: 45,  colour: '#e41a1c'}, // diagonal
    {angle: 135, colour: '#377eb8'}, // diagonal
    {angle: 0,   colour: '#4daf4a'}, // vertical
    {angle: 90,  colour: '#984ea3'}  // horizontal
];

symmetryData.forEach(function(s){

    var A1 = board.create('point', [
        cx + r*Math.cos(deg(s.angle + angle)),
        cy + r*Math.sin(deg(s.angle + angle))
    ], {visible:false, fixed:true});

    var B1 = board.create('point', [
        cx - r*Math.cos(deg(s.angle + angle)),
        cy - r*Math.sin(deg(s.angle + angle))
    ], {visible:false, fixed:true});

    board.create('segment', [A1, B1], {
        strokeColor: s.colour,
        strokeWidth: 2,
        fixed: true
    });
});

// Label
board.create('text', [
    cx - 2.2,
    cy - r - 1.3,
    // "Lines of symmetry of a square"
], {
    fontSize: 15,
    fontWeight: 'bold',
    fixed: true
});

/* Prevent interaction */
board.renderer.svgRoot.style.pointerEvents = "none";
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ta:random_permutation([["2", false],["5", true],["0", false],["4", false]]);

ta1:random_permutation([["2", false],["4", true],["1", false],["3", false]]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>lines of symmetry for the regular pentagon = {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>radio</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta1)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2049988493</deployedseed>
    <deployedseed>1029871755</deployedseed>
    <deployedseed>935762523</deployedseed>
    <deployedseed>369071232</deployedseed>
    <deployedseed>570600636</deployedseed>
    <deployedseed>1305793817</deployedseed>
    <deployedseed>1627827484</deployedseed>
    <deployedseed>351637518</deployedseed>
    <deployedseed>956469304</deployedseed>
    <deployedseed>271497181</deployedseed>
    <deployedseed>1237799226</deployedseed>
    <deployedseed>1477996121</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
