<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the Area of a Wheel Region</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A wheel consists of a circular outer rim and a circular inner hub, both sharing the same centre \(O\), as shown below. The outer radius of the wheel is \({@R@}\,\text{cm}\), and the radius of the inner hub is \({@r@}\,\text{cm}\).</p>

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
        boundingbox: [-7, 7, 7, -7],
        axis: false,
        showNavigation: false,
        showCopyright: false,
        keepAspectRatio: true
    });

    // Centre point, labeled so students see both radii originate here
    var O = board.create('point', [0, 0], {
        fixed: true,
        visible: true,
        name: 'O',
        size: 2,
        label: {offset: [-14, -12]}
    });

    // Outer circle (rim) drawn first, fully shaded — represents the wheel material
    var outer = board.create('circle', [O, 5], {
        strokeWidth: 2,
        strokeColor: 'black',
        fillColor: '#f5c518',
        fillOpacity: 1
    });

    // Inner circle (hub) drawn on top, filled white to "cut out" the centre
    // What remains shaded is the annular rim region only
    var inner = board.create('circle', [O, 3], {
        strokeWidth: 2,
        strokeColor: 'black',
        fillColor: 'white',
        fillOpacity: 1
    });

    // Points marking the two radii
    var P_inner = board.create('point', [3, 0], {size: 2, fixed: true, visible: true, name: ''});
    var P_outer = board.create('point', [-4, 3], {size: 2, fixed: true, visible: true, name: ''});

    // Dashed radius segments, drawn last so they sit above the fills
    board.create('segment', [O, P_inner], {dash: 2, strokeColor: 'black', strokeWidth: 2});
    board.create('segment', [O, P_outer], {dash: 2, strokeColor: 'black', strokeWidth: 2});

    // Radius labels
    board.create('text', [1.2, -0.9, '{#r#} cm']);
    board.create('text', [-2.6, 0.9, '{#R#} cm']);
[[/jsxgraph]]

<p>Using \(\pi = 3.142\), determine the area of the wheel region (the shaded rim, excluding the hub). Give your answer correct to \(3\,\text{decimal places}\).</p>

<p>\(\text{Area of the wheel region}=\)[[input:ans1]]\(\,\text{cm}^2\)[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The outer radius (rim) is \(R={@R@}\,\text{cm}\).<br>
The inner radius (hub) is \(r={@r@}\,\text{cm}\).</p>

<p>The area of the wheel region is the area of the outer circle minus the area of the inner circle:</p>

<p>
\[
\begin{align}
A &= \pi(R^2-r^2)\\
&=3.142\left({@R@}^2-{@r@}^2\right) \\
&=3.142\left({@R_sq@}-{@r_sq@}\right) \\
&=3.142 \times {@diff_sq@} \\
&={@A@}\,\text{cm}^2
\end{align}
\]
</p>

<p>Therefore, the area of the wheel region \(\approx {@significantfigures(A,3)@}\,\text{cm}^2\) (to 3 s.f.).</p>]]></text>
    </generalfeedback>
    <defaultgrade>15</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomize inner radius (hub) */
r: rand_with_step(7, 70, 7);

/* Randomize outer radius (rim), guaranteed greater than r */
R: r + rand_with_step(7, 140, 7);

/* Use pi = 3.142 as specified */
Pi_val: 3.142;

/* Intermediate working values, named to match what they represent */
R_sq: R^2;
r_sq: r^2;
diff_sq: R_sq - r_sq;

/* Final area, to be matched to 3 significant figures in the PRT */
A: Pi_val * diff_sq;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The radius of the outer hub}={@R@}\,\text{cm}\)</p>
<p>\(\text{The radius of the inner hub}={@r@}\,\text{cm}\)</p>
<p>\(\text{The area of the wheel region}={@A@}\,\text{cm}^2\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct area of the wheel region.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>A</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>381937522</deployedseed>
    <deployedseed>1214204477</deployedseed>
    <deployedseed>1626521967</deployedseed>
    <deployedseed>1302563729</deployedseed>
    <deployedseed>1039865122</deployedseed>
    <deployedseed>1562986115</deployedseed>
    <deployedseed>2054144308</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
