<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a triangle using Heron’s Formula</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
The sides of a triangle measure <b>a</b> = {@a@} cm, <b>b</b> = {@b@} cm, and <b>c</b> = {@c@} cm. Find the area of the triangle using Heron’s formula.
</p>

[[jsxgraph]]
    var a = {#a#};
    var b = {#b#};
    var c = {#c#};
    
    // Calculate the coordinates of the third vertex using the Law of Cosines
    var x = (b*b + c*c - a*a) / (2*c);
    var y = Math.sqrt(b*b - x*x);
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-1.5, y + 2, c + 1.5, -2],
        keepAspectRatio: true,
        showNavigation: false,
        showCopyright: false
    });
    
    // Vertices
    var pA = board.create('point', [0, 0], {name: 'A', size: 3, color: 'black', label: {offset: [-15, -15]}});
    var pB = board.create('point', [c, 0], {name: 'B', size: 3, color: 'black', label: {offset: [5, -15]}});
    var pC = board.create('point', [x, y], {name: 'C', size: 3, color: 'black', label: {offset: [0, 15]}});
    
    // Triangle
    var poly = board.create('polygon', [pA, pB, pC], {
        fillColor: '#fff8e1', fillOpacity: 0.8, 
        borders: {strokeColor: '#ff8f00', strokeWidth: 2.5}
    });
    
    // Clean dynamic text labels
    board.create('text', [c/2, -0.5, "c = " + c + " cm"], {anchorX: 'middle', anchorY: 'top', fontSize: 14});
    board.create('text', [x/2 - 0.5, y/2, "b = " + b + " cm"], {anchorX: 'right', anchorY: 'middle', fontSize: 14});
    board.create('text', [(x+c)/2 + 0.5, y/2, "a = " + a + " cm"], {anchorX: 'left', anchorY: 'middle', fontSize: 14});
[[/jsxgraph]]

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{cm}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
The sides of the triangle are \(a = {@a@}\) cm, \(b = {@b@}\) cm and \(c = {@c@}\) cm.
</p>

<p>
<strong>Step 1: Find the semi-perimeter (\(s\)).</strong>
</p>

<p>
\[
s = \frac{a+b+c}{2}
\]
</p>

<p>
\[
s = \frac{{@a@}+{@b@}+{@c@}}{2}
\]
</p>

<p>
\[
s = {@s@}
\]
</p>

<p>
<strong>Step 2: Write Heron’s formula.</strong>
</p>

<p>
\[
A=\sqrt{s(s-a)(s-b)(s-c)}
\]
</p>

<p>
<strong>Step 3: Substitute the values.</strong>
</p>

<p>
\[
A=\sqrt{{@s@}({@s@}-{@a@})({@s@}-{@b@})({@s@}-{@c@})}
\]
</p>

<p>
<strong>Step 4: Simplify the terms in the brackets.</strong>
</p>

<p>
\[
A=\sqrt{{@s@} \times {@s-a@} \times {@s-b@} \times {@s-c@}}
\]
</p>

<p>
<strong>Step 5: Evaluate the square root.</strong>
</p>

<p>
\[
A \approx {@A_disp@}\text{ cm}^2
\]
</p>

<p>
Therefore, the area of the triangle is:
</p>

<p>
\[
\boxed{{@A_disp@} \text{ cm}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Sides of triangle */
a: rand(6) + 4;
b: rand(6) + 5;
c: rand(6) + 6;

/* Ensure strict triangle inequality */
while a + b <= c or a + c <= b or b + c <= a do (
  a: rand(6) + 4,
  b: rand(6) + 5,
  c: rand(6) + 6
);

/* Semi-perimeter (kept exact for standard algebraic display in feedback) */
s: (a + b + c)/2;

/* Area calculated natively as float */
A_float: float(sqrt(s*(s-a)*(s-b)*(s-c)));

/* Display variable formatted to exactly 2 decimal places */
A_disp: dispdp(A_float, 2);

/* Mistakes evaluated as floats for PRT grading */
m1: float(s*(s-a)*(s-b)*(s-c)); /* Missing square root */
m2: float((a+b+c)/3);           /* Calculated average instead of semi-perimeter */
m3: float(sqrt(a*b*c));         /* Wrong formula */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>a: {@a@}, b: {@b@}, c: {@c@}, Area: {@A_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of triangles using Heron’s formula</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to find the area of a triangle from its three sides by first determining the semi-perimeter and then applying Heron’s formula.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A_disp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>A_float</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Missing square root</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the product inside the formula perfectly, but remember to take the <strong>square root</strong> at the final step.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>wrong semi-perimeter</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The semi-perimeter (\(s\)) is exactly <strong>half</strong> the perimeter, not the average of the three sides. Ensure you are dividing by 2, not 3.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>292959087</deployedseed>
    <deployedseed>1604042663</deployedseed>
    <deployedseed>959978202</deployedseed>
    <deployedseed>1738546462</deployedseed>
    <deployedseed>146926383</deployedseed>
    <deployedseed>733333160</deployedseed>
    <deployedseed>1784636726</deployedseed>
    <deployedseed>1579235447</deployedseed>
    <deployedseed>559872600</deployedseed>
    <deployedseed>2037991164</deployedseed>
    <deployedseed>1430680303</deployedseed>
    <deployedseed>745386406</deployedseed>
    <deployedseed>2100473957</deployedseed>
    <deployedseed>1640370723</deployedseed>
    <deployedseed>920767870</deployedseed>
    <deployedseed>911718319</deployedseed>
    <deployedseed>787836760</deployedseed>
    <deployedseed>2029875748</deployedseed>
    <deployedseed>1343094378</deployedseed>
    <deployedseed>176458070</deployedseed>
    <deployedseed>612533071</deployedseed>
    <deployedseed>930129059</deployedseed>
    <deployedseed>1094328764</deployedseed>
    <deployedseed>1565566461</deployedseed>
    <deployedseed>1234704419</deployedseed>
    <deployedseed>1230374712</deployedseed>
    <deployedseed>527715615</deployedseed>
    <deployedseed>914586540</deployedseed>
    <deployedseed>1517796028</deployedseed>
    <deployedseed>665868205</deployedseed>
    <deployedseed>1993737909</deployedseed>
    <deployedseed>587959705</deployedseed>
    <deployedseed>1603712172</deployedseed>
    <deployedseed>940796283</deployedseed>
    <deployedseed>1223463832</deployedseed>
    <deployedseed>1227406627</deployedseed>
    <deployedseed>596506697</deployedseed>
    <deployedseed>32453885</deployedseed>
    <deployedseed>380128885</deployedseed>
    <deployedseed>1757568556</deployedseed>
    <deployedseed>1120143646</deployedseed>
    <deployedseed>1651294029</deployedseed>
    <deployedseed>1728704396</deployedseed>
    <deployedseed>1962912674</deployedseed>
    <deployedseed>412774387</deployedseed>
    <deployedseed>1573551247</deployedseed>
    <deployedseed>646248789</deployedseed>
    <deployedseed>919775676</deployedseed>
    <deployedseed>1161934182</deployedseed>
    <deployedseed>317388820</deployedseed>
    <deployedseed>1514303209</deployedseed>
    <deployedseed>1276544975</deployedseed>
    <deployedseed>2140940416</deployedseed>
    <deployedseed>953980589</deployedseed>
    <deployedseed>790891243</deployedseed>
    <deployedseed>501842175</deployedseed>
    <deployedseed>519250282</deployedseed>
    <deployedseed>1194343085</deployedseed>
    <deployedseed>1576168117</deployedseed>
    <deployedseed>1684098404</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
