<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a triangular land using Heron’s Formula</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A triangular piece of land has sides measuring <b>{@a@} m</b>, <b>{@b@} m</b>, and <b>{@c@} m</b>. Find the area of the land using Heron’s formula.
</p>

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{m}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
The sides of the land are \(a = {@a@}\) m, \(b = {@b@}\) m and \(c = {@c@}\) m.
</p>

<p>
<strong>Step 1: Find the semi-perimeter (\(s\)).</strong>
</p>

<p>
\[
s = \frac{a+b+c}{2}
\]
</p>

<p>
\[
s = \frac{{@a@}+{@b@}+{@c@}}{2}
\]
</p>

<p>
\[
s = {@s@}
\]
</p>

<p>
<strong>Step 2: Write Heron’s formula.</strong>
</p>

<p>
\[
A=\sqrt{s(s-a)(s-b)(s-c)}
\]
</p>

<p>
<strong>Step 3: Substitute the values.</strong>
</p>

<p>
\[
A=\sqrt{{@s@}({@s@}-{@a@})({@s@}-{@b@})({@s@}-{@c@})}
\]
</p>

<p>
<strong>Step 4: Simplify the brackets.</strong>
</p>

<p>
\[
A=\sqrt{{@s@}\times {@s-a@}\times {@s-b@}\times {@s-c@}}
\]
</p>

<p>
<strong>Step 5: Evaluate the square root.</strong>
</p>

<p>
\[
A \approx {@A_disp@}\text{ m}^2
\]
</p>

<p>
Therefore, the area of the land is:
</p>

<p>
\[
\boxed{{@A_disp@}\text{ m}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Field sides */
a: rand(10) + 15;
b: rand(10) + 18;
c: rand(10) + 20;

/* Ensure strict triangle inequality */
while a + b <= c or a + c <= b or b + c <= a do (
  a: rand(10) + 15,
  b: rand(10) + 18,
  c: rand(10) + 20
);

/* Semi-perimeter (kept exact for standard algebraic display in feedback) */
s: (a + b + c)/2;

/* Area calculated natively as float */
A_float: float(sqrt(s*(s-a)*(s-b)*(s-c)));

/* Display variable formatted to exactly 2 decimal places */
A_disp: dispdp(A_float, 2);

/* Mistakes evaluated as floats for PRT grading */
m1: float(a*b*c);
m2: float(s*(s-a)*(s-b)*(s-c)); /* Forgot square root */
m3: float((a+b+c)/2);           /* Just semi-perimeter */
m4: float(sqrt(a*b));           /* Wrong formula */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>a: {@a@}, b: {@b@}, c: {@c@}, Area: {@A_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of triangles using Heron’s formula</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to determine the area of a triangular piece of land from its three sides by applying Heron’s formula in a real-life context.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A_disp</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>A_float</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Forgot square root</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the product inside the formula correctly, but you forgot to take the <strong>square root</strong> at the end.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>used wrong formula</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You simply multiplied the three sides together. You need to find the semi-perimeter first, then apply Heron's formula.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Just semi-perimeter</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You only calculated the semi-perimeter (\(s\)). You need to substitute this value into Heron's formula to find the area.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1185440436</deployedseed>
    <deployedseed>74470588</deployedseed>
    <deployedseed>889474159</deployedseed>
    <deployedseed>1777018425</deployedseed>
    <deployedseed>1093341235</deployedseed>
    <deployedseed>903526</deployedseed>
    <deployedseed>1679766717</deployedseed>
    <deployedseed>191027217</deployedseed>
    <deployedseed>60570797</deployedseed>
    <deployedseed>1922502483</deployedseed>
    <deployedseed>259958585</deployedseed>
    <deployedseed>2034443852</deployedseed>
    <deployedseed>1850751332</deployedseed>
    <deployedseed>1105599091</deployedseed>
    <deployedseed>362449695</deployedseed>
    <deployedseed>1559461890</deployedseed>
    <deployedseed>978749852</deployedseed>
    <deployedseed>1119705224</deployedseed>
    <deployedseed>66766933</deployedseed>
    <deployedseed>1907901877</deployedseed>
    <deployedseed>1645779676</deployedseed>
    <deployedseed>1828049246</deployedseed>
    <deployedseed>1944048893</deployedseed>
    <deployedseed>99154271</deployedseed>
    <deployedseed>26867384</deployedseed>
    <deployedseed>674372148</deployedseed>
    <deployedseed>1420648257</deployedseed>
    <deployedseed>772389492</deployedseed>
    <deployedseed>413356585</deployedseed>
    <deployedseed>115004215</deployedseed>
    <deployedseed>1674147882</deployedseed>
    <deployedseed>808472132</deployedseed>
    <deployedseed>1478244989</deployedseed>
    <deployedseed>620564224</deployedseed>
    <deployedseed>1143039004</deployedseed>
    <deployedseed>1227159319</deployedseed>
    <deployedseed>1192024405</deployedseed>
    <deployedseed>1351218245</deployedseed>
    <deployedseed>1999174685</deployedseed>
    <deployedseed>1455262796</deployedseed>
    <deployedseed>451708699</deployedseed>
    <deployedseed>1990984926</deployedseed>
    <deployedseed>1930335012</deployedseed>
    <deployedseed>1427076833</deployedseed>
    <deployedseed>520838402</deployedseed>
    <deployedseed>1278947265</deployedseed>
    <deployedseed>1495007367</deployedseed>
    <deployedseed>758274385</deployedseed>
    <deployedseed>1190126183</deployedseed>
    <deployedseed>525526124</deployedseed>
    <deployedseed>808586762</deployedseed>
    <deployedseed>1500248754</deployedseed>
    <deployedseed>1943775170</deployedseed>
    <deployedseed>662266261</deployedseed>
    <deployedseed>481390107</deployedseed>
    <deployedseed>776432202</deployedseed>
    <deployedseed>2064188857</deployedseed>
    <deployedseed>705864508</deployedseed>
    <deployedseed>2058119152</deployedseed>
    <deployedseed>790227306</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
