<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of an irregular piece of land</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A surveyor recorded the following measurements for an irregular piece of land <b>AEDF</b> using a baseline <b>AD</b>. Vertices <b>E</b> and <b>F</b> are on opposite sides of the baseline.</p>

<ul>
    <li>Baseline segments: AB = <b>{@d1@} m</b>, BC = <b>{@d2@} m</b>, and CD = <b>{@d3@} m</b>.</li>
    <li>Left Offset: BE = <b>{@hL@} m</b> (perpendicular to AD at B).</li>
    <li>Right Offset: CF = <b>{@hR@} m</b> (perpendicular to AD at C).</li>
</ul>

[[jsxgraph width="400px" height="500px"]]
    var d1 = {#d1#};
    var d2 = {#d2#};
    var d3 = {#d3#};
    var hL = {#hL#};
    var hR = {#hR#};
    
    var AD = d1 + d2 + d3;

    // Dynamic Bounding Box ensuring a perfect fit
    var totalW = hL + hR;
    var spanY = AD * 1.25; 
    var ratio = 400 / 500;
    var spanX = spanY * ratio;
    
    if (totalW * 1.25 > spanX) {
        spanX = totalW * 1.25;
        spanY = spanX / ratio;
    }
    
    var midX = (hR - hL) / 2;
    var midY = AD / 2;

    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [midX - spanX/2, midY + spanY/2, midX + spanX/2, midY - spanY/2], 
        axis: false, 
        showCopyright: false,
        keepAspectRatio: true,
        showNavigation: false
    });

    // Vertical Baseline AD
    var A = board.create('point', [0, 0], {name: 'A', fixed: true, color: 'black'});
    var B = board.create('point', [0, d1], {name: 'B', fixed: true, color: 'black'});
    var C = board.create('point', [0, d1+d2], {name: 'C', fixed: true, color: 'black'});
    var D = board.create('point', [0, AD], {name: 'D', fixed: true, color: 'black'});

    // Offsets
    var E = board.create('point', [-hL, d1], {name: 'E', color: '#1565c0'});
    var F = board.create('point', [hR, d1+d2], {name: 'F', color: '#d32f2f'});

    // Polygon AEDF
    board.create('polygon', [A, E, D, F], {fillColor: '#fdf2e9', borders: {strokeWidth: 2, strokeColor: '#212121'}});

    // Construction Lines
    board.create('segment', [A, D], {dash: 1, strokeColor: '#757575'});
    board.create('segment', [B, E], {dash: 2, strokeColor: '#1976d2'});
    board.create('segment', [C, F], {dash: 2, strokeColor: '#c62828'});

    // Right angle markers
    var rad = spanY * 0.04;
    board.create('angle', [E, B, A], {type: 'square', radius: rad, strokeColor: '#1976d2', name: ''});
    board.create('angle', [F, C, D], {type: 'square', radius: rad, strokeColor: '#c62828', name: ''});

    // Dynamic Labels with Pixel Gaps
    var gapX = spanX * 0.02;
    var gapY = spanY * 0.02;

    // Baseline segment labels
    board.create('text', [gapX, d1/2, d1 + " m"], {anchorX: 'left', anchorY: 'middle', fontSize: 13, strokeColor: '#757575'});
    board.create('text', [gapX, d1 + d2/2, d2 + " m"], {anchorX: 'left', anchorY: 'middle', fontSize: 13, strokeColor: '#757575'});
    board.create('text', [gapX, d1 + d2 + d3/2, d3 + " m"], {anchorX: 'left', anchorY: 'middle', fontSize: 13, strokeColor: '#757575'});

    // Offset labels
    board.create('text', [-hL/2, d1 + gapY, hL + " m"], {anchorX: 'middle', anchorY: 'bottom', fontSize: 14, strokeColor: '#1976d2'});
    board.create('text', [hR/2, d1+d2 + gapY, hR + " m"], {anchorX: 'middle', anchorY: 'bottom', fontSize: 14, strokeColor: '#c62828'});
[[/jsxgraph]]

<p>Calculate the total area of the land <b>AEDF</b>.</p>

<p>Area = [[input:ans1]] <b>\(\text{m}^2\)</b> [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>To calculate the total area, we find the area of the large triangles on the left and right sides of the baseline <b>AD</b> separately.</p>

<p>First, find the total length of the baseline <b>AD</b>:</p>
<p>
\[ \begin{aligned} 
\text{Base AD} &= {@d1@} + {@d2@} + {@d3@} \\
&= {@AD@} \text{ m} 
\end{aligned} \]
</p>

<p><b>1. Area of Left Side (Triangle AED)</b></p>
<p>The vertex <b>E</b> forms one large triangle with base <b>AD</b>.</p>
<p>
\[ \begin{aligned} 
\text{Area (Left)} &= \frac{1}{2} \times \text{Base (AD)} \times \text{Height (BE)} \\
&= \frac{1}{2} \times {@AD@} \times {@hL@} \\
&= {@area_left@} \text{ m}^2 
\end{aligned} \]
</p>

<p><b>2. Area of Right Side (Triangle AFD)</b></p>
<p>The vertex <b>F</b> forms one large triangle with base <b>AD</b>.</p>
<p>
\[ \begin{aligned} 
\text{Area (Right)} &= \frac{1}{2} \times \text{Base (AD)} \times \text{Height (CF)} \\
&= \frac{1}{2} \times {@AD@} \times {@hR@} \\
&= {@area_right@} \text{ m}^2 
\end{aligned} \]
</p>

<p><b>3. Total Area</b></p>
<p>
\[ \begin{aligned} 
\text{Total Area} &= \text{Area (Left)} + \text{Area (Right)} \\
&= {@area_left@} + {@area_right@} \\
&= {@ta@} \text{ m}^2 
\end{aligned} \]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Baseline segments along AD */
d1: 5 + rand(3); 
d2: 6 + rand(4); 
d3: 4 + rand(3); 
AD: d1 + d2 + d3;

/* Offsets */
hL: 4 + rand(3); 
hR: 5 + rand(3); 

/* Calculations */
area_left: float((1/2) * AD * hL);
area_right: float((1/2) * AD * hR);
ta: float(area_left + area_right);

/* Mistakes for PRT */
m1: float((AD * hL) + (AD * hR)); /* Forgot the 1/2 factor */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>AD: {@AD@} | hL: {@hL@} | hR: {@hR@} | Total Area: {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Float-safe grading</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Caught missing 1/2 factor</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You forgot the \(\frac{1}{2}\) factor. The formula for the area of a triangle is \(\frac{1}{2} \times \text{base} \times \text{height}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>212685079</deployedseed>
    <deployedseed>888130700</deployedseed>
    <deployedseed>1772662014</deployedseed>
    <deployedseed>313896979</deployedseed>
    <deployedseed>159640350</deployedseed>
    <deployedseed>654206134</deployedseed>
    <deployedseed>875176505</deployedseed>
    <deployedseed>1333271300</deployedseed>
    <deployedseed>496437411</deployedseed>
    <deployedseed>1655450039</deployedseed>
    <deployedseed>151855583</deployedseed>
    <deployedseed>1947754042</deployedseed>
    <deployedseed>970381506</deployedseed>
    <deployedseed>1723025328</deployedseed>
    <deployedseed>165852735</deployedseed>
    <deployedseed>2100351525</deployedseed>
    <deployedseed>536973615</deployedseed>
    <deployedseed>1005415546</deployedseed>
    <deployedseed>1165019029</deployedseed>
    <deployedseed>2072197046</deployedseed>
    <deployedseed>1981787740</deployedseed>
    <deployedseed>795163800</deployedseed>
    <deployedseed>1732698915</deployedseed>
    <deployedseed>1464887790</deployedseed>
    <deployedseed>1190833090</deployedseed>
    <deployedseed>978472916</deployedseed>
    <deployedseed>1134879019</deployedseed>
    <deployedseed>118983713</deployedseed>
    <deployedseed>2138943827</deployedseed>
    <deployedseed>1546661046</deployedseed>
    <deployedseed>2087002015</deployedseed>
  </question>
</quiz>
