<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Column vectors</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given the column vectors \(\mathbf{a} = {@vec_a@}\) and \(\mathbf{b} = {@vec_b@}\), find the following:</p>
<ol class="HELM_parts">
    <li>
        <div>\( \mathbf{a} + \mathbf{b} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
    </li>
    <li>
        <div>\( {@k1@}\mathbf{a} - {@k2@}\mathbf{b} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p><b>Part (a)</b></p>
<p>To add column vectors, we add the corresponding components (top with top, bottom with bottom):</p>
\[ 
\mathbf{a} + \mathbf{b} = {@vec_a@} + {@vec_b@} 
= \begin{pmatrix} {@ax@} + ({@bx@}) \\ {@ay@} + ({@by@}) \end{pmatrix} 
= {@ta1@} 
\]

<p><b>Part (b)</b></p>
<p>First, we multiply each vector by its scalar:</p>
\[ {@k1@}\mathbf{a} = {@k1@}{@vec_a@} = \begin{pmatrix} {@k1@} \times {@ax@} \\ {@k1@} \times {@ay@} \end{pmatrix} = {@step1_a@} \]
\[ {@k2@}\mathbf{b} = {@k2@}{@vec_b@} = \begin{pmatrix} {@k2@} \times {@bx@} \\ {@k2@} \times {@by@} \end{pmatrix} = {@step1_b@} \]

<p>Then we subtract the results:</p>
\[ 
{@k1@}\mathbf{a} - {@k2@}\mathbf{b} = {@step1_a@} - {@step1_b@} 
= \begin{pmatrix} {@k1*ax@} - ({@k2*bx@}) \\ {@k1*ay@} - ({@k2*by@}) \end{pmatrix} 
= {@ta2@} 
\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Generate Random Components for Vector a */
ax: rand_with_step(-5, 5, 1);
ay: rand_with_step(-5, 5, 1);
vec_a: matrix([ax],[ay]);

/* 2. Generate Random Components for Vector b */
bx: rand_with_step(-5, 5, 1);
by: rand_with_step(-5, 5, 1);
vec_b: matrix([bx],[by]);

/* 3. Generate Random Scalars for Part (b) */
/* We want something like 2a - 3b, so we pick positive integers usually */
k1: rand_with_step(2, 5, 1);
k2: rand_with_step(2, 5, 1);

/* 4. Calculate Answers */

/* Part (a): Addition */
ta1: vec_a + vec_b;

/* Part (b): Linear Combination (k1*a - k2*b) */
ta2: k1*vec_a - k2*vec_b;

/* 5. Intermediate steps for the Worked Solution */
/* We define these to display the multiplication steps clearly in feedback */
step1_a: k1 * vec_a;
step1_b: k2 * vec_b;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\underline{p} = {@vec_a@}\) and \(\underline{q} = {@vec_a@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">Given two random 2D column vectors a and b, the student calculates their sum (a+b) and a linear combination (e.g., 2a - 3b).</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>ta2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001455870</deployedseed>
    <deployedseed>595694920</deployedseed>
    <deployedseed>1199309391</deployedseed>
    <deployedseed>845817541</deployedseed>
    <deployedseed>1034110010</deployedseed>
    <deployedseed>1606427816</deployedseed>
    <deployedseed>718580414</deployedseed>
    <deployedseed>990247914</deployedseed>
    <deployedseed>57699504</deployedseed>
    <deployedseed>1439932943</deployedseed>
    <deployedseed>1978950028</deployedseed>
    <deployedseed>150937084</deployedseed>
    <deployedseed>1815032287</deployedseed>
    <deployedseed>894805138</deployedseed>
    <deployedseed>1664814218</deployedseed>
    <deployedseed>1580163834</deployedseed>
    <deployedseed>408219878</deployedseed>
    <deployedseed>416989048</deployedseed>
    <deployedseed>1924980549</deployedseed>
    <deployedseed>33210280</deployedseed>
    <deployedseed>1215382149</deployedseed>
    <deployedseed>1345541290</deployedseed>
    <deployedseed>220270431</deployedseed>
    <deployedseed>62798957</deployedseed>
    <deployedseed>788059648</deployedseed>
    <deployedseed>981853072</deployedseed>
    <deployedseed>1639624772</deployedseed>
    <deployedseed>1004402025</deployedseed>
    <deployedseed>1811121382</deployedseed>
    <deployedseed>875784885</deployedseed>
    <deployedseed>1377811199</deployedseed>
    <deployedseed>1827662092</deployedseed>
    <deployedseed>367472771</deployedseed>
    <deployedseed>1168088745</deployedseed>
    <deployedseed>741605679</deployedseed>
    <deployedseed>752477490</deployedseed>
    <deployedseed>691986452</deployedseed>
    <deployedseed>786107690</deployedseed>
    <deployedseed>1728462728</deployedseed>
    <deployedseed>1821199061</deployedseed>
  </question>
</quiz>
