<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Using logarithms with standard form</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Write the number \({@value@}\) in standard logarithmic form.</p>

<div>Characteristic \(B = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<div>Mantissa \(\log_{10}(A) = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>
<div>Therefore \(\log_{10}({@value@}) = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>
<p>Give the value of \(\log_{10}({@value@})\) correct to \(4\) decimal places.</p>

<p><b>Hint:</b> (Express the number in standard form \(A \times 10^{B}\). )</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>First write the number in standard form:</p>
\[ {@value@} = {@coefdisp@} \times 10^{@exp@}. \]
<p>The logarithm of a product is the sum of the logarithms, and \(\log_{10}\left(10^{@exp@}\right) = {@exp@}\), so</p>
\[ \log_{10}({@value@}) = \log_{10}\left({@coefdisp@} \times 10^{@exp@}\right) = \log_{10}({@coefdisp@}) + {@exp@}. \]
<p>The power \({@exp@}\) is the <em>characteristic</em>. Because \(1 \le {@coefdisp@} < 10\), the value \(\log_{10}({@coefdisp@})\) lies between \(0\) and \(1\): this is the <em>mantissa</em>, read from logarithm tables. From tables,</p>
\[ \log_{10}({@coefdisp@}) \approx {@mant4@}, \]
<p>and adding the characteristic and the mantissa gives</p>
\[ \log_{10}({@value@}) \approx {@exp@} + {@mant4@} = {@ta4@}. \]]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Design backwards: build N = coef * 10^exp, then everything follows.   */
exp : rand([3,4,5,6]);

/* Coefficient 2.0 .. 9.9 in steps of 0.1: d2 = 0 gives integer          */
/* variants, d2 > 0 gives one-decimal variants (never coef = 1, which    */
/* would make the mantissa zero and the question degenerate).            */
d1 : 2 + rand(8);              /* units digit: 2 .. 9  */
d2 : rand(10);                 /* tenths digit: 0 .. 9 */
coef : d1 + d2/10;

/* The number shown to the student: always an integer since exp >= 3.   */
value : coef * 10^exp;

/* Display form of the coefficient (avoid rationals like 18/5).          */
coefdisp : if d2 = 0 then d1 else float(coef);

/* Model answers.                                                        */
mantf : float(lg(coef));                  /* mantissa, full precision    */
taf   : float(exp + lg(coef));            /* lg(N), full precision       */
mant4 : decimalplaces(mantf, 4);          /* mantissa to 4 d.p.          */
ta4   : decimalplaces(taf, 4);            /* final answer to 4 d.p.      */

/* Anticipated wrong answers.                                            */
wrong_char : exp + 1;          /* counted the digits of N instead        */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\({@value@} = {@coefdisp@} \times 10^{@exp@}\), \(\log_{10}({@value@}) \approx {@ta4@}\)</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description"> The question requires the learner to express a large number in standard form, identify the characteristic and mantissa of its base-ten logarithm, and hence give the logarithm as a single value to 4 decimal places. The final part is marked with error carried forward.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>exp</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>mant4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the characteristic correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>exp</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Diagnostic: counted the digits of N instead?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_char</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you have counted the digits of the number. The characteristic is the power \(B\) when the number is written in standard form \(A \times 10^{B}\) with \(1 \le A < 10\) — for a whole number this is one <em>less</em> than the number of digits.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The characteristic is the power \(B\) when the number is written in standard form \(A \times 10^{B}\) with \(1 \le A < 10\). Write \({@value@}\) in that form first.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the mantissa correct to within tolerance?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>mant4</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Diagnostic: entered the complete logarithm instead of just the mantissa?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>taf</tans>
        <testoptions>0.0005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is the complete logarithm. The mantissa is only the decimal part, \(\log_{10}(A)\), read from the tables — a value between \(0\) and \(1\). The whole-number part belongs in the characteristic box.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The mantissa is \(\log_{10}(A)\), where \(A\) is the coefficient of the standard form \(A \times 10^{B}\). Read it from your logarithm tables, correct to \(4\) decimal places.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the final value correct to within tolerance?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta4</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Error carried forward: does the final value correctly add the student's own characteristic and mantissa?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ans1+ans2</tans>
        <testoptions>0.0005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your final value correctly adds your characteristic and mantissa, so this step earns its mark — but one of those two parts is not right. Check the feedback on the boxes above.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The single value is found by adding the two parts: \(\log_{10}\left(A \times 10^{B}\right) = B + \log_{10}(A)\), the characteristic plus the mantissa. Give the result to \(4\) decimal places.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1010947386</deployedseed>
    <deployedseed>1831155997</deployedseed>
    <deployedseed>833733291</deployedseed>
    <deployedseed>738802563</deployedseed>
    <deployedseed>468244295</deployedseed>
    <deployedseed>2022690933</deployedseed>
    <deployedseed>1504850441</deployedseed>
    <deployedseed>1553148887</deployedseed>
    <deployedseed>734238376</deployedseed>
    <deployedseed>60733463</deployedseed>
    <deployedseed>836059070</deployedseed>
    <deployedseed>702981637</deployedseed>
    <deployedseed>301329819</deployedseed>
    <deployedseed>1505628615</deployedseed>
    <deployedseed>545307751</deployedseed>
    <deployedseed>123506974</deployedseed>
    <deployedseed>2004358962</deployedseed>
    <deployedseed>1782120530</deployedseed>
    <deployedseed>540219338</deployedseed>
    <deployedseed>1539382092</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All three boxes correct: full marks everywhere.</description>
      <testinput>
        <name>ans1</name>
        <value>exp</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>mant4</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Digit-count characteristic, correct mantissa, consistent sum: diagnostic fires on part 1 and the final box earns its error-carried-forward mark.</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_char</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>mant4</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>wrong_char + mant4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Complete logarithm entered in the mantissa box; final value inconsistent with the entered parts: mantissa diagnostic and both catch-all branches fire.</description>
      <testinput>
        <name>ans1</name>
        <value>exp + 2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta4 + 1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
