<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Forming a quadratic expression from the area of a rectangle</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A rectangular garden has the following dimensions:
</p>
<ul>
<li>Length = \( (x + {@a@}) \) metres</li>
<li>Width = \( (x + {@b@}) \) metres</li>
</ul>
<p>
(a) Write an expression for the area of the garden.
</p>
<p>
(b) Expand and simplify your answer.
</p>
<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Area of a rectangle = length × width.
</p>

<p>
\[
(x + {@a@})(x + {@b@})
\]
</p>

<p>
Step 1: Multiply \(x\) by both terms in the second bracket.
</p>

<p>
\[
x(x + {@b@}) = x^2 + {@b@}x
\]
</p>

<p>
Step 2: Multiply {@a@} by both terms in the second bracket.
</p>

<p>
\[
{@a@}(x + {@b@}) = {@a@}x + {@a*b@}
\]
</p>

<p>
Step 3: Add the results.
</p>

<p>
\[
x^2 + {@b@}x + {@a@}x + {@a*b@}
\]
</p>

<p>
Step 4: Combine like terms.
</p>

<p>
\[
= x^2 + {@a + b@}x + {@a*b@}
\]
</p>

<p>
Therefore, the area in simplest form is:
</p>

<p>
\[
x^2 + {@a + b@}x + {@a*b@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random dimensions */
a: rand(7) + 2;   /* 2 to 8 */
b: rand(7) + 3;   /* 3 to 9 */

/* Area expression */
expr: (x + a)*(x + b);

/* Expanded form */
expanded: expand(expr);

/* Common mistakes */
m1: x^2 + a*x + b;
m2: x^2 + (a + b);
m3: x^2 + a*b;
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@expr@}</p>
<p>{@expanded@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Forming a quadratic expression from the area of a rectangle</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to model a contextual problem using algebra and expand to form a quadratic expression
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>expanded</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>expanded</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>m1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Area is length × width. Multiply both expressions completely.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>m2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You are missing the x-term. Expand carefully, then combine like terms.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>m3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>The constant term is incorrect. Multiply the numbers carefully.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>Area = length × width. Expand step by step.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1741871433</deployedseed>
    <deployedseed>1286162258</deployedseed>
    <deployedseed>2011228068</deployedseed>
    <deployedseed>562081928</deployedseed>
    <deployedseed>1944435314</deployedseed>
    <deployedseed>47393767</deployedseed>
    <deployedseed>1379553831</deployedseed>
    <deployedseed>514252906</deployedseed>
    <deployedseed>556156251</deployedseed>
    <deployedseed>218459970</deployedseed>
    <deployedseed>1770029230</deployedseed>
    <deployedseed>1114660591</deployedseed>
    <deployedseed>1284776151</deployedseed>
    <deployedseed>800948180</deployedseed>
    <deployedseed>1990509621</deployedseed>
    <deployedseed>1177849404</deployedseed>
    <deployedseed>822065425</deployedseed>
    <deployedseed>2119791501</deployedseed>
    <deployedseed>1541394730</deployedseed>
    <deployedseed>966608362</deployedseed>
    <deployedseed>1212603707</deployedseed>
    <deployedseed>960968312</deployedseed>
    <deployedseed>1772052613</deployedseed>
    <deployedseed>1626717091</deployedseed>
    <deployedseed>378796556</deployedseed>
    <deployedseed>353049733</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>expanded</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
