<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of a rhombus</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The perimeter of a rhombus is <b>\({@Perimeter@}\text{ m}\)</b>. One of its diagonals is <b>\({@One_diagonal@}\text{ m}\)</b>.</p>

[[jsxgraph width="500px" height="300px"]]
    var d1 = {#One_diagonal#};
    var d2 = {#other_diagonal#};
    var p = {#Perimeter#};
    
    // Half diagonals for drawing
    var a = d1 / 2;
    var b = d2 / 2;
    
    var ratio = 500 / 300;
    var spanX = Math.max(d1 * 1.3, d2 * 1.3 * ratio);
    var spanY = spanX / ratio;
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-spanX/2, spanY/2, spanX/2, -spanY/2], 
        maxboundingbox: [-spanX*1.5, spanY*1.5, spanX*1.5, -spanY*1.5],
        axis: false, 
        showCopyright: false,
        keepAspectRatio: true,
        showNavigation: true 
    });

    var p1 = board.create('point', [a, 0], {name: '', visible: false});
    var p2 = board.create('point', [0, b], {name: '', visible: false});
    var p3 = board.create('point', [-a, 0], {name: '', visible: false});
    var p4 = board.create('point', [0, -b], {name: '', visible: false});
    var center = board.create('point', [0,0], {name: '', visible: false});

    // Rhombus
    board.create('polygon', [p1, p2, p3, p4], {
        fillColor: 'none', 
        strokeColor: '#1976d2', 
        strokeWidth: 2
    });

    // Diagonals
    board.create('segment', [p3, p1], {dash: 2, strokeColor: '#90caf9'});
    board.create('segment', [p4, p2], {dash: 2, strokeColor: '#90caf9'});

    // Right angle marker at intersection
    var rad = Math.min(a, b) * 0.15;
    board.create('angle', [p1, center, p2], {type: 'square', radius: rad, strokeColor: '#90caf9', name: ''});

    // Tightened pixel gaps to keep labels close to the lines
    var gapY = spanY * 0.015; 
    var gapX = spanX * 0.015; 

    // Dynamic Labels
    board.create('text', [a/2, -gapY, "d_1 = " + d1 + " m"], {anchorX: 'middle', anchorY: 'top', fontSize: 14});
    board.create('text', [-gapX, b/2, "d_2 = ?"], {anchorX: 'right', anchorY: 'middle', fontSize: 14, strokeColor: '#c62828'});
[[/jsxgraph]]

<ol type="a">
  <li>
    <p>Find the length of the other diagonal (\(d_2\)).</p>
    <p>\(d_2 = \) [[input:ans1]] <b>\(\text{m}\)</b> [[validation:ans1]] [[feedback:prt1]]</p>
  </li>
  <li>
    <p>Find the area of the rhombus.</p>
    <p>\(\text{Area} = \) [[input:ans2]] <b>\(\text{m}^2\)</b> [[validation:ans2]] [[feedback:prt2]]</p>
  </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p><b>Step 1: Find the length of one side</b></p>
<p>A rhombus has 4 equal sides. Therefore:</p>
<p>
\[ \begin{aligned}
\text{Perimeter} &= 4 \times \text{side} \\
{@Perimeter@} &= 4 \times \text{side} \\
\text{side} &= \frac{{@Perimeter@}}{4} \\
\text{side} &= {@c@} \text{ m}
\end{aligned} \]
</p>

<p><b>Step 2: Use the diagonal property</b></p>
<p>The diagonals of a rhombus cut each other into two equal parts and intersect at right angles.</p>
<ul>
  <li>One diagonal is \({@One_diagonal@}\text{ m}\), so half of it is \(\frac{{@One_diagonal@}}{2} = {@a@}\text{ m}\).</li>
  <li>Let the other diagonal be \(d_2\). Half of it is \(\frac{d_2}{2}\).</li>
</ul>
<p>Using Pythagoras’ Theorem on one of the inner right-angled triangles:</p>
<p>
\[ \begin{aligned}
\text{side}^2 &= \left( \frac{d_1}{2} \right)^2 + \left( \frac{d_2}{2} \right)^2 \\
{@c@}^2 &= ({@a@})^2 + \left( \frac{d_2}{2} \right)^2 \\
\left( \frac{d_2}{2} \right)^2 &= {@c@}^2 - {@a@}^2 \\
\frac{d_2}{2} &= \sqrt{{@c^2@} - {@a^2@}} \\
\frac{d_2}{2} &= {@b@}
\end{aligned} \]
</p>
<p>Multiply by 2 to find the full length of the other diagonal:</p>
<p>
\[ \begin{aligned}
d_2 &= 2 \times {@b@} \\
d_2 &= {@other_diagonal@} \text{ m}
\end{aligned} \]
</p>

<p><b>Step 3: Find the area of the rhombus</b></p>
<p>The formula for the area of a rhombus is \(\text{Area} = \frac{d_1 \times d_2}{2}\).</p>
<p>
\[ \begin{aligned} 
\text{Area} &= \frac{{@One_diagonal@} \times {@other_diagonal@}}{2} \\
\text{Area} &= {@Area_of_Rhombus@} \text{ m}^2
\end{aligned} \]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Generate Pythagorean triple variables */
m : 2 + rand(15);
n : 1 + rand(m-1);

/* Dimensions of the Rhombus (Half diagonals and side) */
a : m^2 - n^2; 
b : 2*m*n ;  
c : m^2 + n^2; 

Perimeter : c * 4; 
One_diagonal : a * 2; 
other_diagonal : b * 2;
Area_of_Rhombus : (One_diagonal * other_diagonal) / 2;

/* Mistakes */
m_half_diag : b; /* Forgot to multiply half-diagonal by 2 */
area_1 : One_diagonal * other_diagonal; /* Forgot to divide by 2 for area */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Diag 2: {@other_diagonal@}, Area: {@Area_of_Rhombus@}, Perim: {@Perimeter@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>other_diagonal</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>Area_of_Rhombus</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks the correct length of the diagonal.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>other_diagonal</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Caught half-diagonal mistake.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m_half_diag</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You correctly used Pythagoras' theorem, but you only found <b>half</b> of the diagonal (one leg of the inner triangle). Remember to multiply by 2!</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct area of the rhombus.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>Area_of_Rhombus</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for the missing 1/2 in area formula.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>area_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Good effort! You correctly multiplied the diagonals together, but you forgot to divide by 2. The area of a rhombus is \(\frac{1}{2} \times d_1 \times d_2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>25</deployedseed>
    <deployedseed>1786286663</deployedseed>
    <deployedseed>1113624194</deployedseed>
    <deployedseed>1181814256</deployedseed>
    <deployedseed>1523760047</deployedseed>
    <deployedseed>1075853406</deployedseed>
    <deployedseed>1195382226</deployedseed>
    <deployedseed>1483663922</deployedseed>
    <deployedseed>66215198</deployedseed>
    <deployedseed>751570943</deployedseed>
    <deployedseed>1940047813</deployedseed>
    <deployedseed>1338870530</deployedseed>
    <deployedseed>604110427</deployedseed>
    <deployedseed>605333801</deployedseed>
    <deployedseed>1964125578</deployedseed>
    <deployedseed>217002894</deployedseed>
    <deployedseed>1887545119</deployedseed>
    <deployedseed>1073027289</deployedseed>
    <deployedseed>776810112</deployedseed>
    <deployedseed>1234221348</deployedseed>
    <deployedseed>2143451653</deployedseed>
    <deployedseed>1544665985</deployedseed>
    <deployedseed>1585813031</deployedseed>
    <deployedseed>439957085</deployedseed>
    <deployedseed>1247102708</deployedseed>
    <deployedseed>1037918588</deployedseed>
    <deployedseed>1286667582</deployedseed>
    <deployedseed>1502287647</deployedseed>
    <deployedseed>71137346</deployedseed>
    <deployedseed>2115432617</deployedseed>
    <deployedseed>1708142534</deployedseed>
    <deployedseed>719389398</deployedseed>
    <deployedseed>113703</deployedseed>
    <deployedseed>1248502879</deployedseed>
    <deployedseed>2141576985</deployedseed>
    <deployedseed>505875526</deployedseed>
    <deployedseed>1700378135</deployedseed>
    <deployedseed>1642493851</deployedseed>
    <deployedseed>1569467852</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>other_diagonal</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Area_of_Rhombus</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
