<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the side length of a rhombic community garden</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A community garden features a rhombus-shaped plot for pollinator-friendly plants. The total area of the plot is exactly <b>{@Area@} m²</b>. An irrigation line running straight across one of the diagonals measures <b>{@d1@} m</b>.</p>

<p>Calculate the length of one outer side of the plot (\(s\)).</p>

[[jsxgraph width="500px" height="300px"]]
    var d1 = {#d1#};
    var d2 = {#d2#};
    var A = {#Area#};
    
    // Half diagonals for drawing
    var a = d1 / 2;
    var b = d2 / 2;
    
    // Proportional scaling
    var ratio = 500 / 300;
    var spanX = Math.max(d1 * 1.3, d2 * 1.3 * ratio);
    var spanY = spanX / ratio;
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-spanX/2, spanY/2, spanX/2, -spanY/2], 
        maxboundingbox: [-spanX*1.5, spanY*1.5, spanX*1.5, -spanY*1.5],
        axis: false, 
        showCopyright: false,
        keepAspectRatio: true,
        showNavigation: true 
    });

    var p1 = board.create('point', [a, 0], {name: '', visible: false});
    var p2 = board.create('point', [0, b], {name: '', visible: false});
    var p3 = board.create('point', [-a, 0], {name: '', visible: false});
    var p4 = board.create('point', [0, -b], {name: '', visible: false});
    var center = board.create('point', [0,0], {name: '', visible: false});

    // Rhombus Plot
    board.create('polygon', [p1, p2, p3, p4], {
        fillColor: '#fff3e0', 
        strokeColor: '#e65100', 
        strokeWidth: 2
    });

    // Known diagonal (irrigation line)
    board.create('segment', [p3, p1], {dash: 2, strokeColor: '#fb8c00'});
    
    // Unknown diagonal (faint construction line)
    board.create('segment', [p4, p2], {dash: 2, strokeColor: '#ffe0b2'});

    // Right angle marker at intersection
    var rad = Math.min(a, b) * 0.15;
    board.create('angle', [p1, center, p2], {type: 'square', radius: rad, strokeColor: '#fb8c00', name: ''});

    var gapY = spanY * 0.015; 
    var gapX = spanX * 0.015; 

    // Dynamic Labels
    board.create('text', [a/2, -gapY, "d_1 = " + d1 + " m"], {anchorX: 'middle', anchorY: 'top', fontSize: 14});
    board.create('text', [a/2 + gapX*2, b/2 + gapY*2, "s = ?"], {anchorX: 'left', anchorY: 'bottom', fontSize: 16, strokeColor: '#c62828'});
    
    board.create('text', [0, -b - gapY*4, "Area = " + A + " m²"], {anchorX: 'middle', anchorY: 'top', fontSize: 16, strokeColor: '#e65100'});
[[/jsxgraph]]

<p>Side length (\(s\)) = [[input:ans1]] <b>m</b> [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p><b>Step 1: Find the length of the missing diagonal (\(d_2\))</b></p>
<p>The formula for the area of a rhombus is:</p>
<p>
\[ \text{Area} = \frac{d_1 \times d_2}{2} \]
</p>
<p>Substitute the known values (\(\text{Area} = {@Area@}\) and \(d_1 = {@d1@}\)):</p>
<p>
\[ \begin{aligned}
{@Area@} &= \frac{{@d1@} \times d_2}{2} \\
{@Area * 2@} &= {@d1@} \times d_2 \\
d_2 &= \frac{{@Area * 2@}}{{@d1@}} \\
d_2 &= {@d2@} \text{ m}
\end{aligned} \]
</p>

<p><b>Step 2: Use the diagonal property to find the side length</b></p>
<p>The diagonals of a rhombus bisect each other (cut in half) at right angles, forming four right-angled triangles inside.</p>
<ul>
  <li>Half of \(d_1 = \frac{{@d1@}}{2} = {@a@}\text{ m}\)</li>
  <li>Half of \(d_2 = \frac{{@d2@}}{2} = {@b@}\text{ m}\)</li>
</ul>
<p>Using Pythagoras’ Theorem on one of these inner right-angled triangles, the side length (\(s\)) is the hypotenuse:</p>
<p>
\[ \begin{aligned}
s^2 &= ({@a@})^2 + ({@b@})^2 \\
s^2 &= {@a^2@} + {@b^2@} \\
s^2 &= {@a^2 + b^2@} \\
s &= \sqrt{{@a^2 + b^2@}} \\
s &= {@side@} \text{ m}
\end{aligned} \]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Generate Pythagorean triple variables to ensure integer results */
m : 2 + rand(8);
n : 1 + rand(m-1);

/* Half diagonals and side */
a : m^2 - n^2; 
b : 2*m*n ;  
c : m^2 + n^2; 

/* Final Variables */
d1 : a * 2; 
d2 : b * 2;
Area : (d1 * d2) / 2;
side : c;

/* Teacher Answer */
ta : side;

/* Mistakes for PRT */
m1 : c * 4;                 /* Found perimeter instead of side length */
m2 : d2;                    /* Found the missing diagonal and stopped early */
m3 : float(sqrt(d1^2 + d2^2)); /* Used full diagonals in Pythagoras instead of halves */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Area: {@Area@}, d1: {@d1@}, d2: {@d2@}, Side: {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct side length</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Caught Perimeter mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the total perimeter of the plot. The question asked for the length of just one side.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Caught stopping early at missing diagonal</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.3</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You successfully found the length of the missing diagonal (\(d_2\))! But you stopped too early. Use Pythagoras' theorem on half of each diagonal to find the outer side length.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Caught failing to halve diagonals before Pythagoras</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You used the full lengths of the diagonals in Pythagoras' theorem. You must divide each diagonal by 2 first, because the right-angled triangles are formed by the <i>halves</i> of the diagonals.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>485268952</deployedseed>
    <deployedseed>1185720276</deployedseed>
    <deployedseed>1800720921</deployedseed>
    <deployedseed>1561899063</deployedseed>
    <deployedseed>1098482511</deployedseed>
    <deployedseed>1867351408</deployedseed>
    <deployedseed>305895908</deployedseed>
    <deployedseed>1201096622</deployedseed>
    <deployedseed>467886528</deployedseed>
    <deployedseed>1903839676</deployedseed>
    <deployedseed>1176625483</deployedseed>
    <deployedseed>1223538354</deployedseed>
    <deployedseed>324395404</deployedseed>
    <deployedseed>1068401248</deployedseed>
    <deployedseed>1501757165</deployedseed>
    <deployedseed>1778201243</deployedseed>
    <deployedseed>1749652180</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
