<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of an amusement park that is pentagonal</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>An amusement park is the shape of a regular pentagon. Each side of the park measures \({@side@}\,\text{m}\) and the height of the pentagon park is \({@height@} \, \text{m}\).  Calculate the total area of the park.</p>



<p>\(\text{The area of the park} =\) [[input:ans1]]\(\,\text{m}^2\)[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Given:</b><br>
• Side length = {@side@} m <br>
• Height of one triangle (from the center to the midpoint of a side) = {@height@} m
</p>

<p>
\(\text{Area of pentagon} = 5 \times \text{area of one triangle}\)
</p>

<p><b>Find the area of one triangle</b></p>

<p>
Each triangle has: <br>
• Base = side of the pentagon = {@side@} m <br>
• Height = {@height@} m <br>
\(
\begin{align}
\text{Area of one triangle} &= \frac{1}{2} \times \text{base} \times \text{height} \\
&= \frac{1}{2} \times {@side@} \times {@height@} \\
&= {@triangle_area@}\, \text{m}^2
\end{align}
\)
</p>

<p><b>Find the total area of the pentagon</b></p>

<p>
\(
\begin{align}
\text{Total area} &= 5 \times \text{area of one triangle} \\
&= 5 \times {@triangle_area@} \\
&= {@total_area@}\, \text{m}^2
\end{align}
\)
</p>

<p>
\(\text{Therefore, the total area of the pentagon-shaped park is } {@total_area@} \text{ m}^2.\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomize side (even numbers between 12 m and 24 m) */
side : 2 * (rand(59) + 6);

/* Randomize height (even numbers less than side, between 4 m and side - 2) */
height : 2 * (rand(floor(side/2 - 1)) + 2);

/* Calculate area of one triangle */
triangle_area : (1/2) * side * height;

/* Calculate total area of the pentagon */
total_area : 5 * triangle_area;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The side of the park}={@side@}\, \text{m}\)</p>
<p>\(\text{The height of the park}={@height@}\, \text{m}\)</p>
<p>\(\text{The total area of the park} ={@total_area @}\, \text{m}^2\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>total_area</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct area of the Park</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>total_area </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
