<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of a rhombus, given the height and base</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the area of the figure below.</p>
[[jsxgraph]]
     (function() {
      const board = JXG.JSXGraph.initBoard(divid, {
        boundingbox: [-1, 6, 10, -2],
        axis: false,
        showNavigation: false,
        keepaspectratio: false,
        showCopyright: false
      });

      // Coordinates chosen so that:
      // side length = 5 units (representing 5 cm)
      // perpendicular height = 4 units (representing 4 cm)
      // A(0,0), B(5,0), C(8,4), D(3,4)
      const A = board.create('point', [0, 0], {name: 'A', fixed: true, visible: false ,size: 3});
      const B = board.create('point', [5, 0], {name: 'B', fixed: true, visible: false ,size: 3});
      const C = board.create('point', [8, 4], {name: 'C', fixed: true, visible: false ,size: 3});
      const D = board.create('point', [3, 4], {name: 'D', fixed: true,visible: false , size: 3});

      // Rhombus polygon
      const rhombus = board.create('polygon', [A, B, C, D], {
        borders: { strokeWidth: 2 },
        fillColor: '#f0f8ff',
        fillOpacity: 0.5
      });

      // Draw perpendicular height from D to base AB
      // Foot of perpendicular E is (3,0)
      const E = board.create('point', [3, 0], {name: 'E', visible: false, size: 2, face: 'o'});
      const heightSegment = board.create('segment', [D, E], {dash: 2, strokeWidth: 2, strokeColor: 'blue'});

      // Draw a visible right angle marker at E (manually as a small square)
      const RA1 = board.create('point', [E.X(), E.Y()], {visible: false});
      const RA2 = board.create('point', [E.X() + 0.4, E.Y()], {visible: false});
      const RA3 = board.create('point', [E.X() + 0.4, E.Y() + 0.4], {visible: false});
      const RA4 = board.create('point', [E.X(), E.Y() + 0.4], {visible: false});
      board.create('polygon', [RA1, RA2, RA3, RA4], {
        borders: {strokeColor: 'black', strokeWidth: 1},
        fillColor: 'black',
        fillOpacity: 0.3
      });
      // Label side lengths "5 cm" on each side (placed at midpoints)
      function midLabel(P, Q, text) {
        const midx = (P.X() + Q.X()) / 2;
        const midy = (P.Y() + Q.Y()) / 2;
        return board.create('text', [midx, midy, text], {anchorX: 'middle', anchorY: 'bottom', fontSize: 14});
      }
      midLabel(C, D, '{#Base#} cm');
      // Label the perpendicular height (4 cm) near the segment
      board.create('text', [3.2, 2, '{#Height#} cm'], {fontSize: 14})
      // Make sure labels stay readable if board is resized
      board.on('resize', function() { board.update(); });
    })();
[[/jsxgraph]]
<p>\(\text{Area of the Rhombus}=\)[[input:ans1]] \(\text{cm}^2\) [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The area of a rhombus is given by:<br>
\( \text{Area} = \text{Base} \times \text{Height} \)
</p>

<p>Given that:<br>
\(\text{Base} = {@Base@}\,\text{cm}\)<br>
\(\text{Height} = {@Height@}\,\text{cm}\)
</p>

<p>Substitute the given values into the formula to get the area as shown below.</p>

<p>\(
\begin{align}
\text{Area of a rhombus} &= \text{Base} \times \text{Height} \\
&= {@Base@} \times {@Height@} \\
&= {@area@}\,\text{cm}^2
\end{align}
\)
</p>

<p>Therefore:<br>
\(\text{The area of the rhombus} = {@area@}\,\text{cm}^2\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>Base : rand(40) + 2;    /* base between 2 and 16 */
Height : rand_with_prohib(2,30,[base]);  /* height between 2 and 16 */

area :Base * Height;

/*Mistakes*/
/*Using the area of right angle formula instade of base ties height */
area_1:(1/2)*Base * Height;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The height of the rhombus }={@Height@}\,\text{cm}\)</p>
<p>\(\text{The side of the rhombus }={@Base@}\,\text{cm}\)</p>
<p>\(\text{The area of the rhombus }={@area@}\,\text{cm}^2\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>area</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
