<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying and Determining the Dimensions of a Cone</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A cone has a perpendicular height of \({@h_cm@}\,\text{cm}\) and a slant height of
\({@l_cm@}\,\text{cm}\). Take \(\pi = \frac{22}{7}\).
</p>
<p>
Determine the following:
</p>

<ol type="i">
  <li><p>the radius of the cone</p>
   <p>\(\text{Radius}=\)[[input:ans1]]\(\,\text{cm}\)[[validation:ans1]][[feedback:prt1]]</p>
     </li>
  <li><p>the volume of the cone to the nearest whole numbers</p>
   <p>\(\text{Volume}=\)[[input:ans2]]\(\,\text{cm}^3\)[[validation:ans2]][[feedback:prt2]]</p>

    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>
The cone has a perpendicular height of \({@h_cm@}\,\text{cm}\) and a slant height of
\({@l_cm@}\,\text{cm}\).
</p>

<p><strong>Step 1: Determine the radius of the cone</strong></p>

<p>
For a cone, the radius \(r\), the perpendicular height \(h\), and the slant height \(l\)
form a right-angled triangle.
</p>

<p>
\(\begin{align}
l^2 &= r^2 + h^2
\end{align}\)
</p>

<p>
Substituting the given values:
</p>

<p>
\(\begin{align}
({@l_cm@})^2 &= r^2 + ({@h_cm@})^2 \\
r^2 &= ({@l_cm@})^2 - ({@h_cm@})^2 \\
r^2 &= {@l_cm@}^2 - {@h_cm@}^2 \\
r &= \sqrt{{@l_cm@}^2 - {@h_cm@}^2}\\
&={@r_cm@} \,\text{cm}
\end{align}\)
</p>



<p><strong>Step 2: Find the volume of the cone</strong></p>

<p>
The formula for the volume of a cone is:
</p>

<p>
\(\begin{align}
V &= \frac{1}{3} \times \text{base area} \times \text{height}\\
\text{Base area} &= \pi r^2 \\
&= \frac{22}{7} \times ({@r_cm@})^2\\
&={@base_area@}\,\text{cm}^2\\
\end{align}\)
</p>

<p>Therefore,<br>

\(\begin{align}
V &= \frac{1}{3} \times {@base_area@}\times {@h_cm@} \\
&=\frac{1}{3} \times {@B_H@} \\
&={@ta@}\\
&={@ta_1@}\ \text{cm}^3
\end{align}\)
</p>

<p>
Rounded to the nearest whole number:
<br>
\(\begin{align}
V &= {@ta_1@}\ \text{cm}^3
\end{align}\)
</p>

<p>
Therefore:
</p>

<ol type="i">
  <li>
    <p>The radius of the cone is \({@r_cm@}\,\text{cm}\).</p>
  </li>
  <li>
    <p>The volume of the cone is <strong>\({@ta_1@}\,\text{cm}^3\)</strong>.</p>
  </li>
</ol>]]></text>
    </generalfeedback>
    <defaultgrade>15</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* List of Pythagorean triples (radius, height, slant height) */
triples : [
  [3,4,5],
  [5,12,13],
  [6,8,10],
  [7,24,25],
  [8,15,17],
  [9,12,15],
  [10,24,26],
  [12,16,20],
  [12,35,37],
  [15,20,25],
  [16,30,34],
  [18,24,30],
  [20,21,29]
];

/* Randomly select one triple */
T : rand(length(triples)) + 1;

/* Assign cone dimensions in cm */
r_cm : triples[T][1];   /* radius */
h_cm : triples[T][2];   /* perpendicular height */
l_cm : triples[T][3];   /* slant height */


/* Area of the circular base */
base_area : 22/7 * r_cm^2;
B_H:base_area*h_cm;

/* Volume of the cone */
ta : (1/3) * base_area * h_cm;
ta_1:round(ta);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The perpendicular height of the cone}={@h_cm@}\,\text{cm}\)</p>
<p>\(\text{The slant height of the cone}={@l_cm@}\,\text{cm}\)</p>
<p>\(\text{The radius of the cone}={@r_cm@}\,\text{cm}\)</p>
<p>\(\text{The volume of the cone}={@ta_1@}\,\text{cm}^3\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>r_cm</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct radius of the cone.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>r_cm</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct volume of the cone.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1606768212</deployedseed>
    <deployedseed>511194647</deployedseed>
    <deployedseed>907885664</deployedseed>
    <deployedseed>857422780</deployedseed>
    <deployedseed>1435536546</deployedseed>
    <deployedseed>1024781509</deployedseed>
    <deployedseed>567462867</deployedseed>
    <deployedseed>1413201145</deployedseed>
    <deployedseed>1836286095</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>r_cm</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
