<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Evaluating a numerical square using a quadratic identity</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Evaluate \({@n@}^2\) using the quadratic identity
</p>
<p>\((a+b)^2=a^2+2ab+b^2.\)</p>

<p> = [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>
<p>Notice that {@n@} is close to the convenient number {@base@}, so write it as</p>
</p>
\[
{@n@} = {@base@} {@if offset>0 then "+" else ""@} {@offset@}.
\]
<p>
Use the quadratic identity:
</p>
\[
(a+b)^2 = a^2 + 2ab + b^2
\]
<p>
Substitute
</p>
\[
a={@base@}, \qquad
b= {@if offset>0 then "" else "("@} {@offset@} {@if offset>0 then "" else ")"@}.
\]
<p>
Then
</p>
\[
({@base@} {@if offset>0 then "+" else ""@} {@offset@})^2
=
{@base@}^2 + 2({@base@})({@offset@}) + ({@offset@})^2
\]
\[
=
{@base^2@} + ({@2*base*offset@}) + {@offset^2@}
\]
\[
=
\boxed{ {@ta@} }
\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>base : rand_with_step(20,200,10);

offset : if rand(2)=0 then 2 else -2;

n : base + offset;

ta : n^2;

/* Common mistakes */

/* 1. Omits the middle term: a² + b² */
m1 : base^2 + offset^2;

/* 2. Uses ±2a instead of ±2ab */
m2 : base^2 + 2*base*sign(offset) + offset^2;

/* 3. Incorrect sign on the middle term */
m3 : base^2 - 2*base*offset + offset^2;

/* 4. Squares the base only */
m4 : base^2;

/* 5. Treats (a±b)² as a² ± b² */
m5 : base^2 + sign(offset)*offset^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta@}</p>
<p> {@n@}</p>
<p>{@base@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Quadratic identities in numerical cases</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to evaluate a numerical expression of the form (a + b)² by applying the quadratic identity a² + 2ab + b².
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Missing the Middle Term</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You forgot the 2ab term completely. Remember: (𝑎+𝑏)^2=𝑎^2+2𝑎𝑏+𝑏^2</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Incorrect Multiplication</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Check how you expanded the square, you used 2a instead of 2ab.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Incorrect Middle Term</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> You used the wrong sign for the middle term. Check the worked solution below./p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>Use the identity: (𝑎+𝑏)^2=𝑎^2+2𝑎𝑏+𝑏^2</text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You did not simplify your answer.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1222860662</deployedseed>
    <deployedseed>1279341630</deployedseed>
    <deployedseed>1364955254</deployedseed>
    <deployedseed>1401621509</deployedseed>
    <deployedseed>1035152784</deployedseed>
    <deployedseed>1304848110</deployedseed>
    <deployedseed>771620108</deployedseed>
    <deployedseed>784142095</deployedseed>
    <deployedseed>2027149288</deployedseed>
    <deployedseed>482063284</deployedseed>
    <deployedseed>216898668</deployedseed>
    <deployedseed>625790673</deployedseed>
    <deployedseed>1912696106</deployedseed>
    <deployedseed>1767374472</deployedseed>
    <deployedseed>867120790</deployedseed>
    <deployedseed>1868238799</deployedseed>
  </question>
</quiz>
