<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>28.2 Similar Pyramids</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Here is a diagram of two similar, hexagonal-based pyramids.
<br>The top section of the green pyramid is covered by the red.</p>

[[jsxgraph height="500px" width="500px"]]

    var redHeight = {#redHeight#};
    var greenHeight = {#greenHeight#};

    // Set inital display conditions
    var zoom = 2;
    var cenOfView = [0,1.25];
    var orientation = [5,0.26,0.05];

    // Creating JSXGraph board: change "jxgbox" to divid for STACK questions
    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[cenOfView[0]-zoom, cenOfView[1]+zoom, cenOfView[0]+zoom, cenOfView[1]-zoom],
      showCopyright:false,
      showNavigation:false,
    });

    // Function to project point in 3D onto a plane through the origin normal to [xN,yN,zN];
    function flatten_point([x,y,z]) {
        return([x,y]);
    };

    // Functions to rotate point in 3D around x, y and z axes
    function rotate_x([x,y,z], angle) {
        return([
            x,
            Math.cos(angle)*y - Math.sin(angle)*z,
            Math.sin(angle)*y + Math.cos(angle)*z
        ]);
    };

    function rotate_y([x,y,z], angle) {
        return([
            Math.sin(angle)*z + Math.cos(angle)*x,
            y,
            Math.cos(angle)*z - Math.sin(angle)*x
        ]);
    };

    function rotate_z([x,y,z], angle) {
        return([
            Math.cos(angle)*x - Math.sin(angle)*y,
            Math.sin(angle)*x + Math.cos(angle)*y,
            z
        ]);
    };

    // Function to create a point-projection pair:
    function make_point([x,y,z], style) {

        // Rotate coordinates by display orientation
        var coords = rotate_z(rotate_y(rotate_x([x,y,z], orientation[0]), orientation[1]), orientation[2]);
        
        // Create point, defined by initial coords, display coords and projected coords
        var point = {
            initCoords:coords,
            dispCoords:coords,
            proj:board.create("point", flatten_point(coords), style)
        }
        elements.push(point);
        return(point);

    };
    
    // List containing all 3D elements to be displayed (used to update position for user movement)
    var elements = [];

    // CREATE POLYGON-BASED PYRAMID print_pyram(noSides, sideLength, height, baseCentreCoords, rotation, faceStyle, pointStyle)
    function print_pyram(
        noSides = 4, 
        sideLength = 1, 
        height = 1, 
        [x,y,z] = [0,0,0], 
        [rotX,rotY,rotZ] = [0,0,0],
        faceStyle = {
            fixed:true,
            highlight:false,
            borders:{
                fixed:true,
                highlight:false
            }
        },
        pointStyle = {visible:false}) {

        var pointList = [];
        
        // Create top point without rotation or translation
        pointList.push([0, 0, height]);

        // Create points in n-gon base without rotation or translation
        var i;
        for (i = 0; i < noSides; i++) {
            pointList.push([
                Math.cos(2*i*Math.PI/noSides)*sideLength/Math.sqrt(2*(1-Math.cos(2*Math.PI/noSides))),
                Math.sin(2*i*Math.PI/noSides)*sideLength/Math.sqrt(2*(1-Math.cos(2*Math.PI/noSides))),
                0]);
        };

        // For each defined point; rotate and translate by user specs.
        var j;
        for (i in pointList) {
            j = rotate_z(rotate_y(rotate_x(pointList[i], rotX), rotY), rotZ);
            pointList[i] = make_point([
                x + j[0],
                y + j[1],
                z + j[2]
            ], pointStyle);
        };

        // Create sloping triangular faces
        for (i in pointList.slice(1)) {
            board.create("polygon", [
                pointList[0].proj,
                pointList.slice(1)[i].proj,
                pointList.slice(1)[(Number(i)+1) % (pointList.length-1)].proj,
            ], faceStyle)
        };

        // Create polygon base
        board.create("polygon", pointList.slice(1).map(function(point) {return point.proj}), faceStyle);

        // Draw angle and dashed line
        board.create("angle", [pointList[1].proj, pointList[4].proj, pointList[0].proj], {
            fixed:true,
            highlight:false,
            withLabel:false,
            fillColor:"blue",
            strokeColor:"blue",
            fillOpacity:0.8,
            radius:0.3
        });

        board.create("segment", [pointList[1].proj, pointList[4].proj], {
            fixed:true,
            highlight:false,
            strokeWidth:1,
            dash:1
        });

    };

    print_pyram(6, 1, 3);
    print_pyram(6, redHeight/greenHeight, 3*redHeight/greenHeight, [0,0,3*(1-redHeight/greenHeight)], [0,0,0], {
        fixed:true,
        highlight:false,
        fillColor:"#FF2F2F",
        fillOpacity:1,
        borders:{
            fixed:true,
            highlight:false
        }
    });

    var arrowStyle = {
        fixed:true,
        highlight:false,
        strokeColor:"black",
        withLabel:true,
        label:{
            fixed:true,
            highlight:false,
            useMathJax:true,
            parse:false
        },
        firstArrow:{
            visible:true,
            size:4
        },
        lastArrow:{
            visible:true,
            size:4
        }
    };

    var p1 = make_point([redHeight/greenHeight + 0.1, 0, 3*(1-redHeight/greenHeight)], {visible:false});

    arrowStyle.name = "\\(".concat(redHeight,"cm\\)");
    arrowStyle.label.anchorX = "right";
    arrowStyle.label.offset = [-8,35];
    board.create("segment", [
        p1.proj,
        [p1.proj.X(), 2.9]
    ], arrowStyle);

    arrowStyle.name = "\\(".concat(greenHeight,"cm\\)");
    arrowStyle.label.anchorX = "left";
    arrowStyle.label.offset = [8,0];
    board.create("segment", [
        [p1.proj.X() + 0.4, 2.9],
        [p1.proj.X() + 0.4, 0.08]
    ], arrowStyle);

[[/jsxgraph]]

<p>(a) The angle at the base of the <b>red</b> pyramid is [[input:ans1]] the angle at the base of the <b>green</b> pyramid. [[validation:ans1]]</p>
<p align="right"><b>(1)</b></p>
<p>[[feedback:prt1]]</p>

<p><br>(b) Given that the area of the <b>green</b> hexagonal base is \({@float(greenArea)@}{cm}^2\), find the area of the <b>red</b> hexagonal base.</p>
<p align="right">\(A_R =\) [[input:ans2]] \({cm}^2\) [[validation:ans2]]</p>
<p align="right"><b>(2)</b></p>
<p>[[feedback:prt2]]</p>

<p><br>(c) Given that the volume of the <b>red</b> hexagonal pyramid is \({@float(redVol)@}{cm}^3\), find the volume of the <b>green</b> pyramid.
<br><i>Give your answer to 2dp</i></p>
<p align="right">\(V_G =\) [[input:ans3]] \({cm}^3\) [[validation:ans3]]</p>
<p align="right"><b>(3)</b></p>
<p>[[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>6</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[dp(num, dp) := block(if num*10^dp-floor(num*10^dp)<0.5 then return(float(floor(num*10^dp)/(10^dp))) else return(float(ceiling(num*10^dp)/(10^dp))));

lenRat:rand([12,14,15,16,18])/10;

/* Randomise similar lengths */

redHeight:rand_with_prohib(41,79,[50,60,70])/10;
greenHeight:redHeight*lenRat;

/* Randomise areas and volumes */

redArea: if lenRat=3/2 then 6*rand_with_step(11,19,1)/10 else 15*rand_with_step(3,7,1)/10;
greenArea:redArea*lenRat^2;
redVol:redArea*redHeight/3;
greenVol:greenArea*greenHeight/3;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>Red Height: {@redHeight@}, Green Height: {@greenHeight@}</p>
<p>(b) {@colourB@} area: {@if colourB="red" then redArea else if colourB="green" then greenArea@}</p>
<p>(c) {@colourC@} volume: {@if colourC="red" then redVol else if colourC="green" then greenVol@}</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans><![CDATA[[["less than", false], ["equal to", true], ["greater than", false]]]]></tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>float(redArea)</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>dp(greenVol,2)</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans><![CDATA["equal to"]]></tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Here is a worked solution:</p>
<p>Since the two pyramids are <i>similar</i>, corresponding angles between the shapes are equal.
<br>Therefore the angle at the base of the red pyramid is <b>equal to</b> the angle at the base of the green.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Node 1: Correct answer */
/* Node 2: Worked solution */

/* Node 3: Uses inverse of correct ratio (1) */
node3_test: if abs(ans2-greenArea*lenRat^2)<=0.1 then 1 else 0;

/* Node 4: Does not square ratio (1) */
node4_test: if abs(ans2-greenArea/lenRat)<=0.1 then 1 else 0;

/* Node 5: N3,4 errors combined (feedback) */
node5_test: if abs(ans2-greenArea*lenRat)<=0.1 then 1 else 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>redArea</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Here is a worked solution:</p>
<p>Because the two pyramids are similar, the ratio between corresponding areas is equal to the <b>square</b> of the ratio between corresponding lengths:</p>
<p>\(\frac{A_R}{A_G} = (\frac{h_R}{h_G})^2\)</p>
<p>Using the measurements given:</p>
<p>\(\frac{A_R}{{@float(greenArea)@}} = (\frac{{@float(redHeight)@}}{{@float(greenHeight)@}})^2\)</p>
<p>\(A_R = {@float(greenArea)@} \times ({@redHeight/greenHeight@})^2\)</p>
<p>\(A_R = {@float(redArea)@} {cm}^2\)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have used the <b>inverse</b> of the correct ratio to determine \(A_R\).
<br>Since the red pyramid is smaller than the green, your answer should be <b>less than</b> \({@float(greenArea)@}{cm}^2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have used the <b>square root</b> of the correct ratio to determine \(A_R\).
<br>Remember that the ratio between corresponding areas is equal to the <b>square</b> of the ratio between corresponding lengths.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have used the <b>inverse</b> and the <b>square root</b> of the correct ratio to determine \(A_R\).
<br>Remember that the ratio between corresponding areas is equal to the <b>square</b> of the ratio between corresponding lengths.
<br>Also, since the red pyramid is smaller than the green, your answer should be <b>less than</b> \({@float(greenArea)@}{cm}^2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>3.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Node 1: Correct answer */
/* Node 2: Worked solution */

/* Node 3: Correct answer, rounded incorrectly (2) */
node3_test: if abs(ans3-greenVol)<=0.1 then 1 else 0;

/* Node 4: Uses inverse of correct ratio (1) */
node4_test: if abs(ans3-redVol/lenRat^3)<=0.1 then 1 else 0;

/* Node 5: Does not cube ratio (1) */
node5_test: if abs(ans3-redVol*lenRat)<=0.1 or abs(ans3-redVol*lenRat^2)<=0.1 then 1 else 0;

/* Node 6: N4,5 errors combined (feedback) */
node6_test: if abs(ans3-redVol/lenRat)<=0.1 or abs(ans3-redVol/lenRat^2)<=0.1 then 1 else 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>dp(greenVol,2)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Here is a worked solution:</p>
<p>Because the two pyramids are similar, the ratio between their corresponding volumes is equal to the <b>cube</b> of the ratio between corresponding lengths:</p>
<p>\(\frac{V_G}{V_R} = (\frac{h_G}{h_R})^3\)</p>
<p>Using the measurements given:</p>
<p>\(\frac{V_G}{{@float(redVol)@}} = (\frac{{@float(greenHeight)@}}{{@float(redHeight)@}})^3\)</p>
<p>\(V_G = {@float(redVol)@} \times ({@greenHeight/redHeight@})^3\)</p>
<p>\(V_G = {@dp(greenVol,5)@}...\)</p>
<p>\(V_G = {@dp(greenVol,2)@} {cm}^3\) <i>(tp 2pd)</i></p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node3_test</sans>
        <tans>1</tans>
        <testoptions>1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.6670000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have found the correct volume of the green pyramid.
<br>However, you have not rounded your answer correctly to 2dp.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node4_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3330000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have used the <b>inverse</b> of the correct ratio to determine \(V_G\).
<br>Since the green pyramid is larger than the red, your answer should be <b>greater than</b> \({@float(redVol)@}{cm}^3\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt3-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node5_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.3330000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have used {@if abs(ans3-redVol*lenRat)<=0.1 then "the <b>cube root</b> of the correct ratio" else if abs(ans3-redVol*lenRat^2)<=0.1 then "the <b>square</b> of the ratio between the heights"@} to determine \(V_G\).
<br>Remember that the ratio between corresponding volumes is equal to the <b>cube</b> of the ratio between corresponding lengths.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt3-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>node6_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.0000000</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have used {@if abs(ans3-redVol/lenRat)<=0.1 then "the <b>inverse</b> and the <b>cube root</b> of the correct ratio" else if abs(ans3-redVol/lenRat^2)<=0.1 then "the <b>inverse</b> of the <b>square</b> of the ratio between the heights"@} to determine \(V_G\).
<br>Remember that the ratio between corresponding volumes is equal to the <b>cube</b> of the ratio between corresponding lengths.
<br>Also, since the green pyramid is larger than the red, your answer should be <b>greater than</b> \({@float(redVol)@}{cm}^2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-6-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>800790404</deployedseed>
    <deployedseed>1915192725</deployedseed>
    <deployedseed>812940126</deployedseed>
    <deployedseed>961058168</deployedseed>
    <deployedseed>1207541706</deployedseed>
    <deployedseed>1118925052</deployedseed>
    <deployedseed>182062003</deployedseed>
    <deployedseed>1175158608</deployedseed>
    <deployedseed>731349282</deployedseed>
    <deployedseed>165201793</deployedseed>
    <deployedseed>1221113962</deployedseed>
    <deployedseed>1630603464</deployedseed>
    <deployedseed>827222267</deployedseed>
    <deployedseed>1795253660</deployedseed>
    <deployedseed>1810133566</deployedseed>
    <deployedseed>1646393571</deployedseed>
    <deployedseed>1102141803</deployedseed>
    <deployedseed>172025673</deployedseed>
    <deployedseed>2132198628</deployedseed>
    <deployedseed>1724713669</deployedseed>
  </question>
</quiz>
