<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Capacity of a Larger Container Using Similarity</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Two similar containers have heights of {@a@} cm and {@b@} cm respectively. If the smaller container holds {@c@} ml, what is the capacity of the larger container? <b>Note:</b> Write your answer correct to <b>2</b> decimal places.</p>
<p>[[input:ans1]] ml[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Given the linear scale factor (L.S.F)= \(\frac{{@b@}}{{@a@}}\)</p>
<p>Volume scale factor (V.S.F) = \((L.S.F)^3 = (\frac{{@b@}}{{@a@}})^3 = {@k^3@}\)</p>
<p>Given the capacity of the small container as {@c@} ml, then the capacity of larger container is given by;</p>
<p>\(\quad \quad \quad \quad \quad \quad  \frac{{@b1@}}{{@a1@}}= \frac{x}{{@c@}} \Rightarrow x = \frac{{@c@}\times {@b1@}}{{@a1@}} = {@larger_container@}={@ta@}\) ml</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[ /* Side lengths of similar regular polygons */
a : rand(9) + 3;   /* side length of smaller polygon */
b : a + rand(9) + 1; /* side length of larger polygon, ensure b > a */

/* volume of the smaller container */
c : rand_with_step(25,300,5) ;

/* Scale factor between corresponding sides */
k : b / a;
  /*cubic scale*/
b1:b^3;
a1:a^3;

/* volume of the larger container (area scales with square of side ratio) */
 larger_container : c * k^3;
ta:decimalplaces( larger_container,2);


/*mistakes*/
ta1:k*c;
ta2:decimalplaces((1/k^3)*c,2);
ta3:c * k^2;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>volume of the smaller container={@c@}</p>
<p>volume of the larger container={@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Using the length scale factor directly</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
It looks like you may have used the scale factor for lengths directly to find the new capacity.
Remember, when containers are similar, lengths scale by the scale factor, but <strong>volumes</strong>
scale by the <strong>cube</strong> of the scale factor.
Make sure you cube the ratio of the corresponding heights before multiplying it by the original volume.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Inverting the scale factor</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
It appears that you may have used the ratio of the larger height to the smaller height in the wrong order.
When scaling up from the smaller container to the larger one. Check the direction of scaling before calculating the volume.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Squaring the scale facto</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
It seems like you squared the ratio of the heights when finding the capacity.Squaring the scale factor gives the ratio of <strong>areas</strong>, not volumes.
Since capacity depends on volume, you must use the <strong>cube</strong> of the scale factor instead.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1888312194</deployedseed>
    <deployedseed>1797272023</deployedseed>
    <deployedseed>452151587</deployedseed>
    <deployedseed>297282071</deployedseed>
    <deployedseed>642147960</deployedseed>
    <deployedseed>1387687611</deployedseed>
    <deployedseed>1083938283</deployedseed>
    <deployedseed>1583672595</deployedseed>
    <deployedseed>158736662</deployedseed>
    <deployedseed>1537002686</deployedseed>
    <deployedseed>1486635612</deployedseed>
    <deployedseed>142836434</deployedseed>
    <deployedseed>1258156807</deployedseed>
    <deployedseed>1348420741</deployedseed>
    <deployedseed>384916768</deployedseed>
    <deployedseed>968167599</deployedseed>
    <deployedseed>2024802024</deployedseed>
    <deployedseed>1893331229</deployedseed>
    <deployedseed>1994180150</deployedseed>
    <deployedseed>6799222</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Using the length scale factor directly</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Inverting the scale factor</description>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Squaring the scale factor</description>
      <testinput>
        <name>ans1</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
