<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Describing a transformation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The diagram shows two triangles:</p>
<ul>
  <li>\(\triangle ABC\) (blue) is the original.</li>
  <li>\(\triangle A_1B_1C_1\) (red) is its image after a single transformation.</li>
</ul>

[[jsxgraph height='420px' width='420px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-7, 7, 7, -7],
    keepaspectratio: true,
    showCopyright: false,
    axis: true,
    showNavigation: false,
    pan: { enabled: false },
    zoom: { enabled: false },
    defaultAxes: {
            x: { ticks: { insertTicks: false, ticksDistance: 1, minorTicks: 0, drawLabels: true } },
            y: { ticks: { insertTicks: false, ticksDistance: 1, minorTicks: 0, drawLabels: true } }
        }
});

var A = board.create('point', [{#x1#}, {#y1#}], { name: 'A', size: 2, fixed: true, color: '#1f77b4' });
var B = board.create('point', [{#x2#}, {#y2#}], { name: 'B', size: 2, fixed: true, color: '#1f77b4' });
var C = board.create('point', [{#x3#}, {#y3#}], { name: 'C', size: 2, fixed: true, color: '#1f77b4' });
board.create('polygon', [A, B, C], {
    borders: { strokeColor: '#1f77b4', strokeWidth: 2 },
    fillColor: '#1f77b4', fillOpacity: 0.15, fixed: true, vertices: { visible: false }
});

var A1 = board.create('point', [{#xx1#}, {#yy1#}], { name: 'A_1', size: 2, fixed: true, color: '#d62728' });
var B1 = board.create('point', [{#xx2#}, {#yy2#}], { name: 'B_1', size: 2, fixed: true, color: '#d62728' });
var C1 = board.create('point', [{#xx3#}, {#yy3#}], { name: 'C_1', size: 2, fixed: true, color: '#d62728' });
board.create('polygon', [A1, B1, C1], {
    borders: { strokeColor: '#d62728', strokeWidth: 2 },
    fillColor: '#d62728', fillOpacity: 0.15, fixed: true, vertices: { visible: false }
});
[[/jsxgraph]]

<p>Which of the following best describes the transformation that maps \(\triangle ABC\) onto \(\triangle A_1B_1C_1\)?</p>
<div>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>Answer:</strong> {@description@}.</p>
<p>A reflection flips the shape over a mirror line, so the image is "back to front": each vertex ends up the same perpendicular distance from the line, on the opposite side.</p>
<p><em>Explore:</em> drag the blue vertices to make your own triangle, choose a mirror line, and compare with your group. Use <strong>Clear</strong> to start again.</p>

[[jsxgraph height='420px' width='420px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-9, 9, 9, -9],
    axis: true, keepaspectratio: true, showCopyright: false
});

var A = board.create('point', [{#x1#}, {#y1#}], { name: 'A', size: 2, color: 'blue', snapToGrid: true });
var B = board.create('point', [{#x2#}, {#y2#}], { name: 'B', size: 2, color: 'blue', snapToGrid: true });
var C = board.create('point', [{#x3#}, {#y3#}], { name: 'C', size: 2, color: 'blue', snapToGrid: true });
board.create('polygon', [A, B, C], { borders: { strokeColor: 'blue' }, fillColor: 'lightblue', fillOpacity: 0.3 });

var imgObjs = [];
function clearImg() { imgObjs.forEach(function (o) { board.removeObject(o); }); imgObjs = []; }

function reflectShow(mapFn, lp, lq) {
    clearImg();
    imgObjs.push(board.create('line', [lp, lq], { strokeColor: '#888888', dash: 2, strokeWidth: 2 }));
    var a = board.create('point', mapFn(A), { name: 'A_1', color: 'red', size: 2, fixed: true });
    var b = board.create('point', mapFn(B), { name: 'B_1', color: 'red', size: 2, fixed: true });
    var c = board.create('point', mapFn(C), { name: 'C_1', color: 'red', size: 2, fixed: true });
    var g = board.create('polygon', [a, b, c], { borders: { strokeColor: 'red' }, fillColor: 'pink', fillOpacity: 0.3 });
    imgObjs.push(a, b, c, g);
}

board.create('button', [-8.5, 8, 'Reflect in x-axis', function () {
    reflectShow(function (P) { return [P.X(), -P.Y()]; }, [0, 0], [1, 0]);
}]);
board.create('button', [-8.5, 6.7, 'Reflect in y-axis', function () {
    reflectShow(function (P) { return [-P.X(), P.Y()]; }, [0, 0], [0, 1]);
}]);
board.create('button', [-8.5, 5.4, 'Reflect in y = x', function () {
    reflectShow(function (P) { return [P.Y(), P.X()]; }, [0, 0], [1, 1]);
}]);
board.create('button', [-8.5, 4.1, 'Reflect in y = -x', function () {
    reflectShow(function (P) { return [-P.Y(), -P.X()]; }, [0, 0], [1, -1]);
}]);
board.create('button', [-8.5, 2.8, 'Clear', function () { clearImg(); }]);
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Curated (triangle, line) variants.  line = [kind, off]:
   kind 1 = horizontal y = off ; 2 = vertical x = off ; 3 = y = x + off ; 4 = y = -x + off.
   Each triangle sits fully on one side of its line, so the image is clear of the original,
   and all coordinates stay on the board. */
variants : [
  [[[-4,-3],[-4,-1],[1,-4]],1,0],
  [[[-2,-3],[2,-4],[3,-2]],1,0],
  [[[-4,-4],[-4,1],[-2,-4]],2,0],
  [[[2,0],[3,-3],[4,2]],2,0],
  [[[-2,0],[0,3],[2,4]],3,0],
  [[[-2,-3],[2,-4],[3,-2]],3,0],
  [[[-4,-3],[-4,-1],[1,-4]],4,0],
  [[[-4,-4],[-4,1],[-2,-4]],4,0],
  [[[-2,0],[2,1],[3,0]],1,3],
  [[[-3,2],[0,0],[0,2]],1,3],
  [[[-1,-1],[0,2],[2,-1]],1,-2],
  [[[-2,0],[2,1],[3,0]],1,-2],
  [[[-4,-4],[-4,1],[-2,-4]],2,1],
  [[[-4,-4],[-1,1],[0,-2]],2,1],
  [[[-1,-1],[0,2],[2,-1]],2,-2],
  [[[-1,-1],[0,2],[1,-2]],2,-2],
  [[[-2,0],[0,3],[2,4]],3,1],
  [[[-2,-3],[2,-4],[3,-2]],3,1],
  [[[-2,0],[0,3],[2,4]],3,-1],
  [[[-2,4],[2,2],[3,4]],3,-1],
  [[[-3,0],[1,-3],[1,-1]],4,1],
  [[[-4,-3],[-4,-1],[1,-4]],4,1],
  [[[-4,-3],[-4,-1],[1,-4]],4,-2],
  [[[-4,-4],[-4,1],[-2,-4]],4,-2]
];

V : rand(variants);
tri : V[1];
kind : V[2];
off : V[3];

/* general line reflection */
reflp(p) := if kind = 1 then [p[1], 2*off - p[2]]
            elseif kind = 2 then [2*off - p[1], p[2]]
            elseif kind = 3 then [p[2] - off, p[1] + off]
            else [off - p[2], off - p[1]];

orig : tri;
img : map(reflp, tri);

x1 : orig[1][1]; y1 : orig[1][2];
x2 : orig[2][1]; y2 : orig[2][2];
x3 : orig[3][1]; y3 : orig[3][2];
xx1 : img[1][1]; yy1 : img[1][2];
xx2 : img[2][1]; yy2 : img[2][2];
xx3 : img[3][1]; yy3 : img[3][2];

/* human-readable description of a line given [kind, off] */
desc(k, o) := if k = 1 then (if o = 0 then "Reflection in the x-axis" else sconcat("Reflection in the line y = ", o))
  elseif k = 2 then (if o = 0 then "Reflection in the y-axis" else sconcat("Reflection in the line x = ", o))
  elseif k = 3 then (if o = 0 then "Reflection in the line y = x" else sconcat("Reflection in the line y = x ", if o > 0 then sconcat("+ ", o) else sconcat("- ", -o)))
  else (if o = 0 then "Reflection in the line y = -x" else sconcat("Reflection in the line y = -x ", if o > 0 then sconcat("+ ", o) else sconcat("- ", -o)));

description : desc(kind, off);

/* wrong-line distractor: the "confusable" partner line -- swap the axis
   (y = k  <->  x = k) or flip the diagonal's sign (y = x + c  <->  y = -x + c).
   This is always a different line, so it is a genuine near-miss
   (e.g. correct x = -2  gives the distractor  y = -2). */
swapline(k, o) := if k = 1 then [2, o] elseif k = 2 then [1, o] elseif k = 3 then [4, o] else [3, o];
wsel : swapline(kind, off);
wrongline_desc : desc(wsel[1], wsel[2]);

/* MCQ options: token value, correct?, display text */
ta : random_permutation([
  [opt_correct,   true,  description],
  [opt_wrongline, false, wrongline_desc],
  [opt_rot,       false, "This transformation is a rotation"],
  [opt_trans,     false, "This is a translation (a slide, with no turning or flipping)"],
  [opt_none,      false, "This is not a transformation"]
]);

kindcode : kind;
lineoff : off;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text><![CDATA[<p>The animation below shows this reflection: \(\triangle ABC\) flips across the dashed mirror line onto \(\triangle A_1B_1C_1\).</p>

[[jsxgraph height='420px' width='420px']]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-7, 7, 7, -7],
    axis: true, keepaspectratio: true, showCopyright: false,
    showNavigation: false, pan: { enabled: false }, zoom: { enabled: false }
});

var kind = {#kindcode#};
var off = {#lineoff#};

// mirror line for the actual reflection
var lp, lq;
if (kind === 1) { lp = [0, off]; lq = [1, off]; }
else if (kind === 2) { lp = [off, 0]; lq = [off, 1]; }
else if (kind === 3) { lp = [0, off]; lq = [1, 1 + off]; }
else { lp = [0, off]; lq = [1, off - 1]; }
board.create('line', [lp, lq], { strokeColor: '#888888', dash: 2, strokeWidth: 2, fixed: true, highlight: false });

var ox = [{#x1#}, {#x2#}, {#x3#}], oy = [{#y1#}, {#y2#}, {#y3#}];
var ix = [{#xx1#}, {#xx2#}, {#xx3#}], iy = [{#yy1#}, {#yy2#}, {#yy3#}];

// original (blue) and image (red) fixed
var oP = [], iP = [];
for (var i = 0; i < 3; i++) {
    oP.push(board.create('point', [ox[i], oy[i]], { visible: false, fixed: true }));
    iP.push(board.create('point', [ix[i], iy[i]], { visible: false, fixed: true }));
}
board.create('polygon', oP, { borders: { strokeColor: '#1f77b4', strokeWidth: 2 }, fillColor: '#1f77b4', fillOpacity: 0.12, fixed: true, vertices: { visible: false } });
board.create('polygon', iP, { borders: { strokeColor: '#d62728', strokeWidth: 2, dash: 2 }, fillColor: '#d62728', fillOpacity: 0.10, fixed: true, vertices: { visible: false } });

var labs = ['A', 'B', 'C'];
for (var i = 0; i < 3; i++) {
    board.create('text', [ox[i] + 0.2, oy[i] + 0.3, labs[i]], { fontSize: 11, color: '#1f77b4', fixed: true });
    board.create('text', [ix[i] + 0.2, iy[i] + 0.3, labs[i] + '₁'], { fontSize: 11, color: '#d62728', fixed: true });
}

// animated triangle sliding original -> image and back
var s = 0, dir = 1;
function lerp(i) { return [(1 - s) * ox[i] + s * ix[i], (1 - s) * oy[i] + s * iy[i]]; }
var aP = [];
for (var i = 0; i < 3; i++) {
    (function (j) {
        aP.push(board.create('point',
            [function () { return lerp(j)[0]; }, function () { return lerp(j)[1]; }],
            { visible: false }));
    })(i);
}
board.create('polygon', aP, { borders: { strokeColor: '#2ca02c', strokeWidth: 3 }, fillColor: '#2ca02c', fillOpacity: 0.25, fixed: true, vertices: { visible: false } });

setInterval(function () {
    s += dir * 0.01;
    if (s >= 1) { s = 1; dir = -1; }
    if (s <= 0) { s = 0; dir = 1; }
    board.update();
}, 30);

board.renderer.svgRoot.style.pointerEvents = "none";
[[/jsxgraph]]]]></text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@description@} (kind {@kindcode@}, off {@lineoff@}); triangle {@orig@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct description?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Reflection, but wrong line?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>opt_wrongline</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have correctly spotted that this is a <strong>reflection</strong> &mdash; but that is not the right mirror line. The mirror line lies halfway between each point and its image, at right angles to the line joining them. Look at where just one vertex, say \(A\), maps to.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Chose rotation?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>opt_rot</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Look at the orientation. Under a <strong>rotation</strong> the triangle would keep its sense (going \(A\rightarrow B\rightarrow C\) the same way round). Here \(\triangle A_1B_1C_1\) is a mirror image &mdash; flipped over &mdash; so it is a reflection, not a rotation.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Chose translation?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>opt_trans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>A <strong>translation</strong> slides every point the same distance in the same direction, keeping the shape facing the same way. Here the image has been flipped over, so it is a reflection, not a translation.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description><![CDATA[Chose "not a transformation"?]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>opt_none</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It <em>is</em> a transformation: every point of \(\triangle ABC\) has a matching image point in \(\triangle A_1B_1C_1\), the same shape and size. Decide whether it has been slid, turned, or flipped.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>382083143</deployedseed>
    <deployedseed>2141205395</deployedseed>
    <deployedseed>1877432223</deployedseed>
    <deployedseed>2000490574</deployedseed>
    <deployedseed>145530277</deployedseed>
    <deployedseed>908471222</deployedseed>
    <deployedseed>1633028841</deployedseed>
    <deployedseed>77412650</deployedseed>
    <deployedseed>1290556743</deployedseed>
    <deployedseed>1994871005</deployedseed>
    <deployedseed>560221389</deployedseed>
    <deployedseed>1841002267</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct description gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description><![CDATA[Wrong reflection line -> reflection diagnostic.]]></description>
      <testinput>
        <name>ans1</name>
        <value>opt_wrongline</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description><![CDATA[Rotation distractor -> rotation diagnostic.]]></description>
      <testinput>
        <name>ans1</name>
        <value>opt_rot</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description><![CDATA[Translation distractor -> translation diagnostic.]]></description>
      <testinput>
        <name>ans1</name>
        <value>opt_trans</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description><![CDATA["Not a transformation" -> that diagnostic.]]></description>
      <testinput>
        <name>ans1</name>
        <value>opt_none</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-5-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
