<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the area of a rectangle</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the area of the rectangle below.</p>

[[jsxgraph width="500px" height="300px"]]
    var l = {#length#};
    var w = {#width#};
    
    // Mathematically center the rectangle and force a 5:3 bounding box ratio
    var centerX = l / 2;
    var centerY = w / 2;
    var ratio = 500 / 300;
    
    // Ensure both dimensions fit with a 30% visual buffer
    var spanX = Math.max(l * 1.3, w * 1.3 * ratio);
    var spanY = spanX / ratio;
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [centerX - spanX/2, centerY + spanY/2, centerX + spanX/2, centerY - spanY/2], 
        maxboundingbox: [centerX - spanX*1.5, centerY + spanY*1.5, centerX + spanX*1.5, centerY - spanY*1.5],
        axis: false, 
        showCopyright: false,
        keepAspectRatio: true,
        showNavigation: true 
    });

    // Dynamic vertices linked to the STACK variables
    var A = board.create('point', [0, 0], {name: '', visible: false});
    var B = board.create('point', [l, 0], {name: '', visible: false});
    var C = board.create('point', [l, w], {name: '', visible: false});
    var D = board.create('point', [0, w], {name: '', visible: false});

    // Rectangle
    board.create('polygon', [A, B, C, D], {
        fillColor: 'none', 
        strokeColor: '#1976d2', 
        strokeWidth: 2
    });

    // Right-angle markers with cleaned names to suppress alpha labels
    var rad = Math.min(l, w) * 0.15;
    board.create('angle', [B, A, D], {type: 'square', radius: rad, strokeColor: '#90caf9', name: ''});
    board.create('angle', [C, B, A], {type: 'square', radius: rad, strokeColor: '#90caf9', name: ''});
    board.create('angle', [D, C, B], {type: 'square', radius: rad, strokeColor: '#90caf9', name: ''});
    board.create('angle', [A, D, C], {type: 'square', radius: rad, strokeColor: '#90caf9', name: ''});

    // Calculate exact visual pixel gaps (~15px) based on the canvas spans
    var gapY = spanY * 0.05; 
    var gapX = spanX * 0.03; 

    // Dynamic labels securely anchored with mathematically locked offsets
    board.create('text', [l/2, -gapY, l + " cm"], {anchorX: 'middle', anchorY: 'top', fontSize: 16});
    board.create('text', [l + gapX, w/2, w + " cm"], {anchorX: 'left', anchorY: 'middle', fontSize: 16});
[[/jsxgraph]]

<p>Area = [[input:ans1]] <b>\(\text{cm}^2\)</b> [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To calculate the area of a rectangle, multiply its length by its width.</p>
<p>
\[ \text{Area} = \text{Length} \times \text{width} \]
</p>
<p>
\[ \text{Area} = {@length@} \text{ cm} \times {@width@} \text{ cm} \]
</p>
<p>
\[ \text{Area} = {@area@} \text{ cm}^2 \]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Length and width */
length: rand_with_step(10,30,1);
width: rand_with_step(4,15,1);

/* Area */
area: length * width;

/* Common Mistakes for PRT */
m1: 2 * (length + width); /* Perimeter */
m2: length + width;       /* Added dimensions */
m3: length^2;             /* Squared the length */
m4: width^2;              /* Squared the width */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>L: {@length@}, W: {@width@}, Area: {@area@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>area</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>area</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Perimeter mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the perimeter. To find the area, you must multiply the length by the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Addition mistake</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You added the dimensions. The area is found by multiplying the length and the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Squares length</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is a rectangle, not a square. You squared the length. Use length &times; width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Squares width</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is a rectangle, not a square. You squared the width. Use length &times; width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>809552978</deployedseed>
    <deployedseed>1421145092</deployedseed>
    <deployedseed>1470533754</deployedseed>
    <deployedseed>487051735</deployedseed>
    <deployedseed>1534729632</deployedseed>
    <deployedseed>1413191561</deployedseed>
    <deployedseed>1993783787</deployedseed>
    <deployedseed>1395083328</deployedseed>
    <deployedseed>38681279</deployedseed>
    <deployedseed>718019919</deployedseed>
    <deployedseed>1990593261</deployedseed>
    <deployedseed>99169190</deployedseed>
    <deployedseed>1264986361</deployedseed>
    <deployedseed>1217944158</deployedseed>
    <deployedseed>62501220</deployedseed>
    <deployedseed>1447526856</deployedseed>
    <deployedseed>341695481</deployedseed>
    <deployedseed>1165869081</deployedseed>
    <deployedseed>1182137625</deployedseed>
    <deployedseed>1976342498</deployedseed>
    <deployedseed>155063598</deployedseed>
    <deployedseed>225640140</deployedseed>
    <deployedseed>594414533</deployedseed>
    <deployedseed>909526324</deployedseed>
    <deployedseed>233258209</deployedseed>
    <deployedseed>1627270235</deployedseed>
    <deployedseed>1435950085</deployedseed>
    <deployedseed>1497703288</deployedseed>
    <deployedseed>1790375086</deployedseed>
    <deployedseed>937608534</deployedseed>
    <deployedseed>1455301619</deployedseed>
    <deployedseed>1803489592</deployedseed>
    <deployedseed>1608616780</deployedseed>
    <deployedseed>1392298667</deployedseed>
    <deployedseed>998339330</deployedseed>
    <deployedseed>1280475100</deployedseed>
    <deployedseed>2116867174</deployedseed>
    <deployedseed>377283233</deployedseed>
    <deployedseed>739289029</deployedseed>
    <deployedseed>1800918197</deployedseed>
    <deployedseed>1874660211</deployedseed>
    <deployedseed>1187286197</deployedseed>
    <deployedseed>2119575143</deployedseed>
    <deployedseed>990315418</deployedseed>
    <deployedseed>453900839</deployedseed>
    <deployedseed>1454692051</deployedseed>
    <deployedseed>664578856</deployedseed>
    <deployedseed>710062744</deployedseed>
    <deployedseed>344281719</deployedseed>
    <deployedseed>248619833</deployedseed>
    <deployedseed>899040127</deployedseed>
    <deployedseed>1232925412</deployedseed>
    <deployedseed>1573888649</deployedseed>
    <deployedseed>810862366</deployedseed>
    <deployedseed>303321871</deployedseed>
    <deployedseed>260214599</deployedseed>
    <deployedseed>1124793643</deployedseed>
    <deployedseed>735181292</deployedseed>
    <deployedseed>1127675243</deployedseed>
    <deployedseed>1845521166</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>area</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
