<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Rotation of a triangle about the origin in a clockwise direction</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A triangle \(ABC\) with coordinates \( A\, ({@Ax@}, {@Ay@}), B\,({@Bx@}, {@By@})\, \text{and}\, C\, ({@Cx@}, {@Cy@}) \) is rotated through the centre and angle of \(90^\circ\)  in a clockwise direction.Find the coordinates of its image.</p>
<p>A' = [[input:ans1]] [[validation:ans1]]</p>
<p>B' = [[input:ans2]] [[validation:ans2]]</p>
<p>C' = [[input:ans3]] [[validation:ans3]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>Since triangle \(ABC\) is rotated in a clockwise direction with an angle of \(90^\circ\) through the origin, then the angle of rotation is \(-90^\circ\) </p>
<p>According to the rule, If we have our points \((x,y)\) which will be mapped to \((y,-x)\) if rotated through the centre and angle of rotation is \(-90^\circ\)
</p>
<p>Therefore, we will individually apply the rotation formula to all three given points.</p>
<p>\(A = ({@Ax@},{@Ay@}) \Rightarrow A' = {@ta@}\)</p>
<p>\(B =({@Bx@},{@By@}) \Rightarrow B' = {@ta1@}\)</p>
<p>\(C =({@Cx@},{@Cy@}) \Rightarrow C' = {@ta2@}\)</p>
<p>The coordinates of triangle \(A'B'C'\) are A' {@ta@}, B' {@ta1@} and C' {@ta2@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>Ax:rand_with_step(-9,9,1);
Ay:rand_with_prohib(1,9,[Ax]);

Bx:rand_with_prohib(1,9,[Ax,Ay]);
By:rand_with_prohib(1,9,[Ax,Ay,Bx]);

Cx:rand_with_prohib(1,9,[Ax,Ay,Bx,By]);
Cy:rand_with_prohib(1,9,[Ax,Ay,Bx,By,Ax]);
ta:[Ay,-Ax];
ta1:[By,-Bx];
ta2:[Cy,-Cx];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>The coordinates of triangle \(A'B'C'\) are A' {@ta@}, B' {@ta1@} and C' {@ta2@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>correst_ans : if (ans1 = ta and ans2 = ta1 and ans3 = ta2) then 1 else 0;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>correst_ans</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>679364820</deployedseed>
    <deployedseed>136983132</deployedseed>
    <deployedseed>66944115</deployedseed>
    <deployedseed>1141100948</deployedseed>
    <deployedseed>1044379300</deployedseed>
    <deployedseed>2046297580</deployedseed>
    <deployedseed>326870434</deployedseed>
    <deployedseed>10828010</deployedseed>
    <deployedseed>445285443</deployedseed>
    <deployedseed>660832721</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
