<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a triangle using two sides and an included angle</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Find the area of the triangle below with sides <b>{@a@} cm</b> and <b>{@b@} cm</b> and an included angle of <b>{@C_deg@}&deg;</b>.
</p>

[[jsxgraph]]
    var a = {#a#};
    var b = {#b#};
    var C_deg = {#C_deg#};
    var C_rad = C_deg * Math.PI / 180;
    
    // Calculate triangle vertices
    var p1x = 0; var p1y = 0;
    var p2x = b; var p2y = 0;
    var p3x = a * Math.cos(C_rad);
    var p3y = a * Math.sin(C_rad);
    
    // Dynamic Bounding Box to handle both acute and obtuse angles
    var minX = Math.min(-2, p3x - 2);
    var maxX = Math.max(b + 3, p3x + 3);
    var maxY = p3y + 3;
    var minY = -2;
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [minX, maxY, maxX, minY],
        keepAspectRatio: true,
        showNavigation: false,
        showCopyright: false
    });
    
    var p1 = board.create('point', [p1x, p1y], {visible: false});
    var p2 = board.create('point', [p2x, p2y], {visible: false});
    var p3 = board.create('point', [p3x, p3y], {visible: false});
    
    // Draw Triangle
    var poly = board.create('polygon', [p1, p2, p3], {
        fillColor: '#e0f7fa', fillOpacity: 0.6, 
        borders: {strokeColor: '#006064', strokeWidth: 2.5}
    });
    
    // Draw Included Angle
    var ang = board.create('angle', [p2, p1, p3], {
        radius: Math.min(a, b) * 0.3, 
        name: C_deg + "°",
        fillColor: '#80deea', fillOpacity: 0.8, strokeColor: '#006064'
    });
    
    // Clean Text Labels
    board.create('text', [b/2, -0.5, b + " cm"], {anchorX: 'middle', anchorY: 'top', fontSize: 15});
    board.create('text', [p3x/2 - 0.5, p3y/2 + 0.5, a + " cm"], {anchorX: 'right', anchorY: 'bottom', fontSize: 15});
[[/jsxgraph]]

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{cm}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
<strong>Step 1: Use the sine rule for the area of a triangle.</strong>
</p>

<p>
\[
\text{Area} = \frac{1}{2}ab\sin C
\]
</p>

<p>
<strong>Step 2: Substitute the known values.</strong>
</p>

<p>
The two known sides are \(a = {@a@}\) and \(b = {@b@}\), and the included angle is \(C = {@C_deg@}^\circ\).
</p>

<p>
\[
\text{Area} = \frac{1}{2} \times {@a@} \times {@b@} \times \sin({@C_deg@}^\circ)
\]
</p>

<p>
<strong>Step 3: Evaluate.</strong>
</p>

<p>
\[
\text{Area} \approx {@Area_disp@} \text{ cm}^2
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Sides */
a: rand(6) + 4;
b: rand(6) + 5;

/* Angle */
C_deg: rand([30, 45, 60, 120, 135, 150]);
C_rad: C_deg * %pi / 180;

/* Area calculated natively as float */
Area_float: float(0.5 * a * b * sin(C_rad));

/* Display variable formatted to exactly 2 decimal places */
Area_disp: dispdp(Area_float, 2);

/* Mistakes evaluated as floats */
m1: float(a * b * sin(C_rad));
m2: float(0.5 * a * b * abs(cos(C_rad)));
m3: float(0.5 * (a + b) * sin(C_rad));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>a: {@a@}, b: {@b@}, C: {@C_deg@}, Area: {@Area_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of triangle using two sides and included angle</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to calculate the area of a triangle using the formula involving two sides and the included angle.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Area_disp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Area_float</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Missing ½</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Remember, the formula for the area of a triangle requires you to multiply by \(\frac{1}{2}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>used cosine</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You used cosine. The area formula requires the <strong>sine</strong> of the included angle: \(\frac{1}{2}ab\sin C\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Wrong expression (addition instead of multiplication)</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You added the two side lengths together. The formula requires the <strong>product</strong> of the two sides: \(\frac{1}{2} \times a \times b \times \sin C\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1107839082</deployedseed>
    <deployedseed>2072928253</deployedseed>
    <deployedseed>1375833527</deployedseed>
    <deployedseed>610641191</deployedseed>
    <deployedseed>1345259980</deployedseed>
    <deployedseed>389520153</deployedseed>
    <deployedseed>400334539</deployedseed>
    <deployedseed>578656043</deployedseed>
    <deployedseed>717131728</deployedseed>
    <deployedseed>227223981</deployedseed>
    <deployedseed>1214313650</deployedseed>
    <deployedseed>391875102</deployedseed>
    <deployedseed>335823016</deployedseed>
    <deployedseed>397554682</deployedseed>
    <deployedseed>168029236</deployedseed>
    <deployedseed>1831148845</deployedseed>
    <deployedseed>1930543653</deployedseed>
    <deployedseed>9124657</deployedseed>
    <deployedseed>1111836430</deployedseed>
    <deployedseed>174851630</deployedseed>
    <deployedseed>215064935</deployedseed>
    <deployedseed>1096658477</deployedseed>
    <deployedseed>1615543242</deployedseed>
    <deployedseed>64073623</deployedseed>
    <deployedseed>1833511883</deployedseed>
    <deployedseed>1295203422</deployedseed>
    <deployedseed>1852276783</deployedseed>
    <deployedseed>561326941</deployedseed>
    <deployedseed>565210930</deployedseed>
    <deployedseed>1412009528</deployedseed>
    <deployedseed>1057501135</deployedseed>
    <deployedseed>1786304080</deployedseed>
    <deployedseed>1208874563</deployedseed>
    <deployedseed>1011695114</deployedseed>
    <deployedseed>1986422133</deployedseed>
    <deployedseed>133593616</deployedseed>
    <deployedseed>2006855941</deployedseed>
    <deployedseed>774833652</deployedseed>
    <deployedseed>1692466642</deployedseed>
    <deployedseed>4744047</deployedseed>
    <deployedseed>1649749482</deployedseed>
    <deployedseed>1054840634</deployedseed>
    <deployedseed>392607407</deployedseed>
    <deployedseed>1591652223</deployedseed>
    <deployedseed>2113658487</deployedseed>
    <deployedseed>604238119</deployedseed>
    <deployedseed>954402183</deployedseed>
    <deployedseed>1195537976</deployedseed>
    <deployedseed>738792544</deployedseed>
    <deployedseed>1506951240</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>Area_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
