<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of land using trigonometry</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A piece of land is in the shape of a triangle. Two sides measure <b>{@a@} m</b> and <b>{@b@} m</b>, and the angle between them is <b>{@C@}&deg;</b>. Find the area of the land.
</p>

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{m}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Step 1: Use the formula:
</p>

<p>
\[
A = \frac{1}{2}ab\sin C
\]
</p>

<p>
Step 2: Substitute the values:
</p>

<p>
\[
A = \frac{1}{2} \times {@a@} \times {@b@} \times \sin {@C@}^\circ
\]
</p>

<p>
Step 3: Evaluate:
</p>

<p>
\[
A \approx {@A_disp@}\text{ m}^2
\]
</p>

<p>
Therefore, the area of the land is <b>{@A_disp@} \(\text{m}^2\)</b>.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Sides (land lengths) */
a: rand(10) + 10;
b: rand(10) + 12;

/* Angle safely selected from a native list, expanded to include obtuse angles */
C: rand([40, 50, 60, 70, 80, 100, 110, 120]);

/* Convert to radians */
Cr: C*%pi/180;

/* Area evaluated natively as a float */
A_float: float(0.5*a*b*sin(Cr));

/* Display variable formatted to exactly 2 decimal places */
A_disp: dispdp(A_float, 2);

/* Mistakes evaluated as floats for PRT catching */
m1: float(a*b*sin(Cr));
m2: float(0.5*a*b*abs(cos(Cr)));
m3: float(a*b);
m4: float(0.5*(a+b)*sin(Cr));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>a: {@a@}, b: {@b@}, C: {@C@}, Area: {@A_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of triangle using two sides and included angle</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to apply the formula for the area of a triangle using two sides and the included angle in a real-life context involving land measurement.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A_disp</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>A_float</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Missing ½</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Do not forget the factor \(\frac{1}{2}\) in the area formula.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Used cosine</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You used cosine instead of sine. The formula requires the sine of the included angle.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>wrong structure</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the area as if it were a rectangle. You must include the sine of the angle and the \(\frac{1}{2}\) multiplier.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Multiplied only</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m4</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Use the product of the two sides, not their sum. The formula is \(A = \frac{1}{2}ab\sin C\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>138333729</deployedseed>
    <deployedseed>1960974099</deployedseed>
    <deployedseed>189410767</deployedseed>
    <deployedseed>292073592</deployedseed>
    <deployedseed>1535652151</deployedseed>
    <deployedseed>1340031592</deployedseed>
    <deployedseed>1164727723</deployedseed>
    <deployedseed>653489220</deployedseed>
    <deployedseed>1363991053</deployedseed>
    <deployedseed>389461981</deployedseed>
    <deployedseed>1187477765</deployedseed>
    <deployedseed>1136876649</deployedseed>
    <deployedseed>1618479612</deployedseed>
    <deployedseed>504997224</deployedseed>
    <deployedseed>29199162</deployedseed>
    <deployedseed>1243497406</deployedseed>
    <deployedseed>269898528</deployedseed>
    <deployedseed>1909113756</deployedseed>
    <deployedseed>1896503056</deployedseed>
    <deployedseed>1764859696</deployedseed>
    <deployedseed>470507885</deployedseed>
    <deployedseed>719387673</deployedseed>
    <deployedseed>644070903</deployedseed>
    <deployedseed>956129129</deployedseed>
    <deployedseed>1161241743</deployedseed>
    <deployedseed>1263330346</deployedseed>
    <deployedseed>302932469</deployedseed>
    <deployedseed>16792855</deployedseed>
    <deployedseed>406983824</deployedseed>
    <deployedseed>123873359</deployedseed>
    <deployedseed>2137707579</deployedseed>
    <deployedseed>1089246460</deployedseed>
    <deployedseed>1123122663</deployedseed>
    <deployedseed>732514605</deployedseed>
    <deployedseed>1191199542</deployedseed>
    <deployedseed>1788778986</deployedseed>
    <deployedseed>802730603</deployedseed>
    <deployedseed>1128786583</deployedseed>
    <deployedseed>191866619</deployedseed>
    <deployedseed>891032818</deployedseed>
    <deployedseed>315799618</deployedseed>
    <deployedseed>1846603462</deployedseed>
    <deployedseed>2096307374</deployedseed>
    <deployedseed>258284291</deployedseed>
    <deployedseed>1790107583</deployedseed>
    <deployedseed>878459105</deployedseed>
    <deployedseed>574192321</deployedseed>
    <deployedseed>673521610</deployedseed>
    <deployedseed>1267361290</deployedseed>
    <deployedseed>415453851</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
