<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating sines,cosines and tangents of acute angles using a calculator</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Using a calculator, calculate the sine, cosine, and tangent of {@angle@}째, giving your answers correct to 4 decimal places.</p>
<p>(i). sin \({@angle@}째=\)[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>(ii). cos \({@angle@}째=\)[[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>(iii). tan \({@angle@}째=\)[[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Random angle between 5 and 90 with step 0.003 */
angle : 5 + 0.02 * rand(floor((90-5)/0.02) + 1);

/* Sine value */
sine_value : sin(angle * %pi/180);
sine_value_decimal : float(sine_value);
ta:decimalplaces(sine_value_decimal ,4);

/* Cosine value */
cosine_value : cos(angle * %pi/180);
cosine_value_decimal : float(cosine_value);
ta_2:ta_1:decimalplaces(cosine_value_decimal,4);

/* Tangent value */
tangent_value : tan(angle * %pi/180);
tangent_value_decimal : float(tangent_value);
ta_3:ta_1:decimalplaces(tangent_value_decimal,4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Selected angle={@angle@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta_2</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>3.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta_3</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>780502075</deployedseed>
    <deployedseed>1123984529</deployedseed>
    <deployedseed>417763661</deployedseed>
    <deployedseed>1743653370</deployedseed>
    <deployedseed>1764833450</deployedseed>
    <deployedseed>128053682</deployedseed>
    <deployedseed>1949896930</deployedseed>
    <deployedseed>1251948106</deployedseed>
    <deployedseed>402220183</deployedseed>
    <deployedseed>667663651</deployedseed>
    <deployedseed>32358042</deployedseed>
    <deployedseed>1133661192</deployedseed>
    <deployedseed>1632699816</deployedseed>
    <deployedseed>1859848413</deployedseed>
    <deployedseed>818950146</deployedseed>
    <deployedseed>1987468147</deployedseed>
    <deployedseed>1716541277</deployedseed>
    <deployedseed>2090086938</deployedseed>
    <deployedseed>84875291</deployedseed>
    <deployedseed>270385381</deployedseed>
    <deployedseed>1603380672</deployedseed>
    <deployedseed>1475406943</deployedseed>
    <deployedseed>709155714</deployedseed>
    <deployedseed>1200081430</deployedseed>
    <deployedseed>433444913</deployedseed>
    <deployedseed>1994934809</deployedseed>
    <deployedseed>1028855330</deployedseed>
    <deployedseed>1266323299</deployedseed>
    <deployedseed>1201193688</deployedseed>
    <deployedseed>762909789</deployedseed>
    <deployedseed>1845597738</deployedseed>
    <deployedseed>503206755</deployedseed>
    <deployedseed>773373561</deployedseed>
    <deployedseed>1438421676</deployedseed>
    <deployedseed>826072499</deployedseed>
    <deployedseed>1367815519</deployedseed>
    <deployedseed>673370590</deployedseed>
    <deployedseed>2139505666</deployedseed>
    <deployedseed>1601363967</deployedseed>
    <deployedseed>1098617565</deployedseed>
    <deployedseed>1034233269</deployedseed>
    <deployedseed>213783341</deployedseed>
    <deployedseed>1691342400</deployedseed>
    <deployedseed>1082859976</deployedseed>
    <deployedseed>1029851733</deployedseed>
    <deployedseed>1223172630</deployedseed>
    <deployedseed>1148775549</deployedseed>
    <deployedseed>1417818247</deployedseed>
    <deployedseed>69131106</deployedseed>
    <deployedseed>1703185419</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
