<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding mean, median and modal class for grouped data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p> A researcher collects data on daily rainfall (in mm) over a month and organises it into \(10\) equal class intervals as shown below.</p>
<table border="1" cellpadding="5" cellspacing="0">
<tr><th>Rainfall (mm)</th><th>Frequency</th></tr>
<tr><td>0 – 9</td><td>{@freq[1]@}</td></tr>
<tr><td>10 – 19</td><td>{@freq[2]@}</td></tr>
<tr><td>20 – 29</td><td>{@freq[3]@}</td></tr>
<tr><td>30 – 39</td><td>{@freq[4]@}</td></tr>
<tr><td>40 – 49</td><td>{@freq[5]@}</td></tr>
<tr><td>50 – 59</td><td>{@freq[6]@}</td></tr>
<tr><td>60 – 69</td><td>{@freq[7]@}</td></tr>
<tr><td>70 – 79</td><td>{@freq[8]@}</td></tr>
<tr><td>80 – 89</td><td>{@freq[9]@}</td></tr>
<tr><td>90 – 100</td><td>{@freq[10]@}</td></tr>
</table>

<p>
<p>Analyze the rainfall distribution and determine the central tendency measures listed below.</p></p>
<p>
a) Modal class: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]<br>
b) Mean rainfall (mm): [[input:ans2]] [[validation:ans2]][[feedback:prt2]]<br>
c) Median rainfall (mm): [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
</p>
<p><em>Enter the modal class as an interval in the form <code>[a,b]</code>. Round numerical answers to 2 decimal places.</em></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked Solution</h4>
<p>To calculate the statistics, we first need to complete the frequency distribution table with the Midpoint (\(x\)), Cumulative Frequency (\(CF\)), and the product of frequency and midpoint (\(fx\)).</p>

<table border="1">
    <thead style="background-color: #f2f2f2;">
        <tr>
            <th>Class Interval</th>
            <th>Midpoint (\(x\))</th>
            <th>Frequency (\(f\))</th>
            <th>Cumulative Frequency (\(CF\))</th>
            <th>\(fx\)</th>
        </tr>
    </thead>
    <tbody>
        <tr>
            <td>{@classes[1][1]@} – {@classes[1][2]@}</td>
            <td>{@midpoints[1]@}</td>
            <td>{@freq[1]@}</td>
            <td>{@cumfreq[1]@}</td>
            <td>{@midpoints[1]*freq[1]@}</td>
        </tr>
        <tr>
            <td>{@classes[2][1]@} – {@classes[2][2]@}</td>
            <td>{@midpoints[2]@}</td>
            <td>{@freq[2]@}</td>
            <td>{@cumfreq[2]@}</td>
            <td>{@midpoints[2]*freq[2]@}</td>
        </tr>
        <tr>
            <td>{@classes[3][1]@} – {@classes[3][2]@}</td>
            <td>{@midpoints[3]@}</td>
            <td>{@freq[3]@}</td>
            <td>{@cumfreq[3]@}</td>
            <td>{@midpoints[3]*freq[3]@}</td>
        </tr>
        <tr>
            <td>{@classes[4][1]@} – {@classes[4][2]@}</td>
            <td>{@midpoints[4]@}</td>
            <td>{@freq[4]@}</td>
            <td>{@cumfreq[4]@}</td>
            <td>{@midpoints[4]*freq[4]@}</td>
        </tr>
        <tr>
            <td>{@classes[5][1]@} – {@classes[5][2]@}</td>
            <td>{@midpoints[5]@}</td>
            <td>{@freq[5]@}</td>
            <td>{@cumfreq[5]@}</td>
            <td>{@midpoints[5]*freq[5]@}</td>
        </tr>
        <tr>
            <td>{@classes[6][1]@} – {@classes[6][2]@}</td>
            <td>{@midpoints[6]@}</td>
            <td>{@freq[6]@}</td>
            <td>{@cumfreq[6]@}</td>
            <td>{@midpoints[6]*freq[6]@}</td>
        </tr>
        <tr>
            <td>{@classes[7][1]@} – {@classes[7][2]@}</td>
            <td>{@midpoints[7]@}</td>
            <td>{@freq[7]@}</td>
            <td>{@cumfreq[7]@}</td>
            <td>{@midpoints[7]*freq[7]@}</td>
        </tr>
        <tr>
            <td>{@classes[8][1]@} – {@classes[8][2]@}</td>
            <td>{@midpoints[8]@}</td>
            <td>{@freq[8]@}</td>
            <td>{@cumfreq[8]@}</td>
            <td>{@midpoints[8]*freq[8]@}</td>
        </tr>
        <tr>
            <td>{@classes[9][1]@} – {@classes[9][2]@}</td>
            <td>{@midpoints[9]@}</td>
            <td>{@freq[9]@}</td>
            <td>{@cumfreq[9]@}</td>
            <td>{@midpoints[9]*freq[9]@}</td>
        </tr>
        <tr>
            <td>{@classes[10][1]@} – {@classes[10][2]@}</td>
            <td>{@midpoints[10]@}</td>
            <td>{@freq[10]@}</td>
            <td>{@cumfreq[10]@}</td>
            <td>{@midpoints[10]*freq[10]@}</td>
        </tr>
        <tr style="font-weight: bold; background-color: #f9f9f9;">
            <td colspan="2">Total (\(\Sigma\)):</td>
            <td>\({@N@}\)</td>
            <td></td>
            <td>\({@sum_fx@}\)</td>
        </tr>
    </tbody>
</table>

<br>
<p><strong>a) Modal Class</strong></p>
<p>The modal class is the class interval with the highest frequency.</p>
<p>From the table, the highest frequency is <strong>{@maxf@}</strong>.</p>
<p>Therefore, the modal class is <strong>{@modal_class[1]@} – {@modal_class[2]@}</strong>.</p>

<br>
<p><strong>b) Median Rainfall</strong></p>
<p>First, determine the median position:</p>
\( \text{Position } = \frac{N}{2} = \frac{@N@}{2} = {@median_pos@} \)
<p>Looking at the <em>Cumulative Frequency (CF)</em> column, the value just before {@median_pos@} is {@cf_before@}, and the next class reaches {@if median_index=1 then cumfreq[1] else cumfreq[median_index]@}.</p>
<p>Thus, the Median Class is <strong>{@L@} – {@classes[median_index][2]@}</strong>.</p>

<p>Now, apply the median formula:</p>
\( \text{Median} = L + \left( \frac{\frac{N}{2} - CF}{f} \right) \times c \)
<p>Where:</p>
<ul>
    <li>\(L = {@L@}\) (Lower boundary of the median class)</li>
    <li>\(\frac{N}{2} = {@median_pos@}\)</li>
    <li>\(CF = {@cf_before@}\) (Cumulative frequency before the median class)</li>
    <li>\(f = {@f_med@}\) (Frequency of the median class)</li>
    <li>\(c = {@h@}\) (Class width)</li>
</ul>

\( \text{Median} = {@L@} + \left( \frac{{@median_pos@} - {@cf_before@}}{@f_med@} \right) \times {@h@} \)
\( \text{Median} = {@L@} + \left( \frac{@median_pos - cf_before@}{@f_med@} \right) \times {@h@} \)
\( \text{Median} \approx \mathbf{@decimalplaces(median_rain, 2)@} \text{ mm} \)

<br>
<p><strong>c) Mean Rainfall</strong></p>
<p>The mean is calculated using the formula:</p>
\( \bar{x} = \frac{\Sigma fx}{\Sigma f} \)

<p>From the table totals:</p>
<ul>
    <li>\(\Sigma fx = {@sum_fx@}\)</li>
    <li>\(\Sigma f = {@N@}\)</li>
</ul>

\( \bar{x} = \frac{@sum_fx@}{@N@} \)
\( \bar{x} \approx \mathbf{@decimalplaces(mean_rain, 2)@} \text{ mm} \)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fixed class intervals */
classes: [
 [0,9], [10,19], [20,29], [30,39], [40,49],
 [50,59], [60,69], [70,79], [80,89], [90,100]
];

/* Randomised frequencies */
freq: [rand(3)+1,rand(4)+2,rand(5)+3,rand(6)+4,rand(7)+5,rand(8)+6,rand(7)+5,rand(6)+4,rand(5)+3,rand(4)+2];

/* Basic quantities */
k: length(freq);
N: apply("+", freq);

/* Class midpoints */
midpoints: map(lambda([c], (float(c[1]+c[2])/2)), classes);

/* --- NEW: Calculate fx column explicitly for the table and sum --- */
/* Multiply Frequency * Midpoint element-wise */
fx_list: map("*", freq, midpoints);

/* Sum of fx */
sum_fx: apply("+", fx_list);

/* Mean rainfall */
mean_rain: float(sum_fx / N);

/* Modal class */
maxf: lmax(freq);
modal_index: 1;
for j:1 thru length(freq) do
  if freq[j] = maxf then (
    modal_index: j,
    return()
);
modal_class: classes[modal_index];

/* Cumulative frequencies */
cumfreq: [];
running: 0;
for j:1 thru k do (
  running: running + freq[j],
  cumfreq: endcons(running, cumfreq)
);

/* Median class logic */
median_pos: N/2;
median_index: 1;

for j:1 thru k do
  if cumfreq[j] >= median_pos then (
    median_index: j,
    return()
  );

/* Median formula quantities */
L: classes[median_index][1];
cf_before: if median_index=1 then 0 else cumfreq[median_index-1];
f_med: freq[median_index];
h: classes[median_index][2] - classes[median_index][1] + 1;

/* Median rainfall */
median_rain: float(L + ((median_pos - cf_before)/f_med)*h);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Modal class : \({@modal_class@}\)</p>
<p>Mean rain : \({@mean_rain@}\)</p>
<p>Median rain :\({@median_rain@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description"></p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>modal_class</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>mean_rain</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>median_rain</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>modal_class</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>mean_rain</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>median_rain</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
