<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Classifying numbers as even or odd</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Classify the following numbers as even or odd.</p>

<ol type="a">
<li>{@nums[1]@}: [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</li>
<p></p>
<li>{@nums[2]@}: [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</li>
<p></p>
<li>{@nums[3]@}: [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</li>
<p></p>
<li>{@nums[4]@}: [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</li>
<p></p>
<li>{@nums[5]@}: [[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</li>
<p></p>
<li>{@nums[6]@}: [[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</li>
<p></p>
<li>{@nums[7]@}: [[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</li>
<p></p>
<li>{@nums[8]@}: [[input:ans8]] [[validation:ans8]] [[feedback:prt8]]</li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
A number is <strong>even</strong> if it is divisible by 2 with no remainder.<br>
A number is <strong>odd</strong> if it is not divisible by 2, meaning that when the number is divided by 2, the remainder is 1.
</p>

<table border="1" cellpadding="5" style="border-collapse: collapse;">
  <tr>
    <th>Number</th>
    <th>Division by 2</th>
    <th>Classification</th>
  </tr>

  <tr>
    <td>{@nums[1]@}</td>
    <td>{@nums[1]@} ÷ 2 = {@floor(nums[1]/2)@}, remainder {@mod(nums[1],2)@}</td>
    <td><strong>{@t_ans1@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[2]@}</td>
    <td>{@nums[2]@} ÷ 2 = {@floor(nums[2]/2)@}, remainder {@mod(nums[2],2)@}</td>
    <td><strong>{@t_ans2@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[3]@}</td>
    <td>{@nums[3]@} ÷ 2 = {@floor(nums[3]/2)@}, remainder {@mod(nums[3],2)@}</td>
    <td><strong>{@t_ans3@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[4]@}</td>
    <td>{@nums[4]@} ÷ 2 = {@floor(nums[4]/2)@}, remainder {@mod(nums[4],2)@}</td>
    <td><strong>{@t_ans4@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[5]@}</td>
    <td>{@nums[5]@} ÷ 2 = {@floor(nums[5]/2)@}, remainder {@mod(nums[5],2)@}</td>
    <td><strong>{@t_ans5@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[6]@}</td>
    <td>{@nums[6]@} ÷ 2 = {@floor(nums[6]/2)@}, remainder {@mod(nums[6],2)@}</td>
    <td><strong>{@t_ans6@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[7]@}</td>
    <td>{@nums[7]@} ÷ 2 = {@floor(nums[7]/2)@}, remainder {@mod(nums[7],2)@}</td>
    <td><strong>{@t_ans7@}</strong></td>
  </tr>

  <tr>
    <td>{@nums[8]@}</td>
    <td>{@nums[8]@} ÷ 2 = {@floor(nums[8]/2)@}, remainder {@mod(nums[8],2)@}</td>
    <td><strong>{@t_ans8@}</strong></td>
  </tr>
</table>
<p></p>
<p><strong>Recall:</strong></p>
<ul>
  <li>A number is <strong>even</strong> if its last digit is 0, 2, 4, 6, or 8.</li>
  <li>A number is <strong>odd</strong> if its last digit is 1, 3, 5, 7, or 9.</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Generate 8 random integers */
nums : makelist(rand(30000)+15000, i, 1, 8);

/* Create teacher answers and dropdown matrices */
m_ans1: [
    ["Even", is(mod(nums[1], 2) = 0)],
    ["Odd", is(mod(nums[1], 2) = 1)]
];
t_ans1: if mod(nums[1], 2) = 0 then "Even" else "Odd";

m_ans2: [
    ["Even", is(mod(nums[2], 2) = 0)],
    ["Odd", is(mod(nums[2], 2) = 1)]
];
t_ans2: if mod(nums[2], 2) = 0 then "Even" else "Odd";

m_ans3: [
    ["Even", is(mod(nums[3], 2) = 0)],
    ["Odd", is(mod(nums[3], 2) = 1)]
];
t_ans3: if mod(nums[3], 2) = 0 then "Even" else "Odd";

m_ans4: [
    ["Even", is(mod(nums[4], 2) = 0)],
    ["Odd", is(mod(nums[4], 2) = 1)]
];
t_ans4: if mod(nums[4], 2) = 0 then "Even" else "Odd";

m_ans5: [
    ["Even", is(mod(nums[5], 2) = 0)],
    ["Odd", is(mod(nums[5], 2) = 1)]
];
t_ans5: if mod(nums[5], 2) = 0 then "Even" else "Odd";

m_ans6: [
    ["Even", is(mod(nums[6], 2) = 0)],
    ["Odd", is(mod(nums[6], 2) = 1)]
];
t_ans6: if mod(nums[6], 2) = 0 then "Even" else "Odd";

m_ans7: [
    ["Even", is(mod(nums[7], 2) = 0)],
    ["Odd", is(mod(nums[7], 2) = 1)]
];
t_ans7: if mod(nums[7], 2) = 0 then "Even" else "Odd";

m_ans8: [
    ["Even", is(mod(nums[8], 2) = 0)],
    ["Odd", is(mod(nums[8], 2) = 1)]
];
t_ans8: if mod(nums[8], 2) = 0 then "Even" else "Odd";

/*mistakes*/
/* Wrong answer feedback for each number */
fb_wrong1: if t_ans1 = "Even" then 
    sconcat("The number ", nums[1], " is EVEN. ") 
else 
    sconcat("The number ", nums[1], " is ODD.");

fb_wrong2: if t_ans2 = "Even" then 
    sconcat("The. ", nums[2], " is EVEN.") 
else 
    sconcat( nums[2], " is classified as ODD.");

fb_wrong3: if t_ans3 = "Even" then 
    sconcat( nums[3], " is EVEN. Look at the last digit ", mod(nums[3],10), ": 0,2,4,6,8 are even digits.") 
else 
    sconcat("Incorrect. ", nums[3], " is ODD. Look at the last digit ", mod(nums[3],10), ": 1,3,5,7,9 are odd digits.");

fb_wrong4: if t_ans4 = "Even" then 
    sconcat("Incorrect. ", nums[4], " is EVEN. Look at the last digit ", mod(nums[4],10), ": 0,2,4,6,8 are even digits.") 
else 
    sconcat("Incorrect. ", nums[4], " is ODD. Look at the last digit ", mod(nums[4],10), ": 1,3,5,7,9 are odd digits.");

fb_wrong5: if t_ans5 = "Even" then 
    sconcat("Incorrect. ", nums[5], " is EVEN. Look at the last digit ", mod(nums[5],10), ": 0,2,4,6,8 are even digits.") 
else 
    sconcat("Incorrect. ", nums[5], " is ODD. Look at the last digit ", mod(nums[5],10), ": 1,3,5,7,9 are odd digits.");

fb_wrong6: if t_ans6 = "Even" then 
    sconcat("Incorrect. ", nums[6], " is EVEN. Look at the last digit ", mod(nums[6],10), ": 0,2,4,6,8 are even digits.") 
else 
    sconcat("Incorrect. ", nums[6], " is ODD. Look at the last digit ", mod(nums[6],10), ": 1,3,5,7,9 are odd digits.");

fb_wrong7: if t_ans7 = "Even" then 
    sconcat("Incorrect. ", nums[7], " is EVEN. Look at the last digit ", mod(nums[7],10), ": 0,2,4,6,8 are even digits.") 
else 
    sconcat("Incorrect. ", nums[7], " is ODD. Look at the last digit ", mod(nums[7],10), ": 1,3,5,7,9 are odd digits.");

fb_wrong8: if t_ans8 = "Even" then 
    sconcat("Incorrect. ", nums[8], " is EVEN. Look at the last digit ", mod(nums[8],10), ": 0,2,4,6,8 are even digits.") 
else 
    sconcat("Incorrect. ", nums[8], " is ODD. Look at the last digit ", mod(nums[8],10), ": 1,3,5,7,9 are odd digits.");
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>
Variant: 
{n1:{@nums[1]@}-{@t_ans1@}, 
n2:{@nums[2]@}-{@t_ans2@}, 
n3:{@nums[3]@}-{@t_ans3@}, 
n4:{@nums[4]@}-{@t_ans4@}, 
n5:{@nums[5]@}-{@t_ans5@}, 
n6:{@nums[6]@}-{@t_ans6@}, 
n7:{@nums[7]@}-{@t_ans7@}, 
n8:{@nums[8]@}-{@t_ans8@}}
</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>m_ans1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>m_ans2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>dropdown</type>
      <tans>m_ans3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>dropdown</type>
      <tans>m_ans4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>dropdown</type>
      <tans>m_ans5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>dropdown</type>
      <tans>m_ans6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>dropdown</type>
      <tans>m_ans7</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>dropdown</type>
      <tans>m_ans8</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>t_ans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong1@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>t_ans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong2@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>t_ans3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong3@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>t_ans4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong4@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>t_ans5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong5@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>t_ans6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong6@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>t_ans7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong7@}</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct classification of the number</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>t_ans8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-1-F</falseanswernote>
        <falsefeedback format="html">
          <text>{@fb_wrong8@}</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1542643273</deployedseed>
    <deployedseed>18109183</deployedseed>
    <deployedseed>1251031222</deployedseed>
    <deployedseed>502057859</deployedseed>
    <deployedseed>106587397</deployedseed>
    <deployedseed>1834057214</deployedseed>
    <deployedseed>797236528</deployedseed>
    <deployedseed>1171138601</deployedseed>
    <deployedseed>1358801974</deployedseed>
    <deployedseed>891169076</deployedseed>
    <deployedseed>1185459200</deployedseed>
    <deployedseed>659566141</deployedseed>
    <deployedseed>1301014304</deployedseed>
    <deployedseed>1816703445</deployedseed>
    <deployedseed>1184903523</deployedseed>
    <deployedseed>1065718095</deployedseed>
    <deployedseed>1652660946</deployedseed>
    <deployedseed>1983450235</deployedseed>
    <deployedseed>1169155453</deployedseed>
    <deployedseed>901622370</deployedseed>
    <deployedseed>1769625220</deployedseed>
    <deployedseed>1699302091</deployedseed>
    <deployedseed>1498307240</deployedseed>
    <deployedseed>983163326</deployedseed>
    <deployedseed>305631209</deployedseed>
    <deployedseed>9722638</deployedseed>
    <deployedseed>203370993</deployedseed>
    <deployedseed>1640974126</deployedseed>
    <deployedseed>940653195</deployedseed>
    <deployedseed>1853799665</deployedseed>
    <deployedseed>2135697574</deployedseed>
    <deployedseed>125086341</deployedseed>
    <deployedseed>1805214387</deployedseed>
    <deployedseed>1517292403</deployedseed>
    <deployedseed>1577092096</deployedseed>
    <deployedseed>2064764841</deployedseed>
    <deployedseed>431222451</deployedseed>
    <deployedseed>1184436726</deployedseed>
    <deployedseed>927120686</deployedseed>
    <deployedseed>1627094703</deployedseed>
    <deployedseed>623726343</deployedseed>
    <deployedseed>1685217621</deployedseed>
    <deployedseed>1180576484</deployedseed>
    <deployedseed>946522338</deployedseed>
    <deployedseed>1927072439</deployedseed>
    <deployedseed>1591809365</deployedseed>
    <deployedseed>409050484</deployedseed>
    <deployedseed>996724230</deployedseed>
    <deployedseed>2135597654</deployedseed>
    <deployedseed>419536078</deployedseed>
    <deployedseed>1686099961</deployedseed>
    <deployedseed>704779609</deployedseed>
    <deployedseed>1918034359</deployedseed>
    <deployedseed>1787800133</deployedseed>
    <deployedseed>1935473153</deployedseed>
    <deployedseed>111136966</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(m_ans1))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(m_ans2))</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>first(mcq_correct(m_ans3))</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>first(mcq_correct(m_ans4))</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>first(mcq_correct(m_ans5))</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>first(mcq_correct(m_ans6))</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>first(mcq_correct(m_ans7))</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>first(mcq_correct(m_ans8))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt6</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt6-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt7</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt7-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt8</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt8-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
