<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Applying negative index laws to simplify expressions</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Simplify the expression below using the law of negative indices:</p>
<p>\[ {@coef@}{@var@}^{-{@n@}} \]</p>
<p>Give your answer without negative indices.</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The negative index law states:</p>
<p>\[ a^{-n} = \frac{1}{a^n}. \]</p>
<p>Applying this to the expression:</p>
<p>\[ {@coef@}{@var@}^{-{@n@}} = {@coef@} \times \frac{1}{ {@var@}^{@n@} }. \]</p>
<p>So the simplified form is:</p>
<p>\[ \frac{ {@coef@} }{ {@var@}^{@n@} }. \]</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Choose a variable */
var: rand([x,y,a,b,p,q,r]);

/* Random positive coefficient */
coef: rand_with_step(2,9,1);

/* Random power for negative index */
n: rand([1,2,3]);

/* Expression: coef * var^(-n) */
expr: coef * var^(-n);

/* Correct answer: coef/(var^n) */
ans: coef/(var^n);
ans2: coef*var^n;

/* Wrong answers */
m1: coef * var^n;       /* flips rule */
m2: -coef/(var^n);      /* introduces sign error */
m3: coef * -n;          /* confuses exponent with multiplication */
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ans@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question requires the learner to apply the negative index law to rewrite expressions with negative powers in fractional form.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p><i class="fa fa-check"></i> Correct answer, well done.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p><i class="fa fa-adjust"></i> Your answer is partially correct.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p><i class="fa fa-times"></i> Incorrect answer.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ans</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>&lt;p&gt;You assumed the negative sign. If an index is negative, rewrite the expression using the reciprocal so that all indices become positive. Check your working and ensure your final answer is fully simplified.&lt;/p&gt; </p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>116511134</deployedseed>
    <deployedseed>793018386</deployedseed>
    <deployedseed>970386283</deployedseed>
    <deployedseed>557871768</deployedseed>
    <deployedseed>369360253</deployedseed>
    <deployedseed>1614505220</deployedseed>
    <deployedseed>1445570925</deployedseed>
    <deployedseed>1432743297</deployedseed>
    <deployedseed>1862326723</deployedseed>
    <deployedseed>155041232</deployedseed>
    <deployedseed>351880555</deployedseed>
    <deployedseed>1312442142</deployedseed>
    <deployedseed>1211915773</deployedseed>
    <deployedseed>2119077635</deployedseed>
    <deployedseed>797620100</deployedseed>
    <deployedseed>510244332</deployedseed>
    <deployedseed>1587191953</deployedseed>
    <deployedseed>273228055</deployedseed>
    <deployedseed>959281729</deployedseed>
    <deployedseed>1006467059</deployedseed>
    <deployedseed>223278109</deployedseed>
    <deployedseed>686254375</deployedseed>
    <deployedseed>1633404680</deployedseed>
    <deployedseed>1332610717</deployedseed>
    <deployedseed>915304489</deployedseed>
    <deployedseed>650671208</deployedseed>
    <deployedseed>1455321712</deployedseed>
    <deployedseed>1788327262</deployedseed>
    <deployedseed>411701322</deployedseed>
    <deployedseed>1622299094</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ans</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
