<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Congruence of Figures Using Transformations</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<h3 style="text-align:center;">
Consider the two figures shown on the coordinate grid.
</h3>

<div style="text-align:center;">

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10,8,10,-8],
    axis:true,
    showCopyright:false
});

/* Original triangle */
var A = board.create('point', [{#Ax#},{#Ay#}], {name:'A', size:3});
var B = board.create('point', [{#Bx#},{#By#}], {name:'B', size:3});
var C = board.create('point', [{#Cx#},{#Cy#}], {name:'C', size:3});
board.create('polygon',[A,B,C],{
    fillColor:'#a5d6a7',
    fillOpacity:0.6,
    borders:{strokeWidth:2}
});

/* Image triangle */
var S = board.create('point', [{#Sx#},{#Sy#}], {name:'S', size:3});
var T = board.create('point', [{#Tx#},{#Ty#}], {name:'T', size:3});
var V = board.create('point', [{#Vx#},{#Vy#}], {name:'V', size:3});
board.create('polygon',[S,T,V],{
    fillColor:'#90caf9',
    fillOpacity:0.6,
    borders:{strokeWidth:2}
});
[[/jsxgraph]]

</div>

<p>
<b>a)</b> Are the two figures congruent?
<br>
[[input:ans1]]
[[validation:ans1]]
[[feedback:prt1]]
</p>

<p>
<b>b)</b> If they are congruent, describe the transformation that maps the first figure onto the second.
<br>
[[input:ans2]]
[[validation:ans2]]
[[feedback:prt2]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p>
The two figures are <b>congruent</b> because all corresponding side lengths
and angles are equal.
</p>

<p>
One way to map the first figure onto the second is:
</p>

<ul>
<li>Apply {@rot_text@}</li>
<li>Then translate the figure by the vector
\(\langle {@tx@}, {@ty@} \rangle\)</li>
</ul>

<p>
Since only rigid motions (translations and rotations) are used,
the figures are congruent.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Base triangle */
Ax: -3;
Ay: 1;

Bx: -1;
By: 4;

Cx: 2;
Cy: 2;

/* Random translation */
tx: rand(4) + 3;
ty: rand(4) - 2;

/* Mode selector
   0 = translation
   1 = rotation + translation
   2 = not congruent
*/
mode: rand(3);

/* Image triangle */
if mode=0 then (
    Sx: Ax+tx,  Sy: Ay+ty,
    Tx: Bx+tx,  Ty: By+ty,
    Vx: Cx+tx,  Vy: Cy+ty
)
elseif mode=1 then (
    Sx: -Ax+tx,  Sy: -Ay+ty,
    Tx: -Bx+tx,  Ty: -By+ty,
    Vx: -Cx+tx,  Vy: -Cy+ty
)
else (
    /* Not congruent */
    Sx: Ax+tx,  Sy: 2*Ay+ty,
    Tx: Bx+tx,  Ty: 2*By+ty,
    Vx: Cx+tx,  Vy: 2*Cy+ty
);

/* Congruence flag */
is_cong: if mode=2 then false else true;

/* Part (a) permutation */
ta_cong_opts:
    random_permutation([
        ["Yes", is_cong],
        ["No", not is_cong]
    ]);

/* Part (b) correct transformation */
ta_trans:
    if mode=0 then "translation"
    elseif mode=1 then "rotation and translation"
    else "not possible";

/* Part (b) permutation */
ta_trans_opts:
    random_permutation([
        ["Translation", ta_trans="translation"],
        ["Rotation and translation", ta_trans="rotation and translation"],
        ["Not possible", ta_trans="not possible"]
    ]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta_cong@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_cong_opts</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_trans_opts</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Convert yes/no text answer into numeric value */
ans1n:
  if is(ans1 = "yes") or is(ans1 = "Yes") or is(ans1 = "YES") then 1
  elseif is(ans1 = "no") or is(ans1 = "No") or is(ans1 = "NO") then 0
  else -1;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_cong</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_trans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta_cong</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_trans</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
