<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Congruence of Figures Using Transformations</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<h3 style="text-align:center;">
Consider the two figures shown on the coordinate grid.
</h3>

<div style="text-align:center;">

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10,8,10,-8],
    axis:true,
    showCopyright:false
});

/* Original triangle */
var A = board.create('point', [{#Ax#},{#Ay#}], {name:'A', size:3});
var B = board.create('point', [{#Bx#},{#By#}], {name:'B', size:3});
var C = board.create('point', [{#Cx#},{#Cy#}], {name:'C', size:3});
board.create('polygon',[A,B,C],{
    fillColor:'#a5d6a7',
    fillOpacity:0.6,
    borders:{strokeWidth:2}
});

/* Image triangle */
var S = board.create('point', [{#Sx#},{#Sy#}], {name:'S', size:3});
var T = board.create('point', [{#Tx#},{#Ty#}], {name:'T', size:3});
var V = board.create('point', [{#Vx#},{#Vy#}], {name:'V', size:3});
board.create('polygon',[S,T,V],{
    fillColor:'#90caf9',
    fillOpacity:0.6,
    borders:{strokeWidth:2}
});
[[/jsxgraph]]

</div>

<p>
<b>a)</b> Are the two figures congruent?
<br>
[[input:ans1]]
[[validation:ans1]]
[[feedback:prt1]]
</p>

<p>
<b>b)</b> If they are congruent, describe the transformation that maps the first figure onto the second.
<br>
[[input:ans2]]
[[validation:ans2]]
[[feedback:prt2]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Two figures in the plane are <strong>congruent</strong> if one can be obtained from the other by a <strong>rigid transformation</strong> (also called an isometry).
</p>

<p>
Rigid transformations preserve:
</p>
<ul>
<li>All side lengths</li>
<li>All angles</li>
<li>The overall shape and size of the figure</li>
</ul>

<p>
The main rigid transformations in the plane are:
</p>
<ul>
<li><strong>Translation</strong> – shifting every point by the same vector</li>
<li><strong>Rotation</strong> – turning the figure around a fixed point</li>
<li><strong>Reflection</strong> – flipping the figure across a line</li>
</ul>

<p>
Any combination of these transformations is still rigid.
</p>

<hr>

<p><strong>How to analyze this problem:</strong></p>

<ol>
<li>
Check whether the side lengths of the two triangles are the same.
If distances are preserved, congruence is possible.
</li>

<li>
Determine the type of transformation:
<ul>
<li>If every vertex is shifted by the same vector → it is a <strong>translation</strong>.</li>
<li>If the triangle appears turned (possibly together with a shift) → it is a <strong>rotation and translation</strong>.</li>
<li>If the triangle is stretched (for example, one coordinate is scaled differently) → distances are not preserved and the figures are <strong>not congruent</strong>.</li>
</ul>
</li>
</ol>

<p>
Remember:
</p>

<p style="text-align:center;">
<strong>Congruent figures have exactly the same shape and size.</strong>
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Base triangle */
Ax: -3;
Ay: 1;

Bx: -1;
By: 4;

Cx: 2;
Cy: 2;

/* Random translation */
tx: rand(4) + 3;
ty: rand(4) - 2;

/* Mode selector
   0 = translation
   1 = rotation + translation
   2 = not congruent
*/
mode: rand(3);

/* Image triangle */
if mode=0 then (
    Sx: Ax+tx,  Sy: Ay+ty,
    Tx: Bx+tx,  Ty: By+ty,
    Vx: Cx+tx,  Vy: Cy+ty
)
elseif mode=1 then (
    Sx: -Ax+tx,  Sy: -Ay+ty,
    Tx: -Bx+tx,  Ty: -By+ty,
    Vx: -Cx+tx,  Vy: -Cy+ty
)
else (
    /* Not congruent */
    Sx: Ax+tx,  Sy: 2*Ay+ty,
    Tx: Bx+tx,  Ty: 2*By+ty,
    Vx: Cx+tx,  Vy: 2*Cy+ty
);

/* Congruence flag */
is_cong: if mode=2 then false else true;

/* Part (a) permutation */
ta_cong_opts:
    random_permutation([
        ["Yes", is_cong],
        ["No", not is_cong]
    ]);

/* Part (b) correct transformation */
ta_trans:
    if mode=0 then "translation"
    elseif mode=1 then "rotation and translation"
    else "not possible";

/* Part (b) permutation */
ta_trans_opts:
    random_permutation([
        ["Translation", is(ta_trans="translation")],
        ["Rotation and translation", is(ta_trans="rotation and translation")],
        ["Not possible", is (ta_trans="not possible")]
    ]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta_cong_opts@}
{@ta_trans_opts@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta_cong_opts</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta_trans_opts</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Convert yes/no text answer into numeric value */
ans1n:
  if is(ans1 = "yes") or is(ans1 = "Yes") or is(ans1 = "YES") then 1
  elseif is(ans1 = "no") or is(ans1 = "No") or is(ans1 = "NO") then 0
  else -1;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta_cong_opts)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>2.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>mcq_correct(ta_trans_opts)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>857132351</deployedseed>
    <deployedseed>971388850</deployedseed>
    <deployedseed>1534292840</deployedseed>
    <deployedseed>1601563579</deployedseed>
    <deployedseed>6018392</deployedseed>
    <deployedseed>791449302</deployedseed>
    <deployedseed>1459531543</deployedseed>
    <deployedseed>657126177</deployedseed>
    <deployedseed>938872360</deployedseed>
    <deployedseed>1126133842</deployedseed>
    <deployedseed>198963845</deployedseed>
    <deployedseed>1087731050</deployedseed>
    <deployedseed>756695139</deployedseed>
    <deployedseed>1185236779</deployedseed>
    <deployedseed>330037547</deployedseed>
    <deployedseed>734253068</deployedseed>
    <deployedseed>1766727706</deployedseed>
    <deployedseed>1857425860</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>first(mcq_correct(ta_cong_opts))</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>first(mcq_correct(ta_trans_opts))</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
