<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Representing data using a grouped frequency distribution table</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The following data represents the exam scores of <strong>{@total_students@}</strong> students in a mathematics class:</p>
<p style="background-color: #f9f9f9; padding: 15px; border: 1px solid #ddd; font-family: monospace; font-size: 1.1em; line-height: 1.6;">
    {@data_str@}
</p>
<p>Construct a <strong>grouped frequency distribution table</strong> using the class intervals below.</p>
<table class="table table-bordered table-striped" style="width: 60%; margin-top: 20px;">
    <thead style="background-color: #e0e0e0;">
        <tr><th style="text-align: center;">Score Interval</th><th style="text-align: center;">Frequency</th></tr>
    </thead>
    <tbody>
        <tr>
            <td style="text-align: center;">{@c1_min@} &#8211; {@c1_max@}</td>
            <td style="text-align: center;">[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</td>
        </tr>
        <tr>
            <td style="text-align: center;">{@c2_min@} &#8211; {@c2_max@}</td>
            <td style="text-align: center;">[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</td>
        </tr>
        <tr>
            <td style="text-align: center;">{@c3_min@} &#8211; {@c3_max@}</td>
            <td style="text-align: center;">[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</td>
        </tr>
        <tr>
            <td style="text-align: center;">{@c4_min@} &#8211; {@c4_max@}</td>
            <td style="text-align: center;">[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</td>
        </tr>
        <tr>
            <td style="text-align: center;">{@c5_min@} &#8211; {@c5_max@}</td>
            <td style="text-align: center;">[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</td>
        </tr>
        <tr style="font-weight: bold; background-color: #f2f2f2;">
            <td style="text-align: right; padding-right: 20px;">Total</td>
            <td style="text-align: center;">[[input:ans_total]] [[validation:ans_total]] [[feedback:prt_total]]</td>
        </tr>
    </tbody>
</table>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>
<p>Tally the scores that fall into each range (endpoints included):</p>
<ul>
<li><strong>{@c1_min@} &#8211; {@c1_max@}:</strong> scores in the 40s &rarr; <strong>{@f1@}</strong>.</li>
<li><strong>{@c2_min@} &#8211; {@c2_max@}:</strong> scores in the 50s &rarr; <strong>{@f2@}</strong>.</li>
<li><strong>{@c3_min@} &#8211; {@c3_max@}:</strong> scores in the 60s &rarr; <strong>{@f3@}</strong>.</li>
<li><strong>{@c4_min@} &#8211; {@c4_max@}:</strong> scores in the 70s &rarr; <strong>{@f4@}</strong>.</li>
<li><strong>{@c5_min@} &#8211; {@c5_max@}:</strong> scores in the 80s &rarr; <strong>{@f5@}</strong>.</li>
</ul>
<p><strong>Total:</strong> {@f1@}+{@f2@}+{@f3@}+{@f4@}+{@f5@} = <strong>{@total_students@}</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Fixed class intervals, width 10 */
c1_min:40$ c1_max:49$ c2_min:50$ c2_max:59$ c3_min:60$ c3_max:69$
c4_min:70$ c4_max:79$ c5_min:80$ c5_max:89$
/* Random frequencies (3..8 per group) */
f1:rand(6)+3$ f2:rand(6)+3$ f3:rand(6)+3$ f4:rand(6)+3$ f5:rand(6)+3$
total_students: f1+f2+f3+f4+f5;
/* Generate raw scores matching each frequency, then shuffle */
generate_group(nn,a,b) := makelist(rand(b-a+1)+a, k, 1, nn);
raw_data_list: append(generate_group(f1,c1_min,c1_max),generate_group(f2,c2_min,c2_max),generate_group(f3,c3_min,c3_max),generate_group(f4,c4_min,c4_max),generate_group(f5,c5_min,c5_max));
shuffled_data: random_permutation(raw_data_list);
data_str: sconcat(shuffled_data[1]);
for i:2 thru total_students do data_str: sconcat(data_str, ", ", shuffled_data[i]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>Frequencies: {@f1@}, {@f2@}, {@f3@}, {@f4@}, {@f5@}; total {@total_students@}.</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="reviewer">Ben (STACK audit)</p>
<p id="description">Students construct a grouped frequency distribution table (five width-10 intervals plus a total) by tallying randomised exam scores. Each interval and the total are marked separately (partial credit) with recount hints.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_total</name>
      <type>algebraic</type>
      <tans>total_students</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>f1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>f2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>f3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>f4</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>f5</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt_total</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans_total</sans>
        <tans>total_students</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_total-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_total-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The total should equal the sum of all five interval frequencies (and the count of scores listed).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>f1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Recount the scores in the {@c1_min@}&#8211;{@c1_max@} range (endpoints included).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>f2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Recount the scores in the {@c2_min@}&#8211;{@c2_max@} range (endpoints included).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>f3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Recount the scores in the {@c3_min@}&#8211;{@c3_max@} range (endpoints included).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>f4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Recount the scores in the {@c4_min@}&#8211;{@c4_max@} range (endpoints included).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>f5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Recount the scores in the {@c5_min@}&#8211;{@c5_max@} range (endpoints included).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>690373012</deployedseed>
    <deployedseed>145200781</deployedseed>
    <deployedseed>922014553</deployedseed>
    <deployedseed>377650129</deployedseed>
    <deployedseed>588103477</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All frequencies and total correct.</description>
      <testinput>
        <name>ans_total</name>
        <value>total_students</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>f1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>f2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>f3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>f4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>f5</value>
      </testinput>
      <expected>
        <name>prt_total</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_total-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt5-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>All wrong (one too many): fires each F branch.</description>
      <testinput>
        <name>ans_total</name>
        <value>total_students+1</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>f1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>f2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>f3+1</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>f4+1</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>f5+1</value>
      </testinput>
      <expected>
        <name>prt_total</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt_total-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt5</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt5-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
