<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Comparing earthquake magnitudes using logarithms</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
On a {@day1@} {@time1@}, two earthquakes were recorded in {@region1@} at {@station@}. The first earthquake had magnitude {@m_small@}, and the second had magnitude {@m_large@}.
</p>

<p>
(a) Find the difference in magnitude between the two earthquakes.
</p>

<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>

<p>
(b) Using logarithmic principles, determine how many times more powerful the larger earthquake is compared to the smaller one.
</p>

<p>
[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Step 1: Find the difference in magnitude:
</p>

<p>
\[
{@m_large@} - {@m_small@} = {@diff@}
\]
</p>

<p>
Step 2: Understand the logarithmic scale.
</p>

<p>
Each increase of 1 in magnitude means the intensity becomes 10 times greater.
</p>

<p>
Step 3: Apply the rule:
</p>

<p>
\[
\text{Intensity ratio} = 10^{@diff@}
\]
</p>

<p>
\[
= {@ratio@}
\]
</p>

<p>
Therefore, the larger earthquake is {@ratio@} times more powerful than the smaller one.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Base magnitudes */
m_small: rand(3) + 3;        /* 3–5 */
gap: rand(3) + 1;            /* 1–3 */
m_large: m_small + gap;

/* Difference */
diff: m_large - m_small;

/* Intensity ratio */
ratio: 10^diff;

/* Distractors (core mistakes) */
m1: diff;
m2: 10^(m_small + m_large);
m3: m_large * m_small;
m4: 10^(m_small);
m5: 10^(m_large);

/* Additional random variables (for STACK richness & variation) */
region1: rand(["Nakuru","Turkana","Garissa","Kisumu","Kitale"]);
region2: rand(["Mombasa","Eldoret","Naivasha","Nyeri","Machakos"]);
day1: rand(["Monday","Tuesday","Wednesday","Thursday","Friday"]);
time1: rand(["morning","afternoon","evening"]);
station: rand(["Station A","Station B","Station C","Station D"]);

noise1: rand(5)+1;
noise2: rand(5)+2;
noise3: rand(5)+3;
noise4: rand(5)+4;
noise5: rand(5)+5;
noise6: rand(5)+6;
noise7: rand(5)+7;
noise8: rand(5)+8;
noise9: rand(5)+9;
noise10: rand(5)+10;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ratio@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Logarithms in real-life</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to apply logarithmic concepts to compare earthquake magnitudes and interpret their real-life meaning.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>diff</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ratio</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>diff</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>mistake1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m_small + m_large</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Find the difference, not the sum.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Difference = larger magnitude − smaller magnitude.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ratio</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>used difference only</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>The difference gives the exponent. Raise 10 to that power.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>added magnitudes</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text>Do not add magnitudes. Use their difference.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>used one magnitude</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>m4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text>Use the difference between magnitudes, not a single value.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1828940207</deployedseed>
    <deployedseed>1110301038</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ratio</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
