<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the Surface Area of a cube</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Look at the cube shown below. Each edge of the cube measures {@side@} \(\text{cm}\).</p>

<!DOCTYPE html>
<html>
<head>
  <title>Sector OAB</title>
  <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraphcore.js"></script>
  <link rel="stylesheet" href="https://cdn.jsdelivr.net/npm/jsxgraph/distrib/jsxgraph.css" />
</head>
<body>

    <div id="box" class="jxgbox" style="width: 500px; height: 500px;"></div>
    <script type="text/javascript">
        const board = JXG.JSXGraph.initBoard('box', {
            boundingbox: [-2, 5, 5, -2], showCopyright: false,
            
            axis: false, showNavigation: false, keepAspectRatio:true
        });

        const side = 2.5;

        // Define points
        let A = board.create('point', [0, 0], {visible: false});
        let B = board.create('point', [side, 0], {visible: false});
        let C = board.create('point', [side, side], {visible: false});
        let D = board.create('point', [0, side], {visible: false});

        let A1 = board.create('point', [0.8, 1], {visible: false});
        let B1 = board.create('point', [side + 0.8, 1], {visible: false});
        let C1 = board.create('point', [side + 0.8, side + 1], {visible: false});
        let D1 = board.create('point', [0.8, side + 1], {visible: false});

        // Create edges
        board.create('polygon', [A, B, C, D], {borders: {strokeColor: 'blue'}});
        board.create('polygon', [A1, B1, C1, D1], {borders: {strokeColor: 'blue'}});
        board.create('line', [A, A1], {straightFirst: false, straightLast: false, strokeColor: 'blue'});
        board.create('line', [B, B1], {straightFirst: false, straightLast: false, strokeColor: 'blue'});
        board.create('line', [C, C1], {straightFirst: false, straightLast: false, strokeColor: 'blue'});
        board.create('line', [D, D1], {straightFirst: false, straightLast: false, strokeColor: 'blue'});

        // Label the 3 visible edges with 13 cm
        board.create('text', [1.2, -0.2, "{#side#} cm"]);
        board.create('text', [3.5, 2, "{#side#} cm"]);
        board.create('text', [3, 0.4, "{#side#} cm"]);
    </script>




</body>
</html>
<p>Calculate the total surface area of the cube in:</p>
<p><ol type="a"><li><p>\(\text{Square centimeters} =\)[[input:ans1]]\(\text{cm}^2\) [[feedback:prt1]]</p></li>
<li>\(\text{Square meters}=\)[[input:ans2]]\(\text{m}^2\) [[feedback:prt2]]</li></ol></p>
<p></p>




<p hidden>[[validation:ans1]][[validation:ans2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>Given the side = {@side@} \(\text{cm}\):</p>
<p>Surface Area (SA) of a cube = 6 × (side)² or \(5 l^2\) where \(l\) is the side of the cube.</p>
<p>Substitute the value in the formula of the cube.</p>
<p>\( \begin{align}       \text{SA} &=6 \times {@side@}^2 \\
&= 6 \times  ({@side@} \times  {@side@})\\
&= 6 \times  {@side@} \\
&= {@surface_area_cm2@}\, \text{cm}^2
  \end{align} \)</p>

<p>Convert cm² to m²</p>
<p>Since:</p>
<p> \(1\, \text{m}^2= 10\,000 \, \text{cm^2}\)</p>
<p>Therefore,</p>

<p>\(\begin{align} \text{Surface Area in} \, \text{m}^2  &= \frac{{@surface_area_cm2@}}{10\,000  } \\
&= {@surface_area_m2@} \, \text{m}^2
\end{align}\)</p>
<p>Which can be simplified as, {@Surface_area_m2@} \(\text{m}^2\)</p>


<p>The surface of the cube in</p>
<p><ol type="a"><li>\(\text{Square centimeters}= \){@surface_area_cm2@} \(\text{cm}^2\)</li>
<li>\(\text{Square meters }= \){@surface_area_m2@} \(\text{m}^2\) or {@Surface_area_m2@} \(\text{m}^2\)</li></ol></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025012100</text>
    </stackversion>
    <questionvariables>
      <text>/* Random cube side length between 5 and 20 cm */
side : rand(50) + 5;

/* Surface area of a cube = 6 × side² */
surface_area_cm2 : 6 * side^2;

/* Convert cm² to m² (1 m² = 10,000 cm²) */
surface_area_m2 : surface_area_cm2 / 10000;

Surface_area_m2:float(surface_area_m2);



/*Mistake 1   Forgetting to Square the Side Length*/
Msurface1_cm2 : 6 * side;
Msurface1_m2 :Msurface1_cm2 / 10000;
MSurface1_m2:float(Msurface1_m2);

/*Mistake 2  Incorrect Unit Conversion (cm² to m²) divding by 100 instade of 10 000*/
Msurface2_m2 :surface_area_cm2 / 100;
MSurface2_m2:float(Msurface2_m2);

/*Mistake 3  Using wrong formula . Using SA= 6l^3 instade of  SA= 6l^2*/

/* Surface area of a cube = 6 × side² */
Msurface3_cm2 : 6 * side^3;

/* Convert cm² to m² (1 m² = 10,000 cm²) */
Msurface3_m2 : Msurface3_cm2 / 10000;

MSurface3_m2:float(Msurface3_m2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The side of the cube is;}\) {@side@}<p>
<p>\(\text{Area in cm² is ;}\) {@surface_area_cm2@}<p>
<p>\(\text{Area in m² is ;}\) {@Surface_area_m2@}<p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Henry Onyango</p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">The question provided above is for a cube. The student is required to find the surface area of the cube in the given units, and then convert the area to \(\text{m}^2\). They should recall that the surface area of a cube is given by
\(\text{SA}=6 l^2\)
because the cube is opened, where \(\l) is the side length of the cube.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>surface_area_cm2</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>Surface_area_m2</tans>
      <boxsize>9</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct area of the cube in cm squared.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>surface_area_cm2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Forgetting to Square the Side Length</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Msurface1_cm2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Good try! 😊 However, you multiplied \(6 \times {@side@}\) instead of squaring the side length first. </p>
<p>The correct formula is:</p>
<p>\(\text{Surface Area}=6 l^2\)</p>
<p>\(\textbf{What to Do:}\)</p>
<p>Always square the side before multiplying.</p>
<p>You’re improving—keep it up! 👏</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Using incorect formular that is 6l^3 instade of 6l^2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Msurface3_cm2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Nice effort! 😊 I see you used the formula \(6  \times  l^3\), which actually gives the volume of the cube multiplied by 6.</p>
<p>For surface area, we use:</p>
<p>Surface Area = \(6 \times l^2\) (because there are \(6\) square faces).</p>
<p>\(\textbf{What to do:}\)</p>
<p>Remember:</p>
<p>\(l^2\) is for area (surface),</p>
<p>\(l^3\) is for volume.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct area of the cube in m squared.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>Surface_area_m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Forgetting to Square the Side Length</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>MSurface1_m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You are almost there. Kindly square the side length then for the conversion will be automatically correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description> Incorrect Unit Conversion (cm² to m²) divding by 100 instade of 10 000</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>MSurface2_m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.2</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Good try! 😊 However, for the conversion, you divided by 100 instead of 10,000.</p>
<p>Remember:</p>
<p>\(1\,\text{ m²} = 10\,000 \,\text{cm²}\), so</p>
<p>\(\textbf{What to Do:}\)</p>
<p>When converting to \(\text{m}^2\), divide by \(10\,000\), not \(100\).</p>
<p>You’re improving—keep it up! 👏</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Using incorect formular that is 6l^3 instade of 6l^2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>MSurface3_m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You are almost there. Kindly square the length don't cube it.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1293364333</deployedseed>
    <deployedseed>525120757</deployedseed>
    <deployedseed>649846224</deployedseed>
    <deployedseed>1093109287</deployedseed>
    <deployedseed>1008716881</deployedseed>
    <deployedseed>1119116582</deployedseed>
    <deployedseed>1550712488</deployedseed>
    <deployedseed>690732707</deployedseed>
    <deployedseed>945264847</deployedseed>
    <deployedseed>1898838649</deployedseed>
    <deployedseed>1244203784</deployedseed>
    <deployedseed>1652939468</deployedseed>
    <deployedseed>2010231012</deployedseed>
    <deployedseed>1999154502</deployedseed>
    <deployedseed>1104531435</deployedseed>
    <deployedseed>1417448100</deployedseed>
    <deployedseed>353985780</deployedseed>
    <deployedseed>64182190</deployedseed>
    <deployedseed>938220842</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>surface_area_cm2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>Surface_area_m2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Forgetting to Square the Side Length</description>
      <testinput>
        <name>ans1</name>
        <value>Msurface1_cm2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>MSurface1_m2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description> Incorrect Unit Conversion (cm² to m²) divding by 100 instade of 10 000</description>
      <testinput>
        <name>ans1</name>
        <value>surface_area_cm2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>MSurface2_m2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.2000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Using incorect formular that is 6l^3 instade of 6l^2</description>
      <testinput>
        <name>ans1</name>
        <value>Msurface3_cm2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>MSurface3_m2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt2-4-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
