<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Application of linear motion in emergency braking</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A vehicle is travelling along a straight road at a constant velocity of \({@v_initial@}\) m/s. The driver observes a hazard and has a reaction time (thinking time) of \({@t_thinking@}\) s before applying the brakes. Once the brakes are applied, the vehicle comes to a complete rest in \({@t_braking@}\) s.</p>

<ol>
    <li>Calculate the distance travelled during the thinking time (the period of constant velocity). [[input:ans1]]  [[validation:ans1]] [[feedback:prt1]] m</li>
    <li>Determine the acceleration of the vehicle during the braking phase. [[input:ans2]] m/s² [[validation:ans2]] [[feedback:prt2]] </li>
    <li>Calculate the total displacement of the vehicle from the moment the hazard was seen until it came to a stop. [[input:ans3]] m [[validation:ans3]] [[feedback:prt3]] </li>
</ol>

<p><i>( Ensure you include the correct sign for acceleration.)</i></p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>1. Thinking Distance (Constant Velocity)</p>
<p>Since the velocity is constant during the reaction time, we use the standard speed-distance formula:</p>
\(\begin{align}
s_1 &= v \times t \\
&= {@v_initial@} \times {@t_thinking@} \\
&= {@ta1@} \text{ m}
\end{align}\)

<p>2. Acceleration</p>
<p>Acceleration is the rate of change of velocity. Given that the final velocity \(v\) is 0:</p>
\(\begin{align}
a &= \frac{v - u}{t} \\
&= \frac{0 - {@v_initial@}}{{@t_braking@}} \\
&= \mathbf{{@ta2@} \text{ m/s}^2}
\end{align}\)

<p>3. Total Displacement</p>
<p>The total displacement is the sum of the distance travelled while thinking and the distance travelled while braking:</p>
\(\begin{align}
\text{Total } s &= s_{\text{thinking}} + s_{\text{braking}} \\
s_{\text{braking}} &= ut + \frac{1}{2}at^2 \\
&= {@v_initial@} \times {@t_braking@} +\frac{1}{2} \times {@ta2@} \times {@t_braking@}^2\\
\text{Total } s &= {@ta3@} \text{ m}
\end{align}\)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Randomized Initial Conditions */
v_initial: 15 + rand(11);      /* Initial velocity: 15 to 25 m/s */
t_thinking: 0.6 + (rand(6)/10); /* Thinking time: 0.6 to 1.1 s */
t_braking: 4 + rand(4);        /* Time to stop: 4 to 7 s */

/* 2. Linear Motion Calculations */
/* Phase 1: Constant Velocity (Distance = Speed * Time) */
dist_thinking: v_initial * t_thinking;

/* Phase 2: Acceleration (a = (v-u)/t) */
/* Final velocity v is 0 (comes to rest) */
accel: (0 - v_initial) / t_braking;

/* Phase 3: Total Displacement (Sum of both phases) */
/* Using s = ut + 0.5at^2 for the braking phase */
dist_braking: (v_initial * t_braking) + (0.5 * accel * t_braking^2);
total_displacement: dist_braking;

/* 3. Teacher Answers (Teacher values) */
ta1: dist_thinking;
ta2: accel;
ta3: total_displacement;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@t_thinking@}</p>
<p>{@v_initial@}</p>
<p>{@t_braking@}</p>
<p>{@ta1@}</p>
<p>{@ta2@}</p>
<p>{@ta3@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1272627558</deployedseed>
    <deployedseed>1172675312</deployedseed>
    <deployedseed>95906685</deployedseed>
    <deployedseed>752009154</deployedseed>
    <deployedseed>246889776</deployedseed>
    <deployedseed>1263806797</deployedseed>
    <deployedseed>703663778</deployedseed>
    <deployedseed>1060179341</deployedseed>
    <deployedseed>415339778</deployedseed>
    <deployedseed>1191841716</deployedseed>
    <deployedseed>501801492</deployedseed>
    <deployedseed>1791644662</deployedseed>
    <deployedseed>74017537</deployedseed>
    <deployedseed>720923124</deployedseed>
    <deployedseed>743983032</deployedseed>
    <deployedseed>2083749830</deployedseed>
    <deployedseed>1136830868</deployedseed>
    <deployedseed>1313378756</deployedseed>
    <deployedseed>1090926344</deployedseed>
    <deployedseed>536547393</deployedseed>
  </question>
</quiz>
