<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Reflect shape in a 45-degree line</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
The red shape is shown on the coordinate plane, together with the dashed mirror line \( y = {@lineexpr@} \).
</p>
<p>Reflect the shape in the line \( y = {@lineexpr@} \) by dragging the blue vertices to their image positions.</p>

[[jsxgraph
  input-ref-sP1="sP1"
  input-ref-sP2="sP2"
  input-ref-sP3="sP3"
  input-ref-sP4="sP4"
  input-ref-sP5="sP5"
  input-ref-sP6="sP6"
]]

var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [-6, 9, 9, -6],
  axis: true,
  keepaspectratio: true,
  showCopyright: false,
  grid: false,
  pan: { enabled: false },
  zoom: { enabled: false },
  defaultAxes: {
    x: { ticks: { insertTicks: false, ticksDistance: 1, minorTicks: 0, drawLabels: true } },
    y: { ticks: { insertTicks: false, ticksDistance: 1, minorTicks: 0, drawLabels: true } }
  }
});
board.create('grid', [], { gridX: 1, gridY: 1 });

// ----------------------------
// Mirror line  y = x + c  (45 degrees)
// ----------------------------
var c = {#c#};
board.create('line', [[0, c], [1, 1 + c]], {
  dash: 2, strokeColor: '#888888', strokeWidth: 2,
  fixed: true, highlight: false
});
board.create('text', [6.2, 6.2 + c, 'mirror line'], {
  fontSize: 12, color: '#888888', fixed: true, anchorX: 'left'
});

// ----------------------------
// Teacher shape (fixed, red)
// ----------------------------
var teacherShape = {#base_shape#};
var teacherPts = [];

for (var i = 0; i < teacherShape.length; i++) {
  teacherPts.push(
    board.create('point', teacherShape[i], { visible: false, fixed: true })
  );
}

for (var i = 0; i < teacherPts.length; i++) {
  board.create('segment', [
    teacherPts[i],
    teacherPts[(i + 1) % teacherPts.length]
  ], { strokeColor: 'red', strokeWidth: 3, fixed: true });
}

// ----------------------------
// Student shape (draggable points only)
// ----------------------------
var studentShape = {#student_shape#};
var studentPts = [];

for (var i = 0; i < studentShape.length; i++) {
  studentPts.push(
    board.create('point', studentShape[i], {
      color: 'blue', size: 3, snapToGrid: true, withLabel: false
    })
  );
}

for (var i = 0; i < studentPts.length; i++) {
  board.create('segment', [
    studentPts[i],
    studentPts[(i + 1) % studentPts.length]
  ], { strokeColor: 'blue', strokeWidth: 2, fixed: true, highlight: false });
}

// ----------------------------
// STACK bindings (CRITICAL)
// ----------------------------
stack_jxg.bind_point(sP1, studentPts[0]);
stack_jxg.bind_point(sP2, studentPts[1]);
stack_jxg.bind_point(sP3, studentPts[2]);
stack_jxg.bind_point(sP4, studentPts[3]);
stack_jxg.bind_point(sP5, studentPts[4]);
stack_jxg.bind_point(sP6, studentPts[5]);

board.unsuspendUpdate();

[[/jsxgraph]]


<div style="display:none">
[[input:sP1]] [[validation:sP1]]
[[input:sP2]] [[validation:sP2]]
[[input:sP3]] [[validation:sP3]]
[[input:sP4]] [[validation:sP4]]
[[input:sP5]] [[validation:sP5]]
[[input:sP6]] [[validation:sP6]]
</div>

[[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p>Reflect each vertex straight across the dashed line \( y = {@lineexpr@} \): its image sits the same perpendicular distance away, on the other side. Join the images in the same order to get the green shape.</p>

<p>The thin lines join each vertex to its image. Each one meets the mirror line at a right angle, and the line cuts it exactly in half.</p>

[[jsxgraph height='430px' width='430px']]
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [-6, 9, 9, -6],
  axis: true, keepaspectratio: true, showCopyright: false, grid: false,
  pan: { enabled: false }, zoom: { enabled: false },
  defaultAxes: {
    x: { ticks: { insertTicks: false, ticksDistance: 1, minorTicks: 0, drawLabels: true } },
    y: { ticks: { insertTicks: false, ticksDistance: 1, minorTicks: 0, drawLabels: true } }
  }
});
board.create('grid', [], { gridX: 1, gridY: 1 });

var c = {#c#};
board.create('line', [[0, c], [1, 1 + c]], {
  dash: 2, strokeColor: '#888888', strokeWidth: 2, fixed: true, highlight: false
});
board.create('text', [6.2, 6.2 + c, 'y = x' + (c === 0 ? '' : (c > 0 ? ' + ' + c : ' - ' + (-c)))], {
  fontSize: 12, color: '#888888', fixed: true, anchorX: 'left'
});

function drawShape(coords, col, fillop) {
  var pts = coords.map(function (p) {
    return board.create('point', p, { visible: false, fixed: true });
  });
  board.create('polygon', pts, {
    fillColor: col, fillOpacity: fillop,
    borders: { strokeColor: col, strokeWidth: 3 },
    fixed: true, vertices: { visible: false }
  });
}

var origPts = {#base_shape#};
var reflPts = {#reflected_shape#};
drawShape(origPts, '#d62728', 0.15);   // original (red)
drawShape(reflPts, '#2ca02c', 0.30);   // reflected (green)

// Perpendicular construction lines: each vertex joined to its image
for (var i = 0; i < origPts.length; i++) {
  board.create('segment', [origPts[i], reflPts[i]], {
    strokeColor: '#555555', strokeWidth: 1, dash: 2, fixed: true, highlight: false
  });
  board.create('point', [
    (origPts[i][0] + reflPts[i][0]) / 2,
    (origPts[i][1] + reflPts[i][1]) / 2
  ], { size: 1, color: '#555555', fixed: true, withLabel: false, highlight: false });
}

board.renderer.svgRoot.style.pointerEvents = "none";
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 45-degree mirror line  y = x + c  (never an axis) */
c : rand([-1, 0, 1]);
lineexpr : x + c;

/* L-shape size and placement */
length : rand([3, 4]);
height : rand([3, 4]);
gap : rand([1, 2]);
bx : rand([2, 3]);
by : bx + c - height - gap;   /* guarantees the whole shape sits strictly below the line */
start_x : bx;
start_y : by;

/* Base L-shape (teacher shape) */
base_shape : [
  [start_x, start_y],
  [start_x + length, start_y],
  [start_x + length - 1, start_y + 1],
  [start_x + 1, start_y + 1],
  [start_x + 1, start_y + height],
  [start_x, start_y + height]
];

/* Reflection across y = x + c :  (x,y) -> (y - c, x + c) */
reflected_shape : map(lambda([p], [p[2] - c, p[1] + c]), base_shape);

/* Student's draggable starting shape: a copy on the FAR side of the line
   (upper-left), clear of the red original and not the answer. */
student_shape : map(lambda([p], [p[1] - 7, p[2] + 4]), base_shape);

/* Targets used only by the question tests */
t_trans : map(lambda([p], [p[1] + 2, p[2] + 1]), base_shape);   /* a translation of the original */
t_yaxis : map(lambda([p], [-p[1], p[2]]), base_shape);          /* reflection in the y-axis */
t_xaxis : map(lambda([p], [p[1], -p[2]]), base_shape);          /* reflection in the x-axis */
wrong_shift : map(lambda([p], [p[1] + 1, p[2]]), reflected_shape);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>line y = {@lineexpr@}; base {@base_shape@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>sP1</name>
      <type>algebraic</type>
      <tans>reflected_shape[1]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>sP2</name>
      <type>algebraic</type>
      <tans>reflected_shape[2]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>sP3</name>
      <type>algebraic</type>
      <tans>reflected_shape[3]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>sP4</name>
      <type>algebraic</type>
      <tans>reflected_shape[4]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>sP5</name>
      <type>algebraic</type>
      <tans>reflected_shape[5]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>sP6</name>
      <type>algebraic</type>
      <tans>reflected_shape[6]</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[student_shape : sort([sP1, sP2, sP3, sP4, sP5, sP6]);
cent(P) := block([m : length(P), ax, ay], ax : apply("+", map(first, P))/m, ay : apply("+", map(second, P))/m, sort(map(lambda([q], [q[1] - ax, q[2] - ay]), P)));
stu_c : cent([sP1, sP2, sP3, sP4, sP5, sP6]);
base_c : cent(base_shape);
refl_yaxis : sort(map(lambda([p], [-p[1], p[2]]), base_shape));
refl_xaxis : sort(map(lambda([p], [p[1], -p[2]]), base_shape));
refl_yx : sort(map(lambda([p], [p[2], p[1]]), base_shape));]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct reflection?</description>
        <answertest>AlgEquiv</answertest>
        <sans>student_shape</sans>
        <tans>sort(reflected_shape)</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Translation (moved but not flipped)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>stu_c</sans>
        <tans>base_c</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have <strong>slid</strong> the shape to a new position, but it is still facing the same way. A reflection <strong>flips</strong> the shape over the dashed line, like a mirror image, so it ends up back-to-front. Picture folding the paper along the line: the shape should land on the far side, turned over.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Reflected in the y-axis?</description>
        <answertest>AlgEquiv</answertest>
        <sans>student_shape</sans>
        <tans>refl_yaxis</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have reflected in the <strong>y-axis</strong> (the vertical line), not the dashed 45&deg; line. Use the dashed line as your mirror: each vertex should end up the same perpendicular distance from <em>that</em> line, straight across it on the other side.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Reflected in the x-axis?</description>
        <answertest>AlgEquiv</answertest>
        <sans>student_shape</sans>
        <tans>refl_xaxis</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have reflected in the <strong>x-axis</strong> (the horizontal line), not the dashed 45&deg; line. Reflect in the dashed line instead: look straight across it at right angles, and place each vertex the same distance beyond it.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Reflected in y = x (line through the origin)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>student_shape</sans>
        <tans>refl_yx</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Close! You have reflected in the line \(y = x\) through the origin, but the mirror line here is the <em>parallel</em> line \(y = {@lineexpr@}\). Measure the perpendicular distance to the <strong>dashed</strong> line, which is shifted to one side of \(y = x\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not yet. Picture folding the plane along the dashed line so the two sides meet. Each vertex should map to the point straight across the line, the same perpendicular distance away on the opposite side.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1759191554</deployedseed>
    <deployedseed>610362512</deployedseed>
    <deployedseed>917059148</deployedseed>
    <deployedseed>1870303665</deployedseed>
    <deployedseed>1071327382</deployedseed>
    <deployedseed>1814251858</deployedseed>
    <deployedseed>901952842</deployedseed>
    <deployedseed>1962256078</deployedseed>
    <deployedseed>32215440</deployedseed>
    <deployedseed>392144042</deployedseed>
    <deployedseed>1757093948</deployedseed>
    <deployedseed>1934070349</deployedseed>
    <deployedseed>5250438</deployedseed>
    <deployedseed>486269779</deployedseed>
    <deployedseed>812228446</deployedseed>
    <deployedseed>327212806</deployedseed>
    <deployedseed>1019911420</deployedseed>
    <deployedseed>440968055</deployedseed>
    <deployedseed>1454562098</deployedseed>
    <deployedseed>1199367907</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct reflection gets full marks.</description>
      <testinput>
        <name>sP1</name>
        <value>reflected_shape[1]</value>
      </testinput>
      <testinput>
        <name>sP2</name>
        <value>reflected_shape[2]</value>
      </testinput>
      <testinput>
        <name>sP3</name>
        <value>reflected_shape[3]</value>
      </testinput>
      <testinput>
        <name>sP4</name>
        <value>reflected_shape[4]</value>
      </testinput>
      <testinput>
        <name>sP5</name>
        <value>reflected_shape[5]</value>
      </testinput>
      <testinput>
        <name>sP6</name>
        <value>reflected_shape[6]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description><![CDATA[Translation only -> translation diagnostic.]]></description>
      <testinput>
        <name>sP1</name>
        <value>t_trans[1]</value>
      </testinput>
      <testinput>
        <name>sP2</name>
        <value>t_trans[2]</value>
      </testinput>
      <testinput>
        <name>sP3</name>
        <value>t_trans[3]</value>
      </testinput>
      <testinput>
        <name>sP4</name>
        <value>t_trans[4]</value>
      </testinput>
      <testinput>
        <name>sP5</name>
        <value>t_trans[5]</value>
      </testinput>
      <testinput>
        <name>sP6</name>
        <value>t_trans[6]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description><![CDATA[Reflected in the y-axis -> y-axis diagnostic.]]></description>
      <testinput>
        <name>sP1</name>
        <value>t_yaxis[1]</value>
      </testinput>
      <testinput>
        <name>sP2</name>
        <value>t_yaxis[2]</value>
      </testinput>
      <testinput>
        <name>sP3</name>
        <value>t_yaxis[3]</value>
      </testinput>
      <testinput>
        <name>sP4</name>
        <value>t_yaxis[4]</value>
      </testinput>
      <testinput>
        <name>sP5</name>
        <value>t_yaxis[5]</value>
      </testinput>
      <testinput>
        <name>sP6</name>
        <value>t_yaxis[6]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description><![CDATA[Reflected in the x-axis -> x-axis diagnostic.]]></description>
      <testinput>
        <name>sP1</name>
        <value>t_xaxis[1]</value>
      </testinput>
      <testinput>
        <name>sP2</name>
        <value>t_xaxis[2]</value>
      </testinput>
      <testinput>
        <name>sP3</name>
        <value>t_xaxis[3]</value>
      </testinput>
      <testinput>
        <name>sP4</name>
        <value>t_xaxis[4]</value>
      </testinput>
      <testinput>
        <name>sP5</name>
        <value>t_xaxis[5]</value>
      </testinput>
      <testinput>
        <name>sP6</name>
        <value>t_xaxis[6]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-4-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description><![CDATA[Other wrong shape -> generic hint.]]></description>
      <testinput>
        <name>sP1</name>
        <value>wrong_shift[1]</value>
      </testinput>
      <testinput>
        <name>sP2</name>
        <value>wrong_shift[2]</value>
      </testinput>
      <testinput>
        <name>sP3</name>
        <value>wrong_shift[3]</value>
      </testinput>
      <testinput>
        <name>sP4</name>
        <value>wrong_shift[4]</value>
      </testinput>
      <testinput>
        <name>sP5</name>
        <value>wrong_shift[5]</value>
      </testinput>
      <testinput>
        <name>sP6</name>
        <value>wrong_shift[6]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-5-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
