<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Area Scale Factor - Finding the Linear Scale Factor from Two Areas</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Two similar triangles have areas \( {@area_small@} \,\text{cm}^2\) and \( {@area_large@} \,\text{cm}^2\) respectively.</p>
<p>The blue triangle below has the smaller area. Drag the slider to stretch the orange triangle until its displayed area matches the larger triangle's area, then read off the scale factor \(k\) shown next to the slider.</p>

[[jsxgraph height='420px' width='560px']]
var qden = {#qden#};
var pnum = {#pnum#};
var m = {#m#};
var areaSmall = {#area_small#};
var areaLarge = {#area_large#};
var kTarget = pnum / qden;
var maxK = kTarget * 1.3;
var t = Math.sqrt(m);
var sSmall = qden * t;

function trianglePts(x0, y0, s) {
    return [[x0, y0], [x0 + 2 * s, y0], [x0, y0 + s]];
}

var gapX = 2 * sSmall + 3;
var bbYmax = sSmall * maxK + 4.5;
var bbXmax = gapX + 2 * sSmall * maxK + 3;

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-1, bbYmax, bbXmax, -3],
    keepaspectratio: true, axis: false, showNavigation: false, showCopyright: false,
    pan: { enabled: false }, zoom: { enabled: false }
});

// Fixed blue triangle with the given smaller area
var bp = trianglePts(0, 0, sSmall).map(function (p) {
    return board.create('point', p, { visible: false, fixed: true });
});
board.create('polygon', bp, {
    fillColor: '#4C72B0', fillOpacity: 0.35,
    borders: { strokeColor: '#4C72B0', strokeWidth: 2 },
    vertices: { visible: false }, hasInnerPoints: false
});
board.create('text', [sSmall * 0.6, -1.4, 'Given area = ' + areaSmall + ' cm\u00B2'],
    { fontSize: 13, fixed: true, anchorX: 'middle' });

// Slider controlling the stretch factor of the orange triangle
var slider = board.create('slider',
    [[gapX, sSmall * maxK + 1.8], [gapX + Math.max(sSmall, 3), sSmall * maxK + 1.8], [1, 1, maxK]],
    { name: 'scale factor k', snapWidth: 0.01, size: 4 });

var basePts = trianglePts(0, 0, sSmall);
var op = [];
for (var i = 0; i < 3; i++) {
    (function (i) {
        op.push(board.create('point', [
            function () { return gapX + basePts[i][0] * slider.Value(); },
            function () { return basePts[i][1] * slider.Value(); }
        ], { visible: false, fixed: true }));
    })(i);
}
board.create('polygon', op, {
    fillColor: '#DD8452', fillOpacity: 0.35,
    borders: { strokeColor: '#DD8452', strokeWidth: 2 },
    vertices: { visible: false }, hasInnerPoints: false
});

board.create('text', [function () { return gapX + sSmall * slider.Value() * 0.6; }, -1.4, function () {
    var k = slider.Value();
    var area = (areaSmall * k * k).toFixed(2);
    return 'k = ' + k.toFixed(2) + ', area = ' + area + ' cm\u00B2';
}], { fontSize: 13, fixed: true, anchorX: 'middle' });

board.create('text', [gapX, sSmall * maxK + 3, 'Target to match: ' + areaLarge + ' cm\u00B2'],
    { fontSize: 13, fixed: true, color: '#a00', anchorX: 'left' });

board.create('button', [gapX, sSmall * maxK + 4.2, 'Reveal the matching scale factor', function () {
    slider.setValue(kTarget);
    board.update();
}], { fixed: true });
[[/jsxgraph]]

<p>What is the linear scale factor from the smaller triangle to the larger triangle? Give your answer correct to <b>2</b> decimal places.</p>
<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>The ratio of the two areas is \(\dfrac{{@area_large@}}{{@area_small@}}\). Because area scales with the <strong>square</strong> of the linear scale factor, this area ratio is itself a squared quantity, so the linear scale factor is found by taking the square root of the area ratio.</p>

[[jsxgraph height='360px' width='520px']]
var qden = {#qden#};
var pnum = {#pnum#};
var m = {#m#};
var areaSmall = {#area_small#};
var areaLarge = {#area_large#};
var t = Math.sqrt(m);
var sSmall = qden * t;
var sLarge = pnum * t;

function trianglePts(x0, y0, s) { return [[x0, y0], [x0 + 2 * s, y0], [x0, y0 + s]]; }
function drawTri(x0, y0, s, col, label) {
    var pts = trianglePts(x0, y0, s).map(function (p) {
        return board.create('point', p, { visible: false, fixed: true });
    });
    board.create('polygon', pts, {
        fillColor: col, fillOpacity: 0.35,
        borders: { strokeColor: col, strokeWidth: 2 },
        vertices: { visible: false }, hasInnerPoints: false
    });
    board.create('text', [x0 + s * 0.6, -1.4, label], { fontSize: 13, fixed: true, anchorX: 'middle' });
}

var gapX = 2 * sSmall + 3;
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-1, sLarge + 2, gapX + 2 * sLarge + 2, -3],
    keepaspectratio: true, axis: false, showNavigation: false, showCopyright: false,
    pan: { enabled: false }, zoom: { enabled: false }
});
drawTri(0, 0, sSmall, '#4C72B0', 'Smaller, area ' + areaSmall + ' cm\u00B2');
drawTri(gapX, 0, sLarge, '#DD8452', 'Larger, area ' + areaLarge + ' cm\u00B2');
board.renderer.svgRoot.style.pointerEvents = "none";
[[/jsxgraph]]

<p>Taking the square root of the area ratio gives the linear scale factor from the smaller triangle to the larger triangle.</p>
<p>\(\sqrt{\dfrac{{@area_large@}}{{@area_small@}}} = {@ta@}\)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Linear scale factor p/q in lowest terms, p > q */
qden : rand(4) + 2;
pnum : qden + rand(4) + 1;
g : gcd(pnum, qden);
pnum : pnum / g;
qden : qden / g;

/* Base multiplier so both areas come out as whole numbers */
m : rand(5) + 2;
area_small : m * qden^2;
area_large : m * pnum^2;

/* Correct answer: linear scale factor from the smaller to the larger triangle */
ta : decimalplaces(pnum/qden, 2);

/* mistakes */
/* Forgot to take the square root, used the area ratio directly */
ta1 : decimalplaces(area_large/area_small, 2);
/* Found the scale factor the wrong way round (large to small) */
ta2 : decimalplaces(qden/pnum, 2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Areas {@area_small@} and {@area_large@}, linear scale factor = {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct linear scale factor?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Used the area ratio directly, without a square root</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you gave the ratio of the areas, \(\dfrac{{@area_large@}}{{@area_small@}}\), as your final answer. That ratio is the <strong>area</strong> scale factor, not the linear scale factor. Since area scales with the square of the linear scale factor, take the square root of the area ratio to get the scale factor for the sides.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Found the scale factor in the wrong direction</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you found the scale factor from the <strong>larger</strong> triangle to the <strong>smaller</strong> one, which gives a value less than 1. The question asks for the scale factor from the smaller triangle to the larger one, so divide the larger area by the smaller area before taking the square root.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Divide the larger area by the smaller area to get the area scale factor, then take the square root of that result to find the linear scale factor.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>Correct answer gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Used the area ratio directly, without a square root.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Found the scale factor in the wrong direction.</description>
      <testinput>
        <name>ans1</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty/>
        <expectedanswernote>prt1-3-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
