<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying data collection methods</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the following research scenario:</p>
<p><em>{@q_text@}</em></p>
<p>Identify the specific method of data collection being used in this situation.</p>
<div>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>Correct answer: {@correct_text@}</strong></p>
<p>Match each scenario to the method by its defining feature:</p>
<ul>
    <li><strong>Observation:</strong> watching and recording behaviour as it happens naturally, without interacting.</li>
    <li><strong>Interview:</strong> a direct conversation to gather detailed, in-depth answers from a person or small group.</li>
    <li><strong>Questionnaire:</strong> a standard set of written questions used to gather data quickly from many people.</li>
    <li><strong>Experimentation:</strong> manipulating one factor while controlling others to test cause and effect.</li>
    <li><strong>Secondary Data Collection:</strong> using existing data already collected and recorded by someone else.</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* --- 1. SCENARIOS (described actions -> method) --- */
/* Methods restricted to those taught in the lesson:                       */
/* Observation, Interview, Questionnaire, Experimentation, Secondary Data. */
scenarios: [
  ["A wildlife ranger hides in a blind to count the elephants visiting a watering hole at sunset, without disturbing them.", "Observation"],
  ["A researcher sits down one-on-one with each shopkeeper in a market and asks them detailed questions about their daily challenges, recording each conversation.", "Interview"],
  ["A university sends a standard form by email to all 10,000 students, asking them to rate the cafeteria food on a scale of 1 to 5.", "Questionnaire"],
  ["A farmer applies a new fertilizer to Plot A and the old fertilizer to Plot B, keeping everything else the same, to see which produces more maize.", "Experimentation"],
  ["An economist analyses Kenya's inflation rates from 1990 to 2020 using annual reports already published by the National Bureau of Statistics.", "Secondary Data Collection"]
];

/* Pick a random scenario */
rand_idx: rand(length(scenarios)) + 1;
chosen_scenario: scenarios[rand_idx];
q_text: chosen_scenario[1];
correct_text: chosen_scenario[2];

/* --- 2. OPTIONS: correct + 3 distinct distractors from taught methods --- */
all_methods: ["Observation", "Interview", "Questionnaire", "Experimentation", "Secondary Data Collection"];
wrong_pool: sublist(all_methods, lambda([x], x # correct_text));
chosen_wrongs: rand_selection(wrong_pool, 3);

correct_answer: [[1, true, correct_text]];
incorrect_answers: [
   [2, false, chosen_wrongs[1]],
   [3, false, chosen_wrongs[2]],
   [4, false, chosen_wrongs[3]]
];
ta: random_permutation(append(correct_answer, incorrect_answers));

/* Helper for question testing: a wrong option's value */
wrongval(t) := first(delete(1, map(first, t)));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@correct_text@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">Methods of data collection: identifying the method (Observation, Interview, Questionnaire, Experimentation, Secondary Data) from a described scenario.</p>
<p id="reviewer">Ben (STACK audit)</p>
<p id="description">Students read a randomised real-world scenario describing a data-gathering activity and identify the method used. Methods restricted to the five taught in the lesson; each wrong answer receives a diagnostic hint pointing at the correct method's defining feature.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[hint: if is(correct_text = "Observation") then "This scenario involves watching and recording behaviour as it naturally happens, without interacting with the subjects."
elseif is(correct_text = "Interview") then "This scenario involves a direct conversation to gather detailed, in-depth answers from a person or small group."
elseif is(correct_text = "Questionnaire") then "This scenario uses a standard set of written questions to gather data quickly from a large number of people."
elseif is(correct_text = "Experimentation") then "This scenario manipulates one factor while keeping the others fixed, to test its effect (cause and effect)."
else "This scenario uses data that was already collected and recorded by someone else for another purpose.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. {@hint@} Which method does that describe?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1939807603</deployedseed>
    <deployedseed>272475855</deployedseed>
    <deployedseed>856172502</deployedseed>
    <deployedseed>45461360</deployedseed>
    <deployedseed>1189181072</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct method selected: full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>A distractor selected: fires the F branch with the diagnostic hint.</description>
      <testinput>
        <name>ans1</name>
        <value>wrongval(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
