<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Selecting appropriate data collection methods</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Read the following research goal:</p>
<p><em>{@q_text@}</em></p>
<p>Which method of data collection is <strong>most appropriate</strong> for achieving this goal?</p>
<div>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>Most appropriate method: {@correct_text@}</strong></p>
<p>Choose the method by matching the goal to what each method does best:</p>
<ul>
    <li><strong>Observation:</strong> best when you need to record real behaviour as it naturally happens.</li>
    <li><strong>Interview:</strong> best for deep, detailed understanding from a few people.</li>
    <li><strong>Questionnaire:</strong> best for collecting standardised data quickly from many people.</li>
    <li><strong>Experimentation:</strong> best for testing cause and effect under controlled conditions.</li>
    <li><strong>Secondary Data Collection:</strong> best when the data you need already exists in records.</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* --- 1. RESEARCH GOALS (goal -> most appropriate method) --- */
/* Methods restricted to those taught in the lesson.                       */
/* Scenarios deliberately different from the 'Identifying' question.       */
scenarios: [
  ["A school counsellor wants to understand, in depth, how a few students coped with exam stress and to explore their personal feelings about it.", "Interview"],
  ["A county government wants to quickly collect the ages and occupations of 5,000 residents using one standard form.", "Questionnaire"],
  ["A road-safety officer wants to find out how many drivers actually stop at a particular junction, watching them without their knowledge.", "Observation"],
  ["A science teacher wants to find out whether adding sugar to water makes beans germinate faster, keeping all other conditions equal.", "Experimentation"],
  ["A student writing a geography report needs rainfall figures for their region going back thirty years.", "Secondary Data Collection"]
];

/* Pick a random goal */
rand_idx: rand(length(scenarios)) + 1;
chosen_scenario: scenarios[rand_idx];
q_text: chosen_scenario[1];
correct_text: chosen_scenario[2];

/* --- 2. OPTIONS: correct + 3 distinct distractors from taught methods --- */
all_methods: ["Observation", "Interview", "Questionnaire", "Experimentation", "Secondary Data Collection"];
wrong_pool: sublist(all_methods, lambda([x], x # correct_text));
chosen_wrongs: rand_selection(wrong_pool, 3);

correct_answer: [[1, true, correct_text]];
incorrect_answers: [
   [2, false, chosen_wrongs[1]],
   [3, false, chosen_wrongs[2]],
   [4, false, chosen_wrongs[3]]
];
ta: random_permutation(append(correct_answer, incorrect_answers));

/* Helper for question testing: a wrong option's value */
wrongval(t) := first(delete(1, map(first, t)));]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@correct_text@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept">Methods of data collection: selecting the most appropriate method for a stated research goal.</p>
<p id="reviewer">Ben (STACK audit)</p>
<p id="description">Students are given a randomised research goal and select the most appropriate data-collection method. Methods restricted to the five taught in the lesson; framing and scenarios are distinct from the 'Identifying data collection methods' question. Each wrong answer receives a diagnostic hint about what the correct method is best suited for.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>radio</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>nonotanswered</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[hint: if is(correct_text = "Observation") then "This goal needs you to record real behaviour as it naturally happens, without influencing it."
elseif is(correct_text = "Interview") then "This goal needs deep, detailed understanding from just a few people, exploring their feelings."
elseif is(correct_text = "Questionnaire") then "This goal needs standardised data collected quickly from a large number of people."
elseif is(correct_text = "Experimentation") then "This goal needs you to test cause and effect by controlling conditions and changing one factor."
else "This goal can be met using data that already exists in records collected by someone else.";]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>mcq_correct(ta)[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. {@hint@} Which method fits that best?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>296407186</deployedseed>
    <deployedseed>1481155070</deployedseed>
    <deployedseed>1937597875</deployedseed>
    <deployedseed>175433249</deployedseed>
    <deployedseed>1116838791</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct method selected: full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>mcq_correct(ta)[1]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>A distractor selected: fires the F branch with the diagnostic hint.</description>
      <testinput>
        <name>ans1</name>
        <value>wrongval(ta)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
