<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating sines,cosines and tangents of acute angles using a calculator</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Using a calculator, calculate the sine, cosine, and tangent of \({@angle@}^\circ\), giving your answers correct to exactly 4 decimal places.</p>

<p>i. \(\sin({@angle@}^\circ) =\) [[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>
<p>ii. \(\cos({@angle@}^\circ) =\) [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>
<p>iii. \(\tan({@angle@}^\circ) =\) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>Worked Solution:</strong></p>

<p>Using a calculator, we evaluate the trigonometric functions for \({@angle@}^\circ\) and round each result to exactly 4 decimal places:</p>

<ul>
    <li>\(\sin({@angle@}^\circ) = {@disp_1@}\)</li>
    <li>\(\cos({@angle@}^\circ) = {@disp_2@}\)</li>
    <li>\(\tan({@angle@}^\circ) = {@disp_3@}\)</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Random angle between 5 and 90 with step 0.02 */
angle : 5 + 0.02 * rand(floor((90-5)/0.02) + 1);

/* Evaluate to floating-point decimals natively */
ans_1 : float(sin(angle * %pi/180));
ans_2 : float(cos(angle * %pi/180));
ans_3 : float(tan(angle * %pi/180));

/* Display variables formatted to exactly 4 decimal places */
disp_1 : dispdp(ans_1, 4);
disp_2 : dispdp(ans_2, 4);
disp_3 : dispdp(ans_3, 4);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Selected angle={@angle@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>disp_1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>disp_2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>disp_3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>ans_1</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans2</sans>
        <tans>ans_2</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans3</sans>
        <tans>ans_3</tans>
        <testoptions>4</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>780502075</deployedseed>
    <deployedseed>1123984529</deployedseed>
    <deployedseed>417763661</deployedseed>
    <deployedseed>1743653370</deployedseed>
    <deployedseed>1764833450</deployedseed>
    <deployedseed>128053682</deployedseed>
    <deployedseed>1949896930</deployedseed>
    <deployedseed>1251948106</deployedseed>
    <deployedseed>402220183</deployedseed>
    <deployedseed>667663651</deployedseed>
    <deployedseed>32358042</deployedseed>
    <deployedseed>1133661192</deployedseed>
    <deployedseed>1632699816</deployedseed>
    <deployedseed>1859848413</deployedseed>
    <deployedseed>818950146</deployedseed>
    <deployedseed>1987468147</deployedseed>
    <deployedseed>1716541277</deployedseed>
    <deployedseed>2090086938</deployedseed>
    <deployedseed>84875291</deployedseed>
    <deployedseed>270385381</deployedseed>
    <deployedseed>1603380672</deployedseed>
    <deployedseed>1475406943</deployedseed>
    <deployedseed>709155714</deployedseed>
    <deployedseed>1200081430</deployedseed>
    <deployedseed>433444913</deployedseed>
    <deployedseed>1994934809</deployedseed>
    <deployedseed>1028855330</deployedseed>
    <deployedseed>1266323299</deployedseed>
    <deployedseed>1201193688</deployedseed>
    <deployedseed>762909789</deployedseed>
    <deployedseed>1845597738</deployedseed>
    <deployedseed>503206755</deployedseed>
    <deployedseed>773373561</deployedseed>
    <deployedseed>1438421676</deployedseed>
    <deployedseed>826072499</deployedseed>
    <deployedseed>1367815519</deployedseed>
    <deployedseed>673370590</deployedseed>
    <deployedseed>2139505666</deployedseed>
    <deployedseed>1601363967</deployedseed>
    <deployedseed>1098617565</deployedseed>
    <deployedseed>1034233269</deployedseed>
    <deployedseed>213783341</deployedseed>
    <deployedseed>1691342400</deployedseed>
    <deployedseed>1082859976</deployedseed>
    <deployedseed>1029851733</deployedseed>
    <deployedseed>1223172630</deployedseed>
    <deployedseed>1148775549</deployedseed>
    <deployedseed>1417818247</deployedseed>
    <deployedseed>69131106</deployedseed>
    <deployedseed>1703185419</deployedseed>
    <deployedseed>1725374102</deployedseed>
    <deployedseed>1878507318</deployedseed>
    <deployedseed>1657613792</deployedseed>
    <deployedseed>1655790074</deployedseed>
    <deployedseed>1692740068</deployedseed>
    <deployedseed>1113199002</deployedseed>
    <deployedseed>2012649597</deployedseed>
    <deployedseed>99248585</deployedseed>
    <deployedseed>122729168</deployedseed>
    <deployedseed>603920935</deployedseed>
    <deployedseed>2075939184</deployedseed>
    <deployedseed>225432171</deployedseed>
    <deployedseed>505891651</deployedseed>
    <deployedseed>1845874731</deployedseed>
    <deployedseed>1799373218</deployedseed>
    <deployedseed>550302014</deployedseed>
    <deployedseed>45563217</deployedseed>
    <deployedseed>569701851</deployedseed>
    <deployedseed>1370074051</deployedseed>
    <deployedseed>1012892464</deployedseed>
    <deployedseed>122398208</deployedseed>
    <deployedseed>166875566</deployedseed>
    <deployedseed>388159076</deployedseed>
    <deployedseed>145919132</deployedseed>
    <deployedseed>1573526374</deployedseed>
    <deployedseed>1107152156</deployedseed>
    <deployedseed>1924337979</deployedseed>
    <deployedseed>1431532403</deployedseed>
    <deployedseed>1610933130</deployedseed>
    <deployedseed>1566514245</deployedseed>
    <deployedseed>1946403970</deployedseed>
    <deployedseed>1445645006</deployedseed>
    <deployedseed>813118727</deployedseed>
    <deployedseed>2134374387</deployedseed>
    <deployedseed>1585825946</deployedseed>
    <deployedseed>739282570</deployedseed>
    <deployedseed>1727807632</deployedseed>
    <deployedseed>526037165</deployedseed>
    <deployedseed>118026036</deployedseed>
    <deployedseed>2082790991</deployedseed>
    <deployedseed>17010554</deployedseed>
    <deployedseed>2059873236</deployedseed>
    <deployedseed>1063324493</deployedseed>
    <deployedseed>1956056814</deployedseed>
    <deployedseed>72268241</deployedseed>
    <deployedseed>974614094</deployedseed>
    <deployedseed>1039192855</deployedseed>
    <deployedseed>1127662377</deployedseed>
    <deployedseed>1291847995</deployedseed>
    <deployedseed>2110552900</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>disp_1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>disp_2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>disp_3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt2-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt3-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
