<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the height of a building using angle of elevation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A student stands <b>{@d@} m</b> away from a building. The angle of elevation from the ground to the top of the building is <b>{@theta@}&deg;</b>. 
</p>

<p>
Use the diagram below to find the height of the building.
</p>

[[jsxgraph]]
    var d = {#d#};
    var theta = {#theta#};
    var h = {#h#};
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-4, h + 6, d + 8, -4],
        keepAspectRatio: true,
        showNavigation: false,
        showCopyright: false
    });
    
    // Points
    var student = board.create('point', [0, 0], {name: 'Student', size: 3, color: 'black', label: {offset: [-25, -15]}});
    var bldg_base = board.create('point', [d, 0], {visible: false});
    var bldg_top = board.create('point', [d, h], {visible: false});
    
    // Lines and Shapes
    board.create('segment', [[-3, 0], [d+6, 0]], {strokeColor: 'black', strokeWidth: 2}); // Ground
    board.create('segment', [student, bldg_top], {strokeColor: 'black', strokeWidth: 1.5}); // Line of sight
    
    // Building Representation
    board.create('polygon', [
        bldg_base, 
        board.create('point', [d+4, 0], {visible: false}), 
        board.create('point', [d+4, h], {visible: false}), 
        bldg_top
    ], {fillColor: '#cfd8dc', borders: {strokeColor: 'black', strokeWidth: 2}});
    
    // Angles
    board.create('angle', [board.create('point', [d-1, 0], {visible: false}), bldg_base, bldg_top], {radius: d*0.1, type: 'square', strokeColor: 'black'});
    board.create('angle', [bldg_base, student, bldg_top], {radius: d*0.25, strokeColor: 'black', name: theta + "°"});
    
    // Adjusted Labels
    board.create('text', [d/2, -1.5, d + " m"], {anchorX: 'middle', anchorY: 'top', fontSize: 14});
    board.create('text', [d-0.3, h/2, "height = ?"], {anchorX: 'right', anchorY: 'middle', fontSize: 14});
    board.create('text', [d+2, h/2, "Building"], {anchorX: 'middle', anchorY: 'middle', fontSize: 14});
[[/jsxgraph]]

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Height = [[input:ans1]] <b>m</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Step 1: Identify the sides relative to the angle.
</p>

<p>
The angle of elevation is {@theta@}&deg;. <br>
The horizontal distance to the building is the <b>adjacent</b> side ({@d@} m). <br>
The height of the building is the <b>opposite</b> side.
</p>

<p>
Step 2: Use the tangent ratio.
</p>

<p>
\[
\tan(\theta) = \frac{\text{opposite}}{\text{adjacent}}
\]
</p>

<p>
\[
\tan({@theta@}^\circ) = \frac{h}{{@d@}}
\]
</p>

<p>
Step 3: Solve for the height (\(h\)).
</p>

<p>
\[
h = {@d@} \times \tan({@theta@}^\circ)
\]
</p>

<p>
\[
h \approx {@h_disp@}
\]
</p>

<p>
Therefore, the height of the building is:
</p>

<p>
\[
\boxed{@h_disp@} \text{ m}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Distance from object */
d: rand(10) + 10;

/* Angle */
theta: rand([30, 35, 40, 45, 50, 55, 60]);

/* Height calculation evaluated natively as a float */
h: float(d * tan(theta * %pi / 180));

/* Display variable formatted to exactly 2 decimal places */
h_disp: dispdp(h, 2);

/* Mistakes */
m1: float(d / tan(theta * %pi / 180)); /* Inverted tangent */
m2: float(d * sin(theta * %pi / 180)); /* Used sine */
m3: float(d * cos(theta * %pi / 180)); /* Used cosine */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>d: {@d@}, theta: {@theta@}, height: {@h_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Angles of elevation</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to determine the height of a building using angle of elevation and the tangent ratio.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>h_disp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>h</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>mistake 1 - inverted ratio</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You inverted the tangent ratio. Remember, \(\tan(\theta) = \frac{\text{opposite}}{\text{adjacent}}\). Because you want to find the opposite side (height), you should multiply the distance by the tangent of the angle.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>mistake 2 - used sine</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You used sine instead of tangent. The adjacent side is known, and the opposite side is unknown. Use \(\tan(\theta)\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>mistake 3 - used cosine</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You used cosine instead of tangent. The adjacent side is known, and the opposite side is unknown. Use \(\tan(\theta)\).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1915197225</deployedseed>
    <deployedseed>150206404</deployedseed>
    <deployedseed>377226346</deployedseed>
    <deployedseed>788961301</deployedseed>
    <deployedseed>343218966</deployedseed>
    <deployedseed>341426863</deployedseed>
    <deployedseed>912379701</deployedseed>
    <deployedseed>1197127631</deployedseed>
    <deployedseed>871755499</deployedseed>
    <deployedseed>1318344059</deployedseed>
    <deployedseed>1997384458</deployedseed>
    <deployedseed>345915095</deployedseed>
    <deployedseed>603183917</deployedseed>
    <deployedseed>609796159</deployedseed>
    <deployedseed>1813660302</deployedseed>
    <deployedseed>1826899015</deployedseed>
    <deployedseed>1127418023</deployedseed>
    <deployedseed>128538150</deployedseed>
    <deployedseed>2073340704</deployedseed>
    <deployedseed>1778921246</deployedseed>
    <deployedseed>1651636079</deployedseed>
    <deployedseed>1534537287</deployedseed>
    <deployedseed>972110656</deployedseed>
    <deployedseed>2076159027</deployedseed>
    <deployedseed>774389412</deployedseed>
    <deployedseed>1803828749</deployedseed>
    <deployedseed>362976876</deployedseed>
    <deployedseed>619699989</deployedseed>
    <deployedseed>707290157</deployedseed>
    <deployedseed>1481584257</deployedseed>
    <deployedseed>2081870749</deployedseed>
    <deployedseed>1184483246</deployedseed>
    <deployedseed>905493036</deployedseed>
    <deployedseed>453598280</deployedseed>
    <deployedseed>1479031466</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>h_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
