<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the height of a building using angle of elevation</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A student stands {@d@} m away from a building. The angle of elevation of the top of the building is {@theta@}\(^{\circ}\). Find the height of the building.
</p>

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Height = [[input:ans1]] m [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Step 1: Use the tangent ratio:
</p>

<p>
\[
\tan \theta = \frac{\text{height}}{\text{distance}}
\]
</p>

<p>
\[
\tan {@theta@} = \frac{h}{@d@}
\]
</p>

<p>
Step 2: Solve for height:
</p>

<p>
\[
h = {@d@} \times \tan {@theta@}
\]
</p>

<p>
\[
h = {@h_round@} m
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Distance from object */
d: rand(10) + 10;

/* Angle (choose standard angles for clean values) */
angles: [30,45,60];
theta: angles[rand(3)];

/* Height calculation */
h: if theta=30 then d/sqrt(3)
   elseif theta=45 then d
   else d*sqrt(3);

/* Rounded */
h_round: float(round(h*100)/100);

/* mistakes */
m1: d*tan(theta);   /* incorrect use */
m2: d*cos(theta);
m3: d*sin(theta);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@h_round@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Trigonometry in real-life</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to determine the height of a building using angle of elevation and the tangent ratio.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>h_round</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>h_round</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>mistake1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Use tangent: 
tan 𝜃 = \frac {height}{distance}</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text>Use: tan θ = \frac {d}{h}</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>694510577</deployedseed>
    <deployedseed>1036570820</deployedseed>
    <deployedseed>644137346</deployedseed>
    <deployedseed>178913173</deployedseed>
    <deployedseed>1653499507</deployedseed>
    <deployedseed>728611637</deployedseed>
    <deployedseed>1034472301</deployedseed>
    <deployedseed>1542308450</deployedseed>
    <deployedseed>1813291236</deployedseed>
    <deployedseed>519959698</deployedseed>
    <deployedseed>1724844656</deployedseed>
    <deployedseed>484043352</deployedseed>
    <deployedseed>940361341</deployedseed>
    <deployedseed>1143423595</deployedseed>
    <deployedseed>919992326</deployedseed>
    <deployedseed>524823649</deployedseed>
    <deployedseed>1452857397</deployedseed>
    <deployedseed>708585427</deployedseed>
    <deployedseed>973367153</deployedseed>
    <deployedseed>2061025975</deployedseed>
    <deployedseed>1352680960</deployedseed>
    <deployedseed>855932687</deployedseed>
    <deployedseed>180773563</deployedseed>
    <deployedseed>2122735753</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>h_round</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
