<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the Area of a Regular Pentagon</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the area of the figure below, centered at O, and round your answer to the nearest \(\textbf{whole number}\).</p>
<p>
<!DOCTYPE html>
<html>
<head>
  <meta charset="utf-8" />
  <script src="https://jsxgraph.org/distrib/jsxgraphcore.js"></script>
  <link rel="stylesheet" href="https://jsxgraph.org/distrib/jsxgraph.css" />
</head>
<body>
  <div id="box" style="width: 500px; height: 500px;"></div>
  <script>
    var board = JXG.JSXGraph.initBoard('box', {
      boundingbox: [-40, 40, 40, -40],
      axis: false, showCopyright: false
    });

    // Centre point O
    var O = board.create('point', [0, 0], {name: 'O', fixed: true,visible: false});

    // Draw pentagon vertices manually using radius 25.5 cm
    var A = board.create('point', [0, 25.5], {visible: false});
    var B = board.create('point', [24.2, 7.9], {visible: false});
    var C = board.create('point', [14.8, -20.6], {visible: false});
    var D = board.create('point', [-14.8, -20.6], {visible: false});
    var E = board.create('point', [-24.2, 7.9], {visible: false});

    // Draw pentagon
    var pentagon = board.create('polygon', [A, B, C, D, E], {
      borders: {strokeColor: 'black'},
      fillColor: 'none'
    });

    // Draw radius line (25.5 cm) from centre to vertex C
    var radiusLine = board.create('line', [O, C], {straightFirst: false, straightLast: false, strokeColor: 'black'});
    

    // Draw radius line (25.5 cm) from centre to vertex B
    var radiusLine = board.create('line', [O, B], {straightFirst: false, straightLast: false, strokeColor: 'black'});
    board.create('text', [-1, -10, "{#slant#} cm"], {fontSize: 12,fixed: true});
    board.create('text', [-0.2, 2, "O"], {fontSize: 12,fixed: true});

    // Draw side base (approx 30 cm) between B and E
    var base = board.create('segment', [B, C], {strokeColor: 'black'});
    board.create('text', [B.X() , B.Y() -15, "{#base_length#} cm"], {fontSize: 12,fixed: true});

    // Draw perpendicular height from O to base BC
     // Draw perpendicular height from O to base BC
    var midpoint = board.create('midpoint', [B, C], {visible: false});
    var heightLine = board.create('segment', [O, midpoint], {strokeColor: 'black', dash: 1});

    // Draw right angle marker without alpha symbol
    var rightAngle = board.create('nonreflexangle', [B, midpoint, O], {
      radius: 3,
      type: 'square',
      name: '',fixed: true
    });

  </script>
</body>
</html>

</p>


<p>\(\text{Area} =\)[[input:ans1]]\(\text{cm}^2\) [[feedback:prt1]]</p>
<p hidden>[[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The figure above is a regular isosceles pentagon</p>
<p>Each triangle has</p>
<p><ul><li>
\(\text{Base} = \frac{{@base_length @}}{2}={@half_base@} \,\text{cm}\) (side of the pentagon)
</li>
<li>
\(\text{Slant height }= {@slant@}\, \text{ cm}\) (distance from centre O to a vertex).
</li></ul></p>
<p>To find the area of the above pentagon the following steps are followed</p>
<p>\(\textbf{ Step 1. Find the area of one triangle.}\)</p>
<p>Each triangle is isosceles, but we can split it into two right-angled triangles to find its height.</p>
<p>Base\(=  {@base_length @} \,\text{cm}\)</p>
<p>Hypotenuse \(= {@slant@}\, \text{ cm}\)</p>
<p>height \(=\) ?</p>
<p>Use Pythagoras’ theorem to find the perpendicular height (h).</p>
<p>
\(\begin{align}
{@slant@}^2 &=  h^2+({@half_base@})^2 \\
{@slant1@}&=h^2 + {@half_base1@} \\
h^2 &= {@slant1@}-{@half_base1@}  \\
h&=\sqrt{{@height_1@}}\\
&={@height@} \\
&={@height3@} \\

& \approx {@height2@}

\end{align}\)
</p>
<p>\(\text{Step 2. Calculate area of one triangle.}\)</p>
<p>
\(
\begin{align} \text{Area} &= \frac{1}{2} \times \text{base} \times \text{height} \\
&= \frac{1}{2} \times {@base_length@} \times {@height2@} \\
&=\frac{1}{2} \times {@area_1@} \\
&={@area_triangle1@}\, \text{cm}^2
\end{align}\)
</p>
<p>\(\textbf{Step 3. Calculate area of the pentagon.}\)</p>
<p>
\(
\begin{align} \text{Area}&= 5 \times \text{Area of one triangle}\\
&= 5 \times {@area_triangle1@}\\
&= {@area_pentagon@}\\
&={@ta@}\,\text{cm}^2
\end{align}\)
</p>
<p>Area of the pentagon \({@ta@}\,\text{cm}^2\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Title: Random Pentagon Area Calculator with Expanded Range */

/* Generate random base_length divisible by 10 between 20 and 200 (20 options) */
base_length : 10 * (rand(25) + 2); /* random(19)+2 ensures minimum 20 */

/* Calculate maximum possible slant height (less than base_length, greater than half_base) */
half_base : base_length / 2;
max_slant : base_length - 5; /* ensure slant < base_length */

/* If half_base is greater than or equal to max_slant, adjust max_slant to be at least half_base + 5 */
if half_base >= max_slant then max_slant : half_base + 5;

/* Generate random slant height divisible by 5, greater than half_base but less than base_length */
slant : 5 * (floor(half_base/5) + rand(floor((max_slant - half_base)/5)) + 1);

slant1:slant^2;
half_base1:half_base^2;
height_1:slant1-half_base1;

/* Calculate height using Pythagoras theorem */
height : sqrt(height_1);
height3:float(height);
height2:decimalplaces(height3,1);


/* Calculate area of one triangle */
area_1:base_length *height2;
area_triangle : (1/2) * base_length * height2;
area_triangle1:decimalplaces(area_triangle,1);

/* Calculate area of pentagon */
area_pentagon : 5 * area_triangle;
ta1:decimalplaces(area_pentagon,1);
ta:round(area_pentagon );

/* ---------- Randomized mistake generation ---------- */

/* Generate random mistake number between 1 and 4 */
mistake_number : rand(4) + 1;

/* Define mistake answers */
if mistake_number = 1 then (
    mistake_answer : round((1/2) * base_length * slant), /* using slant as height */
    feedback : "Great effort! Remember, slant height is the side of the triangle, not its perpendicular height. Try using Pythagoras to find the height correctly."
)
else if mistake_number = 2 then (
    mistake_answer : round((1/2) * base_length * half_base), /* using half base as height */
    feedback : "Nice try. Half the base helps form the right-angled triangle, but it is not the perpendicular height itself. Please review the triangle splitting method."
)
else if mistake_number = 3 then (
    mistake_answer : round(area_triangle), /* forgetting to multiply by 5 */
    feedback : "Almost there! You calculated the area of one triangle perfectly. Just remember to multiply by 5 to get the total pentagon area."
)
else if mistake_number = 4 then (
    mistake_answer : round((base_length * slant) / 2), /* using (base*slant)/2 directly */
    feedback : "You’re working hard, which is great! However, slant height is not perpendicular to the base. Always use the perpendicular height in the area formula."
);

/* Display outputs for debugging or STACK usage */
/*display(base_length, slant, half_base, height, area_triangle, correct_answer, mistake_number, mistake_answer, feedback);*/
]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>The base \(={@base_length@}\)</p>
<p>The slant height \(={@slant@}\)</p>
<p>The area of the pentagon \(={@ta@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
