<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Expanding a perfect square expression</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Expand and simplify:
</p>

<p>
\[
(x + {@a@})^2
\]
</p>

<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
\[
(x + {@a@})^2 = (x + {@a@})(x + {@a@})
\]
</p>

<p>
Multiply each term:
</p>

<p>
\[
x(x + {@a@}) + {@a@}(x + {@a@})
\]
</p>

<p>
\[
= x^2 + {@a@}x + {@a@}x + {@a*a@}
\]
</p>

<p>
Combine like terms:
</p>

<p>
\[
= x^2 + {@2*a@}x + {@a*a@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random constant: 2 to 41 (40 possible values) */
a: rand(40) + 2;

/* Expression */
expr: (x + a)^2;

/* Expanded form */
expanded: expand(expr);

/* Common mistakes */
m1: x^2 + a^2;          /* missing middle term */
m2: x^2 + 2*a + a^2;    /* missing x in middle term */
m3: x^2 + a*x + a^2;    /* incorrect coefficient */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@expr@}, {@expanded@}, a = {@a@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Quadratic identities – perfect squares</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to apply the perfect square identity to expand and simplify a quadratic expression of the form (x + a)^2.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>expanded</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>expanded</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>m1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>You are missing the middle term.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>m2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>The middle term must include x.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>m3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>The coefficient of x is incorrect.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>Use the identity: (𝑥+𝑎)^2=𝑥^2+2𝑎𝑥+𝑎^2.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>553357896</deployedseed>
    <deployedseed>528665757</deployedseed>
    <deployedseed>178631264</deployedseed>
    <deployedseed>837274260</deployedseed>
    <deployedseed>1070591849</deployedseed>
    <deployedseed>858943511</deployedseed>
    <deployedseed>582783127</deployedseed>
    <deployedseed>508887027</deployedseed>
    <deployedseed>328164897</deployedseed>
    <deployedseed>453439856</deployedseed>
    <deployedseed>181151367</deployedseed>
    <deployedseed>1504839935</deployedseed>
    <deployedseed>531230867</deployedseed>
    <deployedseed>415289008</deployedseed>
    <deployedseed>658799336</deployedseed>
    <deployedseed>2133740763</deployedseed>
    <deployedseed>61197749</deployedseed>
    <deployedseed>771522115</deployedseed>
    <deployedseed>1863515806</deployedseed>
    <deployedseed>2020438389</deployedseed>
    <deployedseed>1270276233</deployedseed>
    <deployedseed>1025302412</deployedseed>
    <deployedseed>1586594128</deployedseed>
    <deployedseed>725326891</deployedseed>
    <deployedseed>874950057</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>expanded</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
