<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Recognising a perfect square trinomial</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Write the following expression as a perfect square:</p>

<p>
\[
{@trin@}
\]
</p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Compare the expression with:
</p>

<p>
\[
(x + a)^2 = x^2 + 2ax + a^2
\]
</p>

<p>
\[
(x - a)^2 = x^2 - 2ax + a^2
\]
</p>

<p>
The constant term is {@const@} = ({@a@})^2,
so the number inside the bracket is {@a@}.
</p>

<p>
The middle term is {@mid@}x,
which matches \(\pm 2ax\).
</p>

<p>
Therefore,
</p>

<p>
\[
\boxed{{@correct@}}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Random constant: 2 to 41 */
a: rand(40) + 2;

/* Random sign */
s: rand([1,-1]);

/* Signed values */
sa: s*a;
mid: 2*sa;
const: a^2;

/* Trinomial */
trin: x^2 + mid*x + const;

/* Correct perfect square */
correct: (x + sa)^2;

/* Common mistakes */
m1: x + sa;
m2: (x - sa)^2;
m3: (x + const)^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@trin@}, {@correct@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Quadratic identities – perfect squares</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to recognise a quadratic trinomial as a perfect square and write it in the form (x ± a)^2 by matching the middle term and constant term to the perfect square identity.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>correct</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>correct</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>m1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Check the sign of the middle term.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>m2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>Check the sign of the middle term.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>m3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>The number inside the bracket is {@a@}, not {@const@}.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>Compare with: (𝑥±𝑎)^2=𝑥^2±2𝑎𝑥+𝑎^2.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>129389200</deployedseed>
    <deployedseed>1229908122</deployedseed>
    <deployedseed>430680843</deployedseed>
    <deployedseed>704201932</deployedseed>
    <deployedseed>1072397286</deployedseed>
    <deployedseed>1084666352</deployedseed>
    <deployedseed>1460843158</deployedseed>
    <deployedseed>1249373945</deployedseed>
    <deployedseed>417342285</deployedseed>
    <deployedseed>370676427</deployedseed>
    <deployedseed>841852821</deployedseed>
    <deployedseed>281366446</deployedseed>
    <deployedseed>1520487520</deployedseed>
    <deployedseed>1207201491</deployedseed>
    <deployedseed>1287965774</deployedseed>
    <deployedseed>959071129</deployedseed>
    <deployedseed>62338219</deployedseed>
    <deployedseed>1614850460</deployedseed>
    <deployedseed>1752899836</deployedseed>
    <deployedseed>1110430050</deployedseed>
    <deployedseed>917105986</deployedseed>
    <deployedseed>969569724</deployedseed>
    <deployedseed>350870581</deployedseed>
    <deployedseed>2082610992</deployedseed>
    <deployedseed>820286218</deployedseed>
    <deployedseed>1842144757</deployedseed>
    <deployedseed>1859048539</deployedseed>
    <deployedseed>1688686827</deployedseed>
    <deployedseed>1356782150</deployedseed>
    <deployedseed>981406327</deployedseed>
    <deployedseed>364836513</deployedseed>
    <deployedseed>1662736470</deployedseed>
    <deployedseed>896253996</deployedseed>
    <deployedseed>1702906688</deployedseed>
    <deployedseed>1218558502</deployedseed>
    <deployedseed>410092232</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>correct</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
