<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Applying product and quotient rules of logarithms</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Given that:
</p>

<p>
\[
\log a = {@loga@}, \quad \log b = {@logb@}, \quad \log c = {@logc@}
\]
</p>

<p>
Find:
\[
\log\left(\frac{ab}{c}\right)
\]
</p>

<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
We apply the logarithm rules step by step.
</p>

<p>
Step 1: Use the product rule:
</p>

<p>
\[
\log(ab) = \log a + \log b
\]
</p>

<p>
Step 2: Use the quotient rule:
</p>

<p>
\[
\log\left(\frac{ab}{c}\right) = \log(ab) - \log c
\]
</p>

<p>
Substitute:
</p>

<p>
\[
= {@loga@} + {@logb@} - {@logc@}
\]
</p>

<p>
\[
= {@ans@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* random logs */
loga: rand(20) + 1;
logb: rand(20) + 1;
logc: rand(20) + 1;

/* expression */
ans: loga + logb - logc;

/* mistakes */
m1: loga + logb + logc;
m2: loga - logb - logc;
m3: loga * logb - logc;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ans@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Logarithms of products and ratios</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to apply both the product and quotient rules of logarithms to simplify a logarithmic expression.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ans</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>adding everything</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Remember: division becomes subtraction, not addition.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>wrong subtraction</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>Apply the rules carefully: product first, then quotient.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text>log(ab/c​)=loga+logb−logc</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>750904512</deployedseed>
    <deployedseed>520102447</deployedseed>
    <deployedseed>1546303773</deployedseed>
    <deployedseed>1882326279</deployedseed>
    <deployedseed>1933586453</deployedseed>
    <deployedseed>1782331771</deployedseed>
    <deployedseed>2010983497</deployedseed>
    <deployedseed>886872448</deployedseed>
    <deployedseed>1774479237</deployedseed>
    <deployedseed>1535357027</deployedseed>
    <deployedseed>2050585722</deployedseed>
    <deployedseed>1596035393</deployedseed>
    <deployedseed>995490081</deployedseed>
    <deployedseed>287411137</deployedseed>
    <deployedseed>570261726</deployedseed>
    <deployedseed>1291008202</deployedseed>
    <deployedseed>1081175663</deployedseed>
    <deployedseed>542391146</deployedseed>
    <deployedseed>1140877008</deployedseed>
    <deployedseed>1383642038</deployedseed>
    <deployedseed>451032416</deployedseed>
    <deployedseed>349270067</deployedseed>
    <deployedseed>1456806817</deployedseed>
    <deployedseed>1907173337</deployedseed>
    <deployedseed>51279824</deployedseed>
    <deployedseed>720930110</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ans</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
