<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Multiplication of vectors by a scalar</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Given the vector \( \mathbf{a} = \begin{pmatrix} {@ax@} \\ {@ay@} \end{pmatrix} \).</p>

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10, 10, 10, -10],
    axis: true,
    showCopyright: false,
    defaultAxes: {
        x: {ticks: {distance: 1, minorTicks:0}},
        y: {ticks: {distance: 1, minorTicks:0}}
    },
    grid: {gridX:1, gridY:1}
});

/* Origin */
var O = board.create('point', [0,0], {name:'O', fixed:true, size:1, visible:false});

/* Vector a */
var A = board.create('point', [{#ax#}, {#ay#}], {name:'A', fixed:true, size:1, visible:false});
var vecA = board.create('arrow', [O, A], {
    strokeColor:'blue', 
    strokeWidth:3, 
    name:'a', 
    withLabel:true,
    label: {position: 'top', offset:[0,10], color:'blue'}
});
[[/jsxgraph]]

<p>Calculate the following scalar multiplications. Enter your answers as column vectors.</p>

<p><b>(a)</b> Find \( {@k1@}\mathbf{a} \).</p>
<div>\( {@k1@}\mathbf{a} = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<br>

<p><b>(b)</b> Find \( {@k2@}\mathbf{a} \).</p>
<div>\( {@k2@}\mathbf{a} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>
<br>

<p><b>(c)</b> Find \( \frac{1}{2}\mathbf{a} \).</p>
<div>\( \frac{1}{2}\mathbf{a} = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>
<br>

<p><b>(d)</b> Find \( -\frac{3}{2}\mathbf{a} \) (or \( -1.5\mathbf{a} \)).</p>
<div>\( -1.5\mathbf{a} = \) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p>To multiply a vector by a scalar \( k \), you multiply <b>each component</b> of the vector by \( k \).</p>

\[ k \begin{pmatrix} x \\ y \end{pmatrix} = \begin{pmatrix} k \times x \\ k \times y \end{pmatrix} \]

<p><b>Part (a):</b> Multiply by {@k1@}:</p>
\[ {@k1@} \begin{pmatrix} {@ax@} \\ {@ay@} \end{pmatrix} = \begin{pmatrix} {@k1@} \times {@ax@} \\ {@k1@} \times {@ay@} \end{pmatrix} = \begin{pmatrix} {@ta1[1][1]@} \\ {@ta1[2][1]@} \end{pmatrix} \]

<p><b>Part (b):</b> Multiply by {@k2@} (Negative reverses the direction):</p>
\[ {@k2@} \begin{pmatrix} {@ax@} \\ {@ay@} \end{pmatrix} = \begin{pmatrix} {@k2@} \times {@ax@} \\ {@k2@} \times {@ay@} \end{pmatrix} = \begin{pmatrix} {@ta2[1][1]@} \\ {@ta2[2][1]@} \end{pmatrix} \]

<p><b>Part (c):</b> Multiply by \(\frac{1}{2}\) (Halves the length):</p>
\[ \frac{1}{2} \begin{pmatrix} {@ax@} \\ {@ay@} \end{pmatrix} = \begin{pmatrix} {@ax@} / 2 \\ {@ay@} / 2 \end{pmatrix} = \begin{pmatrix} {@ta3[1][1]@} \\ {@ta3[2][1]@} \end{pmatrix} \]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Define Base Vector a */
/* Using even numbers (-6 to 6) to make division clean */
ax: rand_with_step(-6, 6, 2);
ay: rand_with_step(-6, 6, 2);

/* Avoid zero vector */
ax: if (ax=0 and ay=0) then 2 else ax;

/* Define Vector a as a Matrix */
vec_a: matrix([ax], [ay]);

/* 2. Define Scalars for the 4 parts */
k1: rand([2, 3, 4]);           /* Part A: Positive Integer */
k2: rand([-2, -3]);            /* Part B: Negative Integer */
k3: 1/2;                       /* Part C: Fraction (Halving) */
k4: -3/2;                      /* Part D: Negative Fraction */

/* 3. Calculate Correct Answers (Renamed to 'ta' for Teacher Answer) */
ta1: k1 * vec_a;
ta2: k2 * vec_a;
ta3: k3 * vec_a;
ta4: k4 * vec_a;

/* 4. Wrong Answers (Common Misconceptions) for Feedback */

/* Error: Adding scalar to components instead of multiplying */
wa_add_1: matrix([ax + k1], [ay + k1]);

/* Error: Multiplying only top or only bottom */
wa_partial_x: matrix([ax * k1], [ay]);
wa_partial_y: matrix([ax], [ay * k1]);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}
{@ta2@}
{@ta3@}
{@ta4@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>ta2</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>matrix</type>
      <tans>ta3</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>matrix</type>
      <tans>ta4</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2078794972</deployedseed>
    <deployedseed>873975579</deployedseed>
    <deployedseed>1250281455</deployedseed>
    <deployedseed>1873794114</deployedseed>
    <deployedseed>1321555689</deployedseed>
    <deployedseed>393544997</deployedseed>
    <deployedseed>423340651</deployedseed>
    <deployedseed>741078104</deployedseed>
    <deployedseed>976460332</deployedseed>
    <deployedseed>1192157725</deployedseed>
    <deployedseed>1026723938</deployedseed>
    <deployedseed>458171403</deployedseed>
    <deployedseed>1418919905</deployedseed>
    <deployedseed>358464162</deployedseed>
    <deployedseed>874949233</deployedseed>
    <deployedseed>448086846</deployedseed>
    <deployedseed>1960118935</deployedseed>
    <deployedseed>152987443</deployedseed>
    <deployedseed>1387628376</deployedseed>
    <deployedseed>1278428243</deployedseed>
  </question>
</quiz>
