<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Similarity and Enlargement Cylinders _Surface Area, Length, and Volume Relationships (copy)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p><strong>Two cylinders, P and Q, are mathematically similar.</strong></p>

[[jsxgraph]]
var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [-60, -60, 90, 90],
    axis: false,
    visible: false,
    showCopyright: false,
    showNavigation: true,
    keepAspectRatio: true
});

/* Cylinder P */
var radiusP = 10;
var heightP = 60;

// Bottom ellipse (P)
board.create('ellipse', [
    [-25, 0],
    [-25 + radiusP, 0],
    [-25, 3]
], {strokeColor: '#000', fillColor: '#ccc', fillOpacity: 0.3, fixed: true});

// Top ellipse (P)
board.create('ellipse', [
    [-25, heightP],
    [-25 + radiusP, heightP],
    [-25, heightP + 3]
], {strokeColor: '#000', fillColor: '#ccc', fillOpacity: 0.3, fixed: true});

// Side lines (P)
board.create('segment', [[-26.7, 0], [-26.7, heightP]], {strokeColor: '#000', fixed: true});
board.create('segment', [[-13.5, 0], [-13.5, heightP]], {strokeColor: '#000', fixed: true});

// Labels (P)
board.create('text', [-25 + radiusP + 3, heightP / 2, 'Length = {#L_P#} cm'], {fixed: true});
board.create('text', [-20, heightP + 5, '<b>Cylinder P</b>'], {fixed: true});

/* Cylinder Q */
var radiusQ = 15;
var heightQ = 80;

// Bottom ellipse (Q)
board.create('ellipse', [
    [20, 0],
    [20 + radiusQ, 0],
    [20, 3]
], {strokeColor: '#000', fillColor: '#cde', fillOpacity: 0.3, fixed: true});

// Top ellipse (Q)
board.create('ellipse', [
    [20, heightQ],
    [20 + radiusQ, heightQ],
    [20, heightQ + 3]
], {strokeColor: '#000', fillColor: '#cde', fillOpacity: 0.3, fixed: true});

// Side lines (Q)
board.create('segment', [[18.3, 0], [18.3, heightQ]], {strokeColor: '#000', fixed: true});
board.create('segment', [[35, 0], [35, heightQ]], {strokeColor: '#000', fixed: true});

// Labels (Q)
board.create('text', [20 + radiusQ + 3, heightQ / 2, 'Length = ?'], {fixed: true});
board.create('text', [25, heightQ + 5, '<b>Cylinder Q</b>'], {fixed: true});
[[/jsxgraph]]

<hr>
<ul>
  <li>Total surface area of cylinder <strong>P</strong> = {@ SA_P @} cm²</li>
  <li>Total surface area of cylinder <strong>Q</strong> = {@ SA_Q @} cm²</li>
  <li>Length of cylinder <strong>P</strong> = {@ L_P @} cm</li>
</ul>

<p><strong>(a)</strong> Work out the length of cylinder Q.</p>
<p>Give your answer in cm.</p>

<p>Length of cylinder Q = [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<hr>

<p>The volume of cylinder <strong>P</strong> is {@ V_P @} cm³.</p>
<p><strong>(b)</strong> Work out the volume of cylinder Q.</p>
<p>Give your answer in cm³.</p>

<p>Volume of cylinder Q = [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

<hr>

<details>
  <summary><strong>Hints (click to expand)</strong></summary>
  <ul>
    <li>For similar solids: area scale factor = (length scale factor)²</li>
    <li>Volume scale factor = (length scale factor)³</li>
    <li>Length scale factor = √(area scale factor)</li>
  </ul>
</details>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h5><strong>Worked Solution</strong></h5>

<p><strong>Lets first look at what we are given:</strong></p>
<ul>
  <li>Total surface area of cylinder <strong>P</strong> = {@ SA_P @} cm²</li>
  <li>Total surface area of cylinder <strong>Q</strong> = {@ SA_Q @} cm²</li>
  <li>Length of cylinder <strong>P</strong> = {@ L_P @} cm</li>
  <li>Volume of cylinder <strong>P</strong> = {@ V_P @} cm³</li>
</ul>

<h5><strong>(a)</strong> Find the length of cylinder Q</h5>

<ol>
  <li>Find the area (surface area) scale factor:<br>
  \( \text{Area scale factor} = \dfrac{S_Q}{S_P} = \dfrac{{@ SA_Q @}}{{@ SA_P @}} = {@ scale_factor_area @} \)</li><br>

  <li>Find the linear (length) scale factor:<br>
  Since area scale factor = (length scale factor)²,<br>
  \( \text{Length scale factor} = \sqrt{{@ scale_factor_area @}} = {@ scale_factor_length @} \)</li><br>

  <li>Find the length of Q:<br>
  \( L_Q = {@ L_P @} \times {@ scale_factor_length @} = {@ L_Q @}\ \text{cm} \)</li>
</ol>

<p><strong>Answer (a):</strong> {@ ta_ans1 @} cm</p>

<hr>

<h5><strong>(b)</strong> Find the volume of cylinder Q</h5>

<ol>
  <li>Find the volume scale factor:<br>
  \( \text{Volume scale factor} = (\text{length scale factor})^3 = {@ scale_factor_length @}^3 = {@ scale_factor_volume @} \)</li><br>

  <li>Find the volume of Q:<br>
  \( V_Q = {@ V_P @} \times {@ scale_factor_volume @} = {@ V_Q @}\ \text{cm}^3 \)</li>
</ol>

<p><strong>Answer (b):</strong> {@ ta_ans2 @} cm³</p>

<hr>

<p><em>Summary:</em></p>
<ul>
  <li><strong>Length of cylinder Q:</strong> {@ ta_ans1 @} cm</li>
  <li><strong>Volume of cylinder Q:</strong> {@ ta_ans2 @} cm³</li>
</ul>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomised variables */

/* Surface areas*/
SA_P : 10*rand_with_step(5,12,1);              /* total surface area of cylinder P (e.g. 50–120 cm²) */
scale_factor_area : rand([4,9,16,25]);  /* area scale factor between 4 and 25 */
SA_Q : SA_P * scale_factor_area;

/* Linear scale factor */
scale_factor_length : sqrt(scale_factor_area);

/* Lengths */
L_P : rand_with_step(3,6,1);                   /* length of cylinder P (3–6 cm) */
L_Q : L_P * scale_factor_length;



/* Volumes */
V_P : 10*rand_with_step(8,15,1);               /*volume of cylinder P (80–150 cm³) */
scale_factor_volume : scale_factor_length^3;
V_Q : V_P * scale_factor_volume;

/* Answers */
ta_ans1 : decimalplaces(L_Q,2);    /* length of Q */
ta_ans2 : decimalplaces(V_Q,2);    /* volume of Q */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<br>{@ta_ans1 @}
<br>{@ta_ans2 @}]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta_ans1 </tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta_ans2 </tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta_ans1 </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta_ans2 </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
