<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Enlargement of a Triangle: Coordinates, Side Lengths, and Scale Factor</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
You have discovered two triangles on the board: the small one, \(ABC\), and the big one, \(AB'C'\), which is an enlargement of the first with centre at \(A\).
</p>
<br>
[[jsxgraph]]
var xmin = {#xmin#};
var xmax = {#xmax#};
var ymin = {#ymin#};
var ymax = {#ymax#};

var board = JXG.JSXGraph.initBoard('jxgbox', {
    boundingbox: [xmin, ymax, xmax, ymin],  
    axis:true,
    showCopyright:false,
    showNavigation:false,
    keepAspectRatio:true
});

// Base triangle coordinates from STACK
var xA = {#xA#};
var yA = {#yA#};
var xB = {#xB#};
var yB = {#yB#};
var xC = {#xC#};
var yC = {#yC#};

// Original triangle ABC
var A = board.create('point', [xA, yA], {name:'A', size:2, fixed:true});
var B = board.create('point', [xB, yB], {name:'B', size:2, fixed:true});
var C = board.create('point', [xC, yC], {name:'C', size:2, fixed:true});
var triABC = board.create('polygon', [A,B,C], {
    borders:{strokeColor:"green", strokeWidth:2},   // green border
    fillColor:"lightgreen",                        // green fill
    fillOpacity:0.4
});

// Scale factor from STACK
var k = {#scale_factor#};

// Image triangle A'B'C' (A' = A)
var xB1 = {#xB1#};
var yB1 = {#yB1#};
var xC1 = {#xC1#};
var yC1 = {#yC1#};

var B1 = board.create('point', [xB1,yB1], {name:"B'", size:2, color:"red", fixed:true});
var C1 = board.create('point', [xC1,yC1], {name:"C'", size:2, color:"red", fixed:true});
var triA1B1C1 = board.create('polygon',[A,B1,C1], {
    borders:{strokeColor:"red", strokeWidth:2, dash:2}, // red border
    fillColor:"pink",                                   // red fill
    fillOpacity:0.3
});

// Optional: show rays from A to dilated points
board.create('line', [A,B1], {straightFirst:true, straightLast:true, strokeColor:"gray", dash:1});
board.create('line', [A,C1], {straightFirst:true, straightLast:true, strokeColor:"gray", dash:1});
[[/jsxgraph]]

<p></p>

<h4>Task 1: Coordinates of the Image Triangle</h4>
<p>Write the coordinates of <strong>B'</strong> and <strong>C'</strong>:</p>
<p> \( B' = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]] </p>
<p> \( C' = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]] </p>

<h4>Task 2: Compute Side Lengths</h4>
<p>Compute the lengths \(2dp\) of side <strong>AB</strong> in triangle ABC and <strong>AB'</strong> in triangle AB'C'.</p>

<p>Use the distance formula: <em>\(\displaystyle d = \sqrt( (x_₂ - x_₁)^2 + (y_₂ - y_₁)^2) \)</em></p>

<ul>
  <li>Length of side AB = [[input:ans2a]] [[validation:ans2a]] [[feedback:prt2a]]</li>
<p></p>
  <li>Length of side AB' = [[input:ans2b]] [[validation:ans2b]] [[feedback:prt2b]]</li>
</ul>

<h4>Task 3: What is the linear scale factor ? </h4>
<p> \( SF = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]] </p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/****************************************************************/
stack_include_contrib("prooflib.mac");

/* Random scale factor */
scale_factor : 2 + rand(2);   /* picks either 2 or 3 */

/* Base triangle coordinates (randomised but structured) */
xA : 0;  
yA : 0;

xB : 2 + rand(3);   /* between 2 and 4 */
yB : 0;

xC : 1 + rand(3);   /* between 1 and 3 */
yC : 1 + rand(3);   /* between 1 and 3 */

/* Image triangle after enlargement from A */
xB1 : scale_factor*xB;
yB1 : scale_factor*yB;

xC1 : scale_factor*xC;
yC1 : scale_factor*yC;


/* Find extremes of all points (original + image) */
xmin : min(xA, xB, xC, xB1, xC1) - 1;
xmax : max(xA, xB, xC, xB1, xC1) + 1;
ymin : min(yA, yB, yC, yB1, yC1) - 1;
ymax : max(yA, yB, yC, yB1, yC1) + 1;

/* -------------------------------*/

/* STACK variables for Task 1 */
ta1: [xB1, yB1]; /* coordinates of B' */
ta2: [xC1, yC1]; /* coordinates of C' */

/* Task 2: Side lengths using distance formula */
AB_length  : decimalplaces(sqrt((xB - xA)^2 + (yB - yA)^2),2);
AB1_length : decimalplaces(sqrt((xB1 - xA)^2 + (yB1 - yA)^2),2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>1</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[x, y]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>[x, y]</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2a</name>
      <type>algebraic</type>
      <tans>AB_length  </tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2b</name>
      <type>algebraic</type>
      <tans>AB1_length </tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>scale_factor </tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2a</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2a</sans>
        <tans>AB_length  </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2a-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2a-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2b</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2b</sans>
        <tans>AB1_length</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2b-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2b-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>scale_factor</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>616842617</deployedseed>
  </question>
</quiz>
