<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Descriptive Statistics Histogram</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>{@lx@}</p><p>{@mean1@}</p><p>{@sd1@}</p><p>{@ds1@}</p><p>[[jsxgraph]]var lx = {#lx#};
var mean = {#mean1#};
var sd = {#sd1#};
var min = Math.floor(lx[0]);
var max = Math.ceil(lx[lx.length - 1]);
var numel = lx.length;
var range = max - min;
var numbars = 10;
var barwidth = range / numbars;
var i = 0;
var tot = 0;
for (i = 0; i <= numel - 1; i++) {
  tot = tot + lx[i]
}
var j = 0;
var relfreq = [0, 0, 0, 0, 0,0,0,0,0,0,0];
for (j = 0; j <= numbars - 1; j++) {
  for (i = 0; i <= numel - 1; i++) {
    if (lx[i] > min + j * barwidth) {
      if (lx[i] <= min + (j + 1) * barwidth) {
        relfreq[j] = relfreq[j] + 1;
      }
    }
  }
}
var ymax = 0;
for (j = 0; j <= numbars - 1; j++) {
  if (relfreq[j] / barwidth >= ymax) {
    ymax = relfreq[j] / barwidth;
  }
}
var board = JXG.JSXGraph.initBoard(divid, {
  boundingbox: [-2, ymax + 3, max + 10, -5],
  axis: true, showCopyright: false
});
var p = [];
p[0] = [0, 0];
p[1] = [0, 0];
p[2] = [0, 0];
p[3] = [0, 0];
for (i = 0; i < numbars; i++) {
  p[0] = [min + (i) * barwidth, 0];
  p[1] = [min + (i) * barwidth, relfreq[i] / barwidth]
  p[2] = [min + (i + 1) * barwidth, 0];
  p[3] = [min + (i + 1) * barwidth, relfreq[i] / barwidth]
  var v1 = board.create('line', [p[0], p[1]], {
    straightFirst: false,
    straightLast: false,
    strokeWidth: 2
  })
  var v2 = board.create('line', [p[2], p[3]], {
    straightFirst: false,
    straightLast: false,
    strokeWidth: 2
  })
  var h1 = board.create('line', [p[1], p[3]], {
    straightFirst: false,
    straightLast: false,
    strokeWidth: 2
  })
}
[[/jsxgraph]]</p><p>[[input:ans1]] [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text/>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2018080600</text>
    </stackversion>
    <questionvariables>
      <text>mean1:rand(5)+10;
sd1:rand(3)+2;
ds1:rand(900)+100;
list:makelist(float(quantile_normal(rand(1.0),mean1,sd1)),i,1,ds1);
lx:sort(list);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>{@list@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1.0000000</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0.0000000</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
