<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Interactive: Frequency table for the time taken and a Histogram and frequency polygon</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p><br>(a) A teacher used the school’s Learning Management System (LMS) to record the time, in minutes, that students took to finish a Math CAT assessment.<br>Complete the table by calculating the frequency density for each time interval.
</p>
<p><br></p>

<table style="width:500px; border:1px solid black; text-align:left">
  <tr style="height:70px; border-bottom:1px solid black">
    <th style="padding:10px"width=40%>Range \([min]\)</th>
    <th style="padding:10px"width=30%>Frequency</th>
    <th style="padding:10px"width=30%>Frequency Density</th>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black"">
    <td style="padding:10px">\(0 < t \leq {@first(second(data))@}\)</td>
    <td style="padding:10px; vertical-align:middle">{@third(first(data))*first(first(widthList))@}</td>
    <td style="padding:10px; vertical-align:bottom">[[input:ans1]]</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black"">
    <td style="padding:10px">\({@first(second(data))@} < t \leq {@first(third(data))@}\)</td>
    <td style="padding:10px; vertical-align:middle">{@third(second(data))*first(second(widthList))@}</td>
    <td style="padding:10px; vertical-align:bottom">[[input:ans2]]</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black"">
    <td style="padding:10px">\({@first(third(data))@} < t \leq {@first(fourth(data))@}\)</td>
    <td style="padding:10px; vertical-align:middle">{@third(third(data))*first(third(widthList))@}</td>
    <td style="padding:10px; vertical-align:bottom">[[input:ans3]]</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black"">
    <td style="padding:10px">\({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\)</td>
    <td style="padding:10px; vertical-align:middle">{@third(fourth(data))*first(fourth(widthList))@}</td>
    <td style="padding:10px; vertical-align:bottom">[[input:ans4]]</td>
  </tr>
  <tr style="height:70px; border-bottom:1px solid black"">
    <td style="padding:10px">\({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\)</td>
    <td style="padding:10px; vertical-align:middle">{@third(fifth(data))*first(fifth(widthList))@}</td>
    <td style="padding:10px; vertical-align:bottom">[[input:ans5]]</td>
  </tr>
  <tr style="height:70px">
    <td style="padding:10px">\({@first(sixth(data))@} < t \leq 60\)</td>
    <td style="padding:10px; vertical-align:middle">{@third(sixth(data))*(60-first(sixth(data)))@}</td>
    <td style="padding:10px; vertical-align:bottom">[[input:ans6]]</td>
  </tr>
</table>

<p hidden>[[validation:ans1]] [[validation:ans2]] [[validation:ans3]] [[validation:ans4]] [[validation:ans5]] [[validation:ans6]]</p>
<p align="right"></p>
<p>[[feedback:prt1]]</p>

<p><br>(b) By dragging the red points, complete the histogram below to represent the distribution of times taken by students to complete the Math CAT Mark assessment.</p>

<p><br></p>
[[jsxgraph width="750px" height="450px"
  input-ref-ans7="sR1"
  input-ref-ans8="sR2"
  input-ref-ans9="sR3"
  input-ref-ans10="sR4"
  input-ref-ans11="sR5"
  input-ref-ans12="sR6"
]]

JXG.Options.text.useMathJax = true;

// Class ranges [min, max]
var ranges = [
  [0, {#second(first(data))#}],
  [{#first(second(data))#}, {#second(second(data))#}],
  [{#first(third(data))#}, {#second(third(data))#}],
  [{#first(fourth(data))#}, {#second(fourth(data))#}],
  [{#first(fifth(data))#}, {#second(fifth(data))#}],
  [{#first(sixth(data))#}, 60]
];

// Create board
var board = JXG.JSXGraph.initBoard(divid, {
  boundingBox: [-5, 50, 65, -3],
  grid: true,
  showNavigation: false,
  showCopyright: false
});

// Axes
var yAxis = board.create("axis", [[0, 0], [0, 50]], {
  name: "\\(fd\\)",
  withLabel: true
});

var xAxis = board.create("axis", [[0, 0], [40, 0]], {
  name: "\\(t\\,[min]\\)",
  withLabel: true
});

// Vertical guide segments
var segList = [];
for (var i = 0; i < ranges.length; i++) {
  segList.push(
    board.create("segment", [
      [(ranges[i][0] + ranges[i][1]) / 2, 0],
      [(ranges[i][0] + ranges[i][1]) / 2, 40]
    ], {
      fixed: true,
      visible: false
    })
  );
}

// Draggable frequency density points
var pointList = [];
for (i = 0; i < ranges.length; i++) {
  pointList.push(
    board.create("glider", [
      (ranges[i][0] + ranges[i][1]) / 2,
      0,
      segList[i]
    ], {
      size: 4,
      snapToGrid: true,
      showInfoBox: false,
      withLabel: false
    })
  );
}

// Histogram rectangles
var lineListHori = [];
for (i = 0; i < ranges.length; i++) {
  lineListHori.push(
    board.create("parallel", [pointList[i], xAxis], {
      fixed: true,
      visible: false
    })
  );
}

var lineListVert = [
  board.create("parallel", [[ranges[0][0], 0], yAxis], {
    fixed: true,
    visible: false
  })
];

for (i = 0; i < ranges.length; i++) {
  lineListVert.push(
    board.create("parallel", [[ranges[i][1], 0], yAxis], {
      fixed: true,
      visible: false
    })
  );
}

for (i = 0; i < pointList.length; i++) {
  board.create("polygon", [
    [ranges[i][0], 0],
    [ranges[i][1], 0],
    board.create("intersection", [lineListHori[i], lineListVert[i + 1]], { visible: false }),
    board.create("intersection", [lineListHori[i], lineListVert[i]], { visible: false })
  ], {
    fixed: true,
    highlight: false,
    vertices: { visible: false }
  });
}

// -----------------------------
// Frequency polygon (FINAL)
// -----------------------------

// Endpoints on x-axis
var startPoint = board.create("point", [ranges[0][0], 0], {
  fixed: true,
  visible: false
});

var endPoint = board.create("point", [ranges[ranges.length - 1][1], 0], {
  fixed: true,
  visible: false
});

// Midpoints of bars
var midPoints = [];
for (i = 0; i < ranges.length; i++) {
  midPoints.push(
    board.create("point", [
      (ranges[i][0] + ranges[i][1]) / 2,
      function (j) {
        return function () {
          return pointList[j].Y();
        };
      }(i)
    ], {
      fixed: true,
      visible: false
    })
  );
}

// Polygon segments
board.create("segment", [startPoint, midPoints[0]], {
  strokeColor: "black",
  strokeWidth: 3,
  fixed: true
});

for (i = 1; i < midPoints.length; i++) {
  board.create("segment", [midPoints[i - 1], midPoints[i]], {
    strokeColor: "black",
    strokeWidth: 3,
    fixed: true
  });
}

board.create("segment", [midPoints[midPoints.length - 1], endPoint], {
  strokeColor: "black",
  strokeWidth: 3,
  fixed: true
});

// STACK bindings
stack_jxg.bind_point(sR1, pointList[0]);
stack_jxg.bind_point(sR2, pointList[1]);
stack_jxg.bind_point(sR3, pointList[2]);
stack_jxg.bind_point(sR4, pointList[3]);
stack_jxg.bind_point(sR5, pointList[4]);
stack_jxg.bind_point(sR6, pointList[5]);

board.unsuspendUpdate();

[[/jsxgraph]]



<p hidden>[[input:ans7]] [[input:ans8]] [[input:ans9]] [[input:ans10]] [[input:ans11]] [[input:ans12]]</p>
<p hidden>[[validation:ans7]] [[validation:ans8]] [[validation:ans9]] [[validation:ans10]] [[validation:ans11]] [[validation:ans12]]</p>
<p align="right"><br> </p>
<p>[[feedback:prt2]]</p>

<p><br>(c) Estimate the number of students who completed the Math CAT Mark assessment in less than 20 minutes.</p>
<p align="left">[[input:ans13]] [[validation:ans13]]</p>
<p align="right"> </p>
<p>[[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<br>(a)

<p>
Using the relationship
\(
\text{frequency density} = \frac{\text{frequency}}{\text{class width}},
\)
</p>

<p>
the <b>frequency density</b> for each time interval can be calculated by dividing the
number of students by the width of the corresponding time interval.
</p>

<p><br></p>

<table style="width:450px; border:1px solid black; text-align:left">
  <tr style="height:70px; border-bottom:1px solid black">
    <th style="padding:10px">Time taken (min)</th>
    <th style="padding:10px">Number of students</th>
    <th style="padding:10px">Frequency density</th>
  </tr>

  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">\(0 < t \leq {@first(second(data))@}\)</td>
    <td style="padding:10px; vertical-align:middle">
      {@third(first(data))*first(first(widthList))@}
    </td>
    <td style="padding:10px; vertical-align:middle">
      \(\frac{{@third(first(data))*first(first(widthList))@}}{{@first(second(data))@}} = {@tans1@}\)
    </td>
  </tr>

  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">
      \({@first(second(data))@} < t \leq {@first(third(data))@}\)
    </td>
    <td style="padding:10px; vertical-align:middle">
      {@third(second(data))*first(second(widthList))@}
    </td>
    <td style="padding:10px; vertical-align:middle">
      \(\frac{{@third(second(data))*first(second(widthList))@}}{{@first(third(data))-first(second(data))@}} = {@tans2@}\)
    </td>
  </tr>

  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">
      \({@first(third(data))@} < t \leq {@first(fourth(data))@}\)
    </td>
    <td style="padding:10px; vertical-align:middle">
      {@third(third(data))*first(third(widthList))@}
    </td>
    <td style="padding:10px; vertical-align:middle">
      \(\frac{{@third(third(data))*first(third(widthList))@}}{{@first(fourth(data))-first(third(data))@}} = {@tans3@}\)
    </td>
  </tr>

  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">
      \({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\)
    </td>
    <td style="padding:10px; vertical-align:middle">
      {@third(fourth(data))*first(fourth(widthList))@}
    </td>
    <td style="padding:10px; vertical-align:middle">
      \(\frac{{@third(fourth(data))*first(fourth(widthList))@}}{{@first(fifth(data))-first(fourth(data))@}} = {@tans4@}\)
    </td>
  </tr>

  <tr style="height:70px; border-bottom:1px solid black">
    <td style="padding:10px">
      \({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\)
    </td>
    <td style="padding:10px; vertical-align:middle">
      {@third(fifth(data))*first(fifth(widthList))@}
    </td>
    <td style="padding:10px; vertical-align:middle">
      \(\frac{{@third(fifth(data))*first(fifth(widthList))@}}{{@first(sixth(data))-first(fifth(data))@}} = {@tans5@}\)
    </td>
  </tr>

  <tr style="height:70px">
    <td style="padding:10px">
      \({@first(sixth(data))@} < t \leq 60\)
    </td>
    <td style="padding:10px; vertical-align:middle">
      {@third(sixth(data))*(60-first(sixth(data)))@}
    </td>
    <td style="padding:10px; vertical-align:middle">
      \(\frac{{@third(sixth(data))*(60-first(sixth(data)))@}}{{@60-first(sixth(data))@}} = {@tans6@}\)
    </td>
  </tr>
</table>

<p><br></p>

<p>
(b) The completed histogram below represents the distribution of the time taken by
students to complete the Math CAT Mark assessment.
</p>


[[jsxgraph width="750px" height="450px" input-ref-ans7="sR1" input-ref-ans8="sR2" input-ref-ans9="sR3" input-ref-ans10="sR4" input-ref-ans11="sR5" input-ref-ans12="sR6"]]

    // [min, max]
    var ranges = [
      [0, {#second(first(data))#}],
      [{#first(second(data))#}, {#second(second(data))#}],
      [{#first(third(data))#}, {#second(third(data))#}],
      [{#first(fourth(data))#}, {#second(fourth(data))#}],
      [{#first(fifth(data))#}, {#second(fifth(data))#}],
      [{#first(sixth(data))#}, 60],
    ];

    JXG.Options.text.useMathJax=true;
    var board = JXG.JSXGraph.initBoard(divid, {
      boundingBox:[-5,15,65,-3],
      showCopyright:false,
      showNavigation:false,
      keepAspectRatio:false,
      grid:true
    });

    // Defining axes

    var yAxis = board.create("axis", [[0,0], [0,50]], {
      highlight:false,
      name:"\\(fd\\)",
      withLabel:true,
      label:{
        fixed:true, // Set fixed to true
        highlight:false,
        fontSize:15,
        position:"rt",
        offset:[-20,-20],
        anchorX:"right",
        useMathJax:true,
        parse:false
      },
      ticks:{
        label:{
          highlight:false,
          anchorX:"right",
          offset:[-10,0]
        }
      }
    });

    var xAxis = board.create("axis", [[0,0], [40,0]], {
      highlight:false,
      name:"\\(t \\ [min]\\)",
      withLabel:true,
      label:{
        fixed:true, // Set fixed to true
        highlight:false,
        fontSize:15,
        position:"rt",
        offset:[0,-40],
        anchorX:"right",
        useMathJax:true,
        parse:false
      },
      ticks:{
        label:{
          highlight:false,
          anchorX:"middle",
          offset:[0,-15]
        }
      }
    });

    var i;

    // Vertical segments along which gliding points can move

    var segList = [];
    for (i in ranges) {
      segList.push(board.create("segment", [
        [(ranges[i][0]+ranges[i][1])/2, 0],
        [(ranges[i][0]+ranges[i][1])/2, 14]
      ], {
        fixed:true, // Set fixed to true
        visible:false
      }));
    };

    // User movable gliders

    var pointList = [];
    for (i in ranges) {
      pointList.push(board.create("glider", [(ranges[i][0]+ranges[i][1])/2, 0, segList[i]], {
        showInfoBox:false,
        size:4,
        snapToGrid:true,
        withLabel:false
      }));
    };
    
    // Constructing rectangles from glider points

    var constStyle = {
      fixed:true, // Set fixed to true
      visible:false
    };

    var lineListHori = [];
    for (i in ranges) {
      lineListHori.push(board.create("parallel", [
        pointList[i], 
        xAxis
      ], constStyle));
    };

    var lineListVert = [
      board.create("parallel", [
        [ranges[0][0], 0],
        yAxis
      ], constStyle)
    ];

    for (i in ranges) {
      lineListVert.push(board.create("parallel", [
        [ranges[i][1], 0], 
        yAxis
      ], constStyle));
    };

    var interList = [];
    for (i = 0; i < pointList.length; i++) {
      interList.push([
        board.create("intersection", [
          lineListHori[i],
          lineListVert[i]
        ], constStyle),
        board.create("intersection", [
          lineListHori[i],
          lineListVert[i+1]
        ], constStyle)
      ]);
    };

    for (i = 0; i < pointList.length; i++) {
      board.create("polygon", [
        [ranges[i][0], 0],
        [ranges[i][1], 0],
        interList[i][1],
        interList[i][0]
      ], {
        fixed:true, // Set fixed to true
        highlight:false,
        vertices:{
          fixed:true, // Set fixed to true
          visible:false
        },
        borders:{
          fixed:true, // Set fixed to true
          highlight:false
        }
      })
    };
// -----------------------------
// Frequency polygon (FINAL)
// -----------------------------

// Endpoints on x-axis
var startPoint = board.create("point", [ranges[0][0], 0], {
  fixed: true,
  visible: false
});

var endPoint = board.create("point", [ranges[ranges.length - 1][1], 0], {
  fixed: true,
  visible: false
});

// Midpoints of bars
var midPoints = [];
for (i = 0; i < ranges.length; i++) {
  midPoints.push(
    board.create("point", [
      (ranges[i][0] + ranges[i][1]) / 2,
      function (j) {
        return function () {
          return pointList[j].Y();
        };
      }(i)
    ], {
      fixed: true,
      visible: false
    })
  );
}

// Polygon segments
board.create("segment", [startPoint, midPoints[0]], {
  strokeColor: "black",
  strokeWidth: 3,
  fixed: true
});

for (i = 1; i < midPoints.length; i++) {
  board.create("segment", [midPoints[i - 1], midPoints[i]], {
    strokeColor: "black",
    strokeWidth: 3,
    fixed: true
  });
}

board.create("segment", [midPoints[midPoints.length - 1], endPoint], {
  strokeColor: "black",
  strokeWidth: 3,
  fixed: true
});

    stack_jxg.bind_point(sR1, pointList[0]);
    stack_jxg.bind_point(sR2, pointList[1]);
    stack_jxg.bind_point(sR3, pointList[2]);
    stack_jxg.bind_point(sR4, pointList[3]);
    stack_jxg.bind_point(sR5, pointList[4]);
    stack_jxg.bind_point(sR6, pointList[5]);

    board.unsuspendUpdate();

[[/jsxgraph]]

<p>
(c) To estimate the number of students who completed the Math CAT Mark assessment
in less than 20 minutes, we consider the interval \(0 < t \leq 20\).
</p>

<p>
This interval consists of complete class intervals from part (a), together with a partial
interval
\(
{@if first(sixth(data))<20 then first(sixth(data)) else if first(fifth(data))<20 then first(fifth(data)) else if first(fourth(data))<20 then first(fourth(data)) else if first(third(data))<20 then first(third(data)) else if first(second(data))<20 then first(second(data)) else if first(first(data))<20 then first(first(data))@} < t \leq 20.
\)
</p>

<p>
The number of students in this partial interval is estimated by finding the area of the
corresponding section of the histogram.
</p>

<p>
Using this approach, the estimated number of students in the partial interval is
\(
{@if first(sixth(data))<20 then (20-first(sixth(data)))*third(sixth(data)) else if first(fifth(data))<20 then (20-first(fifth(data)))*third(fifth(data)) else if first(fourth(data))<20 then (20-first(fourth(data)))*third(fourth(data)) else if first(third(data))<20 then (20-first(third(data)))*third(third(data)) else if first(second(data))<20 then (20-first(second(data)))*third(second(data)) else if first(first(data))<20 then (20-first(first(data)))*third(first(data))@}.
\)
</p>

<p>
Therefore, the estimated total number of students who completed the assessment
in less than 20 minutes is
\(
{@if first(sixth(data))<20 then sconcat(first(first(widthList))*tans1, "+", first(second(widthList))*tans2, "+", first(third(widthList))*tans3, "+", first(fourth(widthList))*tans4, "+", first(fifth(widthList))*tans5 , "+", (20-first(sixth(data)))*third(sixth(data))) else if first(fifth(data))<20 then sconcat(first(first(widthList))*tans1, "+", first(second(widthList))*tans2, "+", first(third(widthList))*tans3, "+", first(fourth(widthList))*tans4, "+", (20-first(fifth(data)))*third(fifth(data))) else if first(fourth(data))<20 then sconcat(first(first(widthList))*tans1, "+", first(second(widthList))*tans2, "+", first(third(widthList))*tans3, "+", (20-first(fourth(data)))*third(fourth(data))) else if first(third(data))<20 then sconcat(first(first(widthList))*tans1, "+", first(second(widthList))*tans2, "+", (20-first(third(data)))*third(third(data))) else if first(second(data))<20 then sconcat(first(first(widthList))*tans1, "+", (20-first(second(data)))*third(second(data))) else if first(first(data))<20 then (20-first(sixth(data)))*third(sixth(data))@}
= {@tans13@}.
\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>15</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Class widths and frequency densities [width, freq den] */

narrowRange:[rand(2)+2, rand(3)+12];
wideRange:[rand(4)+13, rand(3)+3];

medRangeWidths: random_permutation([rand(3)+4, rand(3)+7, rand(3)+10]);
medRangeHeights: [rand(2)+6, rand(2)+8, rand(2)+10];

medRanges:[[first(medRangeWidths), first(medRangeHeights)], [second(medRangeWidths), second(medRangeHeights)], [third(medRangeWidths), third(medRangeHeights)]];

widthList: random_permutation(append([narrowRange], [wideRange], medRanges));

/* Each range given as [min, max, freq den] */

data: [[0, first(first(widthList)), second(first(widthList))], [first(first(widthList)), first(first(widthList))+first(second(widthList)), second(second(widthList))], [first(first(widthList))+first(second(widthList)), first(first(widthList))+first(second(widthList))+first(third(widthList)), second(third(widthList))], [first(first(widthList))+first(second(widthList))+first(third(widthList)), first(first(widthList))+first(second(widthList))+first(third(widthList))+first(fourth(widthList)), second(fourth(widthList))], [first(first(widthList))+first(second(widthList))+first(third(widthList))+first(fourth(widthList)), first(first(widthList))+first(second(widthList))+first(third(widthList))+first(fourth(widthList))+first(fifth(widthList)), second(fifth(widthList))], [first(first(widthList))+first(second(widthList))+first(third(widthList))+first(fourth(widthList))+first(fifth(widthList)), 60, rand(3)+1]];

/* Solutions */

tans1: third(first(data));
tans2: third(second(data));
tans3: third(third(data));
tans4: third(fourth(data));
tans5: third(fifth(data));
tans6: third(sixth(data));

tans7: [float(second(first(data))/2), third(first(data))];
tans8: [float((first(second(data))+second(second(data)))/2), third(second(data))];
tans9: [float((first(third(data))+second(third(data)))/2), third(third(data))];
tans10: [float((first(fourth(data))+second(fourth(data)))/2), third(fourth(data))];
tans11: [float((first(fifth(data))+second(fifth(data)))/2), third(fifth(data))];
tans12: [float((first(sixth(data))+60)/2), third(sixth(data))];

tans13: if first(sixth(data))<20 then (first(sixth(data))-20)*tans6+first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4+first(fifth(widthList))*tans5 else if first(fifth(data))<20 then (first(fifth(data))-20)*tans5+first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4 else if first(fourth(data))<20 then (first(sixth(data))-20)*tans4+first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3 else if first(third(data))<20 then (first(sixth(data))-20)*tans3+first(first(widthList))*tans1+first(second(widthList))*tans2 else if first(second(data))<20 then (first(sixth(data))-20)*tans2+first(first(widthList))*tans1 else if first(first(data))<20 then 20*tans1;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text><![CDATA[<p>{@data@}</p>
<p>\(0 < t \leq {@first(second(data))@}\)</p>
<p>\({@first(second(data))@} < t \leq {@first(third(data))@}\)</p>
<p>\({@first(third(data))@} < t \leq {@first(fourth(data))@}\)</p>
 <p>\({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\)</p>
  <p>\({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\)</p>
 <p>\({@first(sixth(data))@} < t \leq 60\)</p>]]></text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="reviewer"></p>
<p id="reviewer"> </p>
<p id="description">Given a frequency table for the time taken by students to complete a Math CAT Mark assessment on the school’s learning management system, the student is required to complete the histogram shown below to represent the distribution of completion times.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>tans1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>tans10</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans11</name>
      <type>algebraic</type>
      <tans>tans11</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans12</name>
      <type>algebraic</type>
      <tans>tans12</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans13</name>
      <type>algebraic</type>
      <tans>tans13</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>tans2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>tans3</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>tans4</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>tans5</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>tans6</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>tans7</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>tans8</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>tans9</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>6.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Correct answer */
node1_test: if ans1=tans1 and ans2=tans2 and ans3=tans3 and ans4=tans4 and ans5=tans5 and ans6=tans6 then 1 else 0;

/* Node 2: Worked solution */
/* Node 3-8: Check ans1-6 */</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct range 1 answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for correct range 2 answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>tans1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your answer for the \(0 < t \leq {@first(second(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your answer for the \(0 < t \leq {@first(second(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for correct range 3 answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>tans2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(second(data))@} < t \leq {@first(third(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(second(data))@} < t \leq {@first(third(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for correct range 4 answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>tans3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>5</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(third(data))@} < t \leq {@first(fourth(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(third(data))@} < t \leq {@first(fourth(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Checks for correct range 5  answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>tans4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt1-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt1-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Checks for correct range 6 answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>tans5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>7</truenextnode>
        <trueanswernote>prt1-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt1-7-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>Checks for correct range 7 answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>tans6</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt1-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(sixth(data))@} < t \leq 60\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your answer for the \({@first(sixth(data))@} < t \leq 60\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>6.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Node 1: Correct answer */
node1_test: if ans7=tans7 and ans8=tans8 and ans9=tans9 and ans10=tans10 and ans11=tans11 and ans12=tans12 then 1 else 0;

/* Node 2: Worked solution */

/* Node 3: ans7 correct */
/* Node 4: ans7 consistent with (a) */
node4_error: [float(second(first(data))/2), ans1];

/* Node 5: ans8 correct */
/* Node 6: ans8 consistent with (a) */
node6_error: [float((first(second(data))+second(second(data)))/2), ans2];

/* Node 7: ans9 correct */
/* Node 8: ans9 consistent with (a) */
node8_error: [float((first(third(data))+second(third(data)))/2), ans3];

/* Node 9: ans10 correct */
/* Node 10: ans10 consistent with (a) */
node10_error: [float((first(fourth(data))+second(fourth(data)))/2), ans4];

/* Node 11: ans11 correct */
/* Node 12: ans11 consistent with (a) */
node12_error: [float((first(fifth(data))+second(fifth(data)))/2), ans5];

/* Node 13: ans12 correct */
/* Node 14: ans12 consistent with (a) */
node14_error: [float((first(sixth(data))+60)/2), ans6];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>node1_test</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>10</name>
        <description>Checks for correct dragged  histogram bar 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>tans11</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>12</truenextnode>
        <trueanswernote>prt2-11-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>11</falsenextnode>
        <falseanswernote>prt2-11-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>11</name>
        <description>Checks for correct dragged  histogram bar 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans11</sans>
        <tans>node12_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>12</truenextnode>
        <trueanswernote>prt2-12-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\) range is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>12</falsenextnode>
        <falseanswernote>prt2-12-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(fifth(data))@} < t \leq {@first(sixth(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>12</name>
        <description>Checks for correct dragged  histogram bar 3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>tans12</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-13-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(sixth(data))@} < t \leq 60\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>13</falsenextnode>
        <falseanswernote>prt2-13-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>13</name>
        <description>Checks for correct dragged  histogram bar 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans12</sans>
        <tans>node14_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt2-14-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(sixth(data))@} < t \leq 60\) range is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-14-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(sixth(data))@} < t \leq 60\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for correct dragged  histogram bar 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>tans7</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt2-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your histogram bar for the \(0 < t \leq {@first(second(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for correct dragged  histogram bar 6</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>node4_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>4</truenextnode>
        <trueanswernote>prt2-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>Your histogram bar for the \(0 < t \leq {@first(second(data))@}\) range is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt2-4-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p><br>Your histogram bar for the \(0 < t \leq {@first(second(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Checks for correct dragged  histogram bar 3</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>tans8</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt2-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(second(data))@} < t \leq {@first(third(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>5</falsenextnode>
        <falseanswernote>prt2-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>5</name>
        <description>Checks for correct dragged  histogram bar 2</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>node6_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>6</truenextnode>
        <trueanswernote>prt2-6-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(second(data))@} < t \leq {@first(third(data))@}\) range is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>6</falsenextnode>
        <falseanswernote>prt2-6-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(second(data))@} < t \leq {@first(third(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>6</name>
        <description>Checks for correct dragged  histogram bar 1</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>tans9</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt2-7-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(third(data))@} < t \leq {@first(fourth(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>7</falsenextnode>
        <falseanswernote>prt2-7-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>7</name>
        <description>Checks for correct dragged  histogram bar 4</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans9</sans>
        <tans>node8_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>8</truenextnode>
        <trueanswernote>prt2-8-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(third(data))@} < t \leq {@first(fourth(data))@}\) range is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>8</falsenextnode>
        <falseanswernote>prt2-8-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(third(data))@} < t \leq {@first(fourth(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>8</name>
        <description>Checks for correct dragged  histogram bar 5</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>tans10</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>10</truenextnode>
        <trueanswernote>prt2-9-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\) range is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>9</falsenextnode>
        <falseanswernote>prt2-9-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>9</name>
        <description>Checks for correct dragged  histogram bar 6</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans10</sans>
        <tans>node10_error</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0.167</truescore>
        <truepenalty/>
        <truenextnode>10</truenextnode>
        <trueanswernote>prt2-10-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\) range is <b>incorrect</b>.
<br>However, it is consistent with the answer that you gave to part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>10</falsenextnode>
        <falseanswernote>prt2-10-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your histogram bar for the \({@first(fourth(data))@} < t \leq {@first(fifth(data))@}\) range is <b>incorrect</b>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>3.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[/* Node 1: Correct answer */
/* Node 2: Worked solution */

/* Nodes 3-4: Does not use linear interpolation, just sums whole ranges */

node3_error: if first(sixth(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4+first(fifth(widthList))*tans5+first(sixth(data))*tans6 else if first(fifth(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4+first(fifth(widthList))*tans5 else if first(fourth(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4 else if first(third(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3 else if first(second(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2 else if first(first(data))<20 then first(first(widthList))*tans1;

node4_error: if first(sixth(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4+first(fifth(widthList))*tans5 else if first(fifth(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3+first(fourth(widthList))*tans4 else if first(fourth(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2+first(third(widthList))*tans3 else if first(third(data))<20 then first(first(widthList))*tans1+first(second(widthList))*tans2 else if first(second(data))<20 then first(first(widthList))*tans1 else if first(first(data))<20 then 0;]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans13</sans>
        <tans>tans13</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>1</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Checks for all included students </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans13</sans>
        <tans>node3_error</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.333</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have included all of the students in the \({@if first(sixth(data))<20 then first(sixth(data)) else if first(fifth(data))<20 then first(fifth(data)) else if first(fourth(data))<20 then first(fourth(data)) else if first(third(data))<20 then first(third(data)) else if first(second(data))<20 then first(second(data)) else if first(first(data))<20 then 0@} < t \leq {@if first(sixth(data))<20 then 60 else if first(fifth(data))<20 then second(fifth(data)) else if first(fourth(data))<20 then second(fourth(data)) else if first(third(data))<20 then second(third(data)) else if first(second(data))<20 then second(second(data)) else if first(first(data))<20 then second(first(data))@}\) range.
<br>However, <b>some</b> of the students' journeys in this range are longer than 20 minutes.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Checks for not all included students </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans13</sans>
        <tans>node4_error</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.333</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><br>You have not included any of the students in the \({@if first(sixth(data))<20 then first(sixth(data)) else if first(fifth(data))<20 then first(fifth(data)) else if first(fourth(data))<20 then first(fourth(data)) else if first(third(data))<20 then first(third(data)) else if first(second(data))<20 then first(second(data)) else if first(first(data))<20 then 0@} < t \leq {@if first(sixth(data))<20 then 60 else if first(fifth(data))<20 then second(fifth(data)) else if first(fourth(data))<20 then second(fourth(data)) else if first(third(data))<20 then second(third(data)) else if first(second(data))<20 then second(second(data)) else if first(first(data))<20 then second(first(data))@}\) range.
<br>However, <b>some</b> of the students' journeys in this range are shorter than 20 minutes.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>68493</deployedseed>
    <deployedseed>564601512</deployedseed>
    <deployedseed>1568346964</deployedseed>
    <deployedseed>1332604686</deployedseed>
    <deployedseed>818400324</deployedseed>
    <deployedseed>1646466817</deployedseed>
    <deployedseed>410864816</deployedseed>
    <deployedseed>470627697</deployedseed>
    <deployedseed>1626156462</deployedseed>
    <deployedseed>1435552801</deployedseed>
    <deployedseed>1562963000</deployedseed>
    <deployedseed>542968020</deployedseed>
    <deployedseed>979855631</deployedseed>
    <deployedseed>1685975696</deployedseed>
    <deployedseed>1221908091</deployedseed>
    <deployedseed>181920437</deployedseed>
    <deployedseed>915987477</deployedseed>
    <deployedseed>971955865</deployedseed>
    <deployedseed>1662753888</deployedseed>
    <deployedseed>2099429003</deployedseed>
    <deployedseed>1811489932</deployedseed>
    <deployedseed>519047429</deployedseed>
    <deployedseed>1527124257</deployedseed>
    <deployedseed>78380188</deployedseed>
    <deployedseed>1687207903</deployedseed>
    <deployedseed>1769612913</deployedseed>
    <deployedseed>1667302454</deployedseed>
    <deployedseed>1359667857</deployedseed>
    <deployedseed>1267110978</deployedseed>
    <deployedseed>286412896</deployedseed>
    <deployedseed>162599168</deployedseed>
    <deployedseed>1938170020</deployedseed>
    <deployedseed>1281490526</deployedseed>
    <deployedseed>1454108059</deployedseed>
    <deployedseed>1466365485</deployedseed>
    <deployedseed>1527373785</deployedseed>
    <deployedseed>1804134857</deployedseed>
    <deployedseed>49466454</deployedseed>
    <deployedseed>1069532953</deployedseed>
    <deployedseed>583492495</deployedseed>
    <deployedseed>662892985</deployedseed>
    <deployedseed>574168607</deployedseed>
    <deployedseed>1148775761</deployedseed>
    <deployedseed>2085356976</deployedseed>
    <deployedseed>1624186558</deployedseed>
    <deployedseed>300480599</deployedseed>
    <deployedseed>1090754517</deployedseed>
    <deployedseed>1245303539</deployedseed>
    <deployedseed>1176228820</deployedseed>
    <deployedseed>1615268850</deployedseed>
    <deployedseed>136132084</deployedseed>
    <deployedseed>897175852</deployedseed>
    <deployedseed>2136339773</deployedseed>
    <deployedseed>1110512081</deployedseed>
    <deployedseed>1456263125</deployedseed>
    <deployedseed>772194053</deployedseed>
    <deployedseed>1338146547</deployedseed>
    <deployedseed>933817546</deployedseed>
    <deployedseed>828741140</deployedseed>
    <deployedseed>1190023622</deployedseed>
    <deployedseed>860965708</deployedseed>
    <deployedseed>546140717</deployedseed>
    <deployedseed>1151977576</deployedseed>
    <deployedseed>559405238</deployedseed>
    <deployedseed>488650530</deployedseed>
    <deployedseed>546733971</deployedseed>
    <deployedseed>46662752</deployedseed>
    <deployedseed>1219337834</deployedseed>
    <deployedseed>1156542400</deployedseed>
    <deployedseed>205293852</deployedseed>
    <deployedseed>1985424427</deployedseed>
    <deployedseed>2042811877</deployedseed>
    <deployedseed>368028496</deployedseed>
    <deployedseed>1873881736</deployedseed>
    <deployedseed>1852671360</deployedseed>
    <deployedseed>799511099</deployedseed>
    <deployedseed>841251800</deployedseed>
    <deployedseed>1360840155</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>tans1</value>
      </testinput>
      <testinput>
        <name>ans10</name>
        <value>tans10</value>
      </testinput>
      <testinput>
        <name>ans11</name>
        <value>tans11</value>
      </testinput>
      <testinput>
        <name>ans12</name>
        <value>tans12</value>
      </testinput>
      <testinput>
        <name>ans13</name>
        <value>tans13</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>tans2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>tans3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>tans4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>tans5</value>
      </testinput>
      <testinput>
        <name>ans6</name>
        <value>tans6</value>
      </testinput>
      <testinput>
        <name>ans7</name>
        <value>tans7</value>
      </testinput>
      <testinput>
        <name>ans8</name>
        <value>tans8</value>
      </testinput>
      <testinput>
        <name>ans9</name>
        <value>tans9</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-2-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
