<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Surface area of a storage box, made of a hemisphere and a cylinder</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A designer is creating a decorative storage container by joining a hemisphere to the top of a cylinder, as shown below.
</p>

[[jsxgraph]]
const brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10, 10, 10, -12],
    axis: false,
    showNavigation: false,
    showCopyright: false,
    keepAspectRatio: true
});

const r = 4;      // schematic radius
const h = 8;      // schematic cylinder height
const k = 0.30;   // ellipse squash factor

// ---------- Cylinder ----------

// Cylinder sides
const L1 = brd.create('point', [-r,-h], {visible:false});
const L2 = brd.create('point', [-r,0],  {visible:false});
const R1 = brd.create('point', [ r,-h], {visible:false});
const R2 = brd.create('point', [ r,0],  {visible:false});

brd.create('segment',[L1,L2],{
    strokeColor:'black',
    strokeWidth:2
});

brd.create('segment',[R1,R2],{
    strokeColor:'black',
    strokeWidth:2
});

// Bottom ellipse
brd.create('curve',[
    t => r*Math.cos(t),
    t => -h + k*r*Math.sin(t),
    0,2*Math.PI
],{
    strokeColor:'black',
    strokeWidth:1.5
});

// Top join ellipse (visible front)
brd.create('curve',[
    t => r*Math.cos(t),
    t => k*r*Math.sin(t),
    0,Math.PI
],{
    strokeColor:'black',
    strokeWidth:1.5
});

// Top join ellipse (hidden back)
brd.create('curve',[
    t => r*Math.cos(t),
    t => k*r*Math.sin(t),
    Math.PI,2*Math.PI
],{
    strokeColor:'black',
    strokeWidth:1,
    dash:2
});

// ---------- Hemisphere ----------

// Hemisphere dome
brd.create('curve',[
    t => r*Math.cos(t),
    t => r*Math.sin(t),
    0,Math.PI
],{
    strokeColor:'#1f3c88',
    strokeWidth:2,
    fillColor:'#90caf9',
    fillOpacity:0.5
});

// ---------- Dimension lines ----------

// Radius
const O = brd.create('point',[0,-h],{visible:false});
const Rad = brd.create('point',[r,-h],{visible:false});

brd.create('segment',[O,Rad],{
    dash:2
});

brd.create('text',[
    r/2,
    -h-1,
    '{#radius#} cm'
]);

// Height
const H1 = brd.create('point',[-6,0],{visible:false});
const H2 = brd.create('point',[-6,-h],{visible:false});

brd.create('segment',[H1,H2],{
    strokeWidth:1,
    dash:2
});

brd.create('text',[
    -8,
    -h/2,
    '{#height#} cm'
]);

[[/jsxgraph]]

<p>
The cylinder has:
</p>

<ul>
<li>radius <b>{@radius@} cm</b></li>
<li>height <b>{@height@} cm</b></li>
</ul>


<p>
Calculate the total external surface area of the container. Use \(\pi=\frac{22}{7}\).
</p>

<p>
[[input:ans1]]\(\text{cm}^2\)  [[feedback:prt1]]
</p>
<p>[[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Solution</h3>

<p>
To calculate the total external surface area of the composite solid, identify only the surfaces that are exposed.
</p>
<p> The exposed surfaces are the curved surface of the cylinder, the hemisphere and the circle. Therefore the formulas to be used are : </p>

<ul>
<li> surface area of the curved part of the cylinder =\(2\pi rh\)
</li>

<li> Surface area of the hemisphere = \(2\pi r^2\)
</li>

<li> Since the cylinder has a circular base the formula to be used is  \(\pi r^2\)
</li>
</ul>

<p>
The circular face where the hemisphere joins the cylinder is inside the solid and is <strong>not</strong> part of the external surface area.
</p>

<p>
Therefore,
</p>

\[
\begin{aligned}
\text{Total Surface Area}
&=2\pi rh+2\pi r^2+\pi r^2\\
&=2\pi rh+3\pi r^2.
\end{aligned}
\]

<p>
Substituting the given values of  radius(r) and height(h) into the formula
\(r={@radius@}\text{ cm},\qquad h={@height@}\text{ cm}\)

gives; (\(2\times\frac{22}{7}\times {@radius@}\times {@height@})+(3\times\frac{22}{7}\times {@radius@}^2)= {@Total_surface_area@}\text{ cm}^2.\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* List of [radius,height] pairs */
dimensions:[
    [3,8],[4,10],[5,12],[6,15],[7,16],
    [8,20],[9,24],[10,25],[12,30],[15,36],
    [16,40],[18,45],[20,48],[24,60],[25,64]
];

/* Randomly choose one pair */
idx:rand(length(dimensions))+1;
chosen:dimensions[idx];

/* Extract radius and height */
radius:chosen[1];
height:chosen[2];

/* Use π ≈ 22/7 */
pi_value:22/7;

/* Surface areas */
Cylinder_curved_area:2*pi_value*radius*height;
Hemisphere_curved_area:2*pi_value*radius^2;
Cylinder_base_area:pi_value*radius^2;

/* Correct answer */
Total_surface_area:
round(
    Cylinder_curved_area
    + Hemisphere_curved_area
    + Cylinder_base_area
);

ta:Total_surface_area;

/* Common misconceptions */

/* Using full sphere surface area */
Sphere_mistake:
round(
    Cylinder_curved_area
    + 4*pi_value*radius^2
    + Cylinder_base_area
);

/* Omitting the cylinder base */
No_base_mistake:
round(
    Cylinder_curved_area
    + Hemisphere_curved_area
);

/* Adding two circular bases */
Two_bases_mistake:
round(
    Cylinder_curved_area
    + Hemisphere_curved_area
    + 2*Cylinder_base_area
);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Total_surface_area</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Total_surface_area</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Calculating of the correct surface area</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Total_surface_area</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Forgetting to include the circular base of the storage box</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>No_base_mistake</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p> Incorrect! Seems like you forgot to add the base of the storage box to the total surface area. You used the formula: 
</p>
<p> \{Surface area} = 2\times\pi_value\times\(r^2)+pi_value\times\r\times\h</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <deployedseed>6</deployedseed>
    <deployedseed>7</deployedseed>
    <deployedseed>8</deployedseed>
    <deployedseed>9</deployedseed>
    <deployedseed>10</deployedseed>
    <deployedseed>11</deployedseed>
    <deployedseed>12</deployedseed>
    <deployedseed>13</deployedseed>
    <deployedseed>14</deployedseed>
    <deployedseed>15</deployedseed>
    <deployedseed>16</deployedseed>
    <deployedseed>17</deployedseed>
    <deployedseed>18</deployedseed>
    <deployedseed>19</deployedseed>
    <deployedseed>20</deployedseed>
    <deployedseed>21</deployedseed>
    <deployedseed>22</deployedseed>
    <deployedseed>23</deployedseed>
    <deployedseed>24</deployedseed>
    <deployedseed>25</deployedseed>
    <deployedseed>26</deployedseed>
    <deployedseed>27</deployedseed>
    <deployedseed>28</deployedseed>
    <deployedseed>29</deployedseed>
    <deployedseed>30</deployedseed>
    <deployedseed>31</deployedseed>
    <deployedseed>32</deployedseed>
    <deployedseed>33</deployedseed>
    <deployedseed>34</deployedseed>
    <deployedseed>35</deployedseed>
    <deployedseed>36</deployedseed>
    <deployedseed>37</deployedseed>
    <deployedseed>38</deployedseed>
    <deployedseed>39</deployedseed>
    <deployedseed>40</deployedseed>
    <deployedseed>41</deployedseed>
    <deployedseed>42</deployedseed>
    <deployedseed>43</deployedseed>
    <deployedseed>44</deployedseed>
    <deployedseed>45</deployedseed>
    <deployedseed>46</deployedseed>
    <deployedseed>47</deployedseed>
    <deployedseed>48</deployedseed>
    <deployedseed>49</deployedseed>
    <deployedseed>50</deployedseed>
    <deployedseed>51</deployedseed>
    <deployedseed>52</deployedseed>
    <deployedseed>53</deployedseed>
    <deployedseed>54</deployedseed>
    <deployedseed>55</deployedseed>
    <deployedseed>56</deployedseed>
    <deployedseed>57</deployedseed>
    <deployedseed>58</deployedseed>
    <deployedseed>59</deployedseed>
    <deployedseed>60</deployedseed>
    <deployedseed>61</deployedseed>
    <deployedseed>62</deployedseed>
  </question>
</quiz>
