<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Using logarithms to compute a ratio</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
Given that:
</p>

<p>
\[
\text{log } A = {@L_a@}, \quad \text{log } B = {@L_b@}
\]
</p>

<p>
Use logarithms to find the value of:
</p>

<p>
\[
\frac{A}{B}
\]
</p>

<p>
Give your answer correct to 3 decimal places.
</p>

<p>
[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
Step 1: Use the quotient rule:
</p>

<p>
\[
\text{log} \left(\frac{A}{B}\right) = \text{log } A - \text{log } B
\]
</p>

<p>
\[
= {@L_a@} - {@L_b@}
\]
</p>

<p>
\[
= {@L_diff@}
\]
</p>

<p>
Step 2: Take the antilogarithm:
</p>

<p>
\[
\frac{A}{B} = 10^{@L_diff@}
\]
</p>

<p>
\[
= {@ans_round@}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Given logarithm values */
L_a: float((rand(3000) + 2000)/10000);
L_b: float((rand(2000) + 1000)/10000);

/* Difference of logarithms */
L_diff: L_a - L_b;

/* Final answer by antilogarithm */
ans: 10^(L_diff);

/* Rounded answer to 3 decimal places */
ans_round: float(round(ans*1000)/1000);

/* Common mistakes */
m1: L_a + L_b;
m2: L_diff;
m3: 10^(L_a + L_b);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ans_round@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Using logarithms to compute ratios</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to compute a ratio using logarithms by applying the quotient rule and then finding the antilogarithm.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ans_round</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ans_round</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>added instead if substracting</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text>For division, subtract the logarithms</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>stopped at logarithm value</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You found the logarithm correctly, but you must take the antilogarithm to get the final value.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Took antilog of the sum</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text>You used addition instead of subtraction. For ratios subtract the logarithms first</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text>Subtract the logarithms, then take the antilogarithm</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>377744507</deployedseed>
    <deployedseed>2089556668</deployedseed>
    <deployedseed>1363188490</deployedseed>
    <deployedseed>2110815998</deployedseed>
    <deployedseed>1530755388</deployedseed>
    <deployedseed>827312746</deployedseed>
    <deployedseed>110330549</deployedseed>
    <deployedseed>286981797</deployedseed>
    <deployedseed>1534296949</deployedseed>
    <deployedseed>500180559</deployedseed>
    <deployedseed>1791153977</deployedseed>
    <deployedseed>36643497</deployedseed>
    <deployedseed>345621767</deployedseed>
    <deployedseed>67330556</deployedseed>
    <deployedseed>1141028146</deployedseed>
    <deployedseed>1277770257</deployedseed>
    <deployedseed>2136064301</deployedseed>
    <deployedseed>701836891</deployedseed>
    <deployedseed>262603115</deployedseed>
    <deployedseed>725043378</deployedseed>
    <deployedseed>1054086243</deployedseed>
    <deployedseed>1048447618</deployedseed>
    <deployedseed>1727509954</deployedseed>
    <deployedseed>1692424912</deployedseed>
    <deployedseed>1056688384</deployedseed>
    <deployedseed>423099098</deployedseed>
    <deployedseed>959806065</deployedseed>
    <deployedseed>513610022</deployedseed>
    <deployedseed>1016125057</deployedseed>
    <deployedseed>1287531711</deployedseed>
    <deployedseed>804743380</deployedseed>
    <deployedseed>1073761094</deployedseed>
    <deployedseed>825163453</deployedseed>
    <deployedseed>303795951</deployedseed>
    <deployedseed>1819003217</deployedseed>
    <deployedseed>407759628</deployedseed>
    <deployedseed>1427024408</deployedseed>
    <deployedseed>1922254618</deployedseed>
    <deployedseed>433000295</deployedseed>
    <deployedseed>1122185196</deployedseed>
    <deployedseed>982015972</deployedseed>
    <deployedseed>785226536</deployedseed>
    <deployedseed>2087381784</deployedseed>
    <deployedseed>897179221</deployedseed>
    <deployedseed>1386260412</deployedseed>
    <deployedseed>1709868987</deployedseed>
    <deployedseed>1541850323</deployedseed>
    <deployedseed>1605445849</deployedseed>
    <deployedseed>675953208</deployedseed>
    <deployedseed>998411892</deployedseed>
    <deployedseed>1114069269</deployedseed>
    <deployedseed>1627442405</deployedseed>
    <deployedseed>154708202</deployedseed>
    <deployedseed>1031232714</deployedseed>
    <deployedseed>514823027</deployedseed>
    <deployedseed>905305591</deployedseed>
    <deployedseed>1054068252</deployedseed>
    <deployedseed>860102247</deployedseed>
    <deployedseed>594870002</deployedseed>
    <deployedseed>386695734</deployedseed>
    <deployedseed>1980342509</deployedseed>
    <deployedseed>1649749542</deployedseed>
    <deployedseed>1580518881</deployedseed>
    <deployedseed>1887880529</deployedseed>
    <deployedseed>584389964</deployedseed>
    <deployedseed>284971069</deployedseed>
    <deployedseed>2027781082</deployedseed>
    <deployedseed>81128996</deployedseed>
    <deployedseed>2132196307</deployedseed>
    <deployedseed>1452867974</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ans_round</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
