<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a farmer's trapezium-shaped field</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A farmer is designing a trapezium-shaped field. The two parallel sides measure <b>{@a@} m</b> and <b>{@b@} m</b>. The sloping sides are equal in length.</p>

<p>
To calculate how much land is available for planting, the farmer needs to determine the area of the field.
</p>

[[jsxgraph]]
    var a = {#a#};
    var b = {#b#};
    var h = {#h#};
    var slant = {#slant#};
    var diff = (a - b) / 2;
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-2, h + 3, a + 4, -2],
        keepAspectRatio: true,
        showNavigation: false,
        showCopyright: false,
        pan: {enabled: false},
        zoom: {enabled: false}
    });
    
    // Points
    var p1 = board.create('point', [0, 0], {visible: false, fixed: true});
    var p2 = board.create('point', [a, 0], {visible: false, fixed: true});
    var p3 = board.create('point', [a - diff, h], {visible: false, fixed: true});
    var p4 = board.create('point', [diff, h], {visible: false, fixed: true});
    
    // Polygon
    board.create('polygon', [p1, p2, p3, p4], {
        fillColor: '#e6f2ff', fillOpacity: 0.8, fixed: true, highlight: false,
        borders: {strokeColor: 'black', strokeWidth: 2, fixed: true, highlight: false}
    });
    
    // Height line
    var h_base = board.create('point', [diff, 0], {visible: false, fixed: true});
    board.create('segment', [p4, h_base], {
        dash: 2, strokeColor: 'gray', strokeWidth: 2, fixed: true
    });
    
    // Right angle marker for height
    board.create('angle', [p2, h_base, p4], {radius: 1, type: 'square', strokeColor: 'gray', fixed: true, name: ''});
    
    // Labels
    board.create('text', [a/2, -0.7, a + " m"], {anchorX: 'middle', anchorY: 'top', fontSize: 14, fixed: true});
    board.create('text', [a/2, h + 0.7, b + " m"], {anchorX: 'middle', anchorY: 'bottom', fontSize: 14, fixed: true});
    board.create('text', [diff/2 - 0.5, h/2, slant.toFixed(1) + " m"], {anchorX: 'right', anchorY: 'middle', fontSize: 14, fixed: true});
    board.create('text', [a - diff/2 + 0.5, h/2, slant.toFixed(1) + " m"], {anchorX: 'left', anchorY: 'middle', fontSize: 14, fixed: true});
    board.create('text', [diff + 0.3, h/2, "h = " + h + " m"], {anchorX: 'left', anchorY: 'middle', fontSize: 14, fixed: true});
[[/jsxgraph]]

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{m}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p>
The area of a trapezium is given by:
\[
\text{Area} = \frac{1}{2}(a + b) \times h
\]
where:
<ul>
<li>\(a\) and \(b\) are the lengths of the parallel sides</li>
<li>\(h\) is the perpendicular height</li>
</ul>
</p>

<p>
From the diagram:
<ul>
<li>Bottom base: \(a = {@a@}\) m</li>
<li>Top base: \(b = {@b@}\) m</li>
<li>Height: \(h = {@h@}\) m</li>
</ul>
</p>

<p><strong>Substitute into the formula:</strong></p>
<p>
\[
\text{Area} = \frac{1}{2}({@a@} + {@b@}) \times {@h@}
\]
</p>

<p>
\[
\text{Area} \approx {@Area_disp@} \text{ m}^2
\]
</p>

<p>
Therefore, the area of the field is:
</p>

<p>
\[
\boxed{{@Area_disp@}\text{ m}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Parallel sides */
a : rand(6) + 10;
b : rand(5) + 4;

/* Height */
h : rand(5) + 4;

/* Geometry */
diff : (a - b)/2;
slant : float(sqrt(h^2 + diff^2));

/* Area evaluated natively as float */
Area_float : float(0.5*(a + b)*h);
Area_disp : dispdp(Area_float, 2);

/* --- Common mistakes evaluated as floats --- */
/* 1. Missing 1/2 factor */
mistake1 : float((a + b)*h);

/* 2. Using slanted side instead of height */
mistake2 : float(0.5*(a + b)*slant);

/* 3. Adding all four sides (treating like perimeter idea) */
mistake3 : float(0.5*(a + b + 2*slant)*h);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>a: {@a@}, b: {@b@}, h: {@h@}, Area: {@Area_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>Area_disp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>Area_float</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>mistake1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You seem to have forgotten to multiply the product of the base sums and height by the factor \(\frac{1}{2}\). Kindly check the step-by-step worked solution.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>mistake2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have used the slanted side instead of the perpendicular height (\(h\)) to find the area. Please look at the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>mistake3</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have incorrectly used all four sides. Remember, only the parallel sides are used in the formula to find the area. Check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>370274150</deployedseed>
    <deployedseed>813448076</deployedseed>
    <deployedseed>1160264344</deployedseed>
    <deployedseed>1217290343</deployedseed>
    <deployedseed>1353853806</deployedseed>
    <deployedseed>789710838</deployedseed>
    <deployedseed>446452184</deployedseed>
    <deployedseed>1727739254</deployedseed>
    <deployedseed>1676186475</deployedseed>
    <deployedseed>500470621</deployedseed>
    <deployedseed>494494756</deployedseed>
    <deployedseed>851227890</deployedseed>
    <deployedseed>184357055</deployedseed>
    <deployedseed>1859146993</deployedseed>
    <deployedseed>569212681</deployedseed>
    <deployedseed>1285822200</deployedseed>
    <deployedseed>1211029241</deployedseed>
    <deployedseed>1620378927</deployedseed>
    <deployedseed>1334460838</deployedseed>
    <deployedseed>78160054</deployedseed>
    <deployedseed>441841901</deployedseed>
    <deployedseed>1213630383</deployedseed>
    <deployedseed>491561868</deployedseed>
    <deployedseed>268844887</deployedseed>
    <deployedseed>1031605268</deployedseed>
    <deployedseed>2099039315</deployedseed>
    <deployedseed>799306742</deployedseed>
    <deployedseed>56069321</deployedseed>
    <deployedseed>280613737</deployedseed>
    <deployedseed>1007481031</deployedseed>
    <deployedseed>1032015966</deployedseed>
    <deployedseed>1824924962</deployedseed>
    <deployedseed>1570970000</deployedseed>
    <deployedseed>1377017198</deployedseed>
    <deployedseed>1813253849</deployedseed>
    <deployedseed>1878273229</deployedseed>
    <deployedseed>1326406690</deployedseed>
    <deployedseed>1707346004</deployedseed>
    <deployedseed>438583632</deployedseed>
    <deployedseed>511996840</deployedseed>
    <deployedseed>1559667079</deployedseed>
    <deployedseed>412920299</deployedseed>
    <deployedseed>1467938140</deployedseed>
    <deployedseed>942126779</deployedseed>
    <deployedseed>2100572702</deployedseed>
    <deployedseed>1525940984</deployedseed>
    <deployedseed>1867644209</deployedseed>
    <deployedseed>1553284856</deployedseed>
    <deployedseed>1406370617</deployedseed>
    <deployedseed>1788494962</deployedseed>
    <deployedseed>1325210153</deployedseed>
    <deployedseed>685179429</deployedseed>
    <deployedseed>12733665</deployedseed>
    <deployedseed>48909808</deployedseed>
    <deployedseed>462650146</deployedseed>
    <deployedseed>1997460686</deployedseed>
    <deployedseed>514832787</deployedseed>
    <deployedseed>175575623</deployedseed>
    <deployedseed>1209777092</deployedseed>
    <deployedseed>1955059341</deployedseed>
    <deployedseed>911950879</deployedseed>
    <deployedseed>1096815537</deployedseed>
    <deployedseed>7564859</deployedseed>
    <deployedseed>342590378</deployedseed>
    <deployedseed>58525608</deployedseed>
    <deployedseed>633079306</deployedseed>
    <deployedseed>1573629576</deployedseed>
    <deployedseed>76252361</deployedseed>
    <deployedseed>969424300</deployedseed>
    <deployedseed>1993226381</deployedseed>
    <deployedseed>248783559</deployedseed>
    <deployedseed>1170472460</deployedseed>
    <deployedseed>707783702</deployedseed>
    <deployedseed>544280213</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>Area_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
