<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Representing vectors geometrically</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the points <b>P({@x1@}, {@y1@})</b> and <b>Q({@x2@}, {@y2@})</b> in a 2D Cartesian plane. The vector <b>\(\vec{PQ}\)</b> represents the displacement from P to Q.</p>

[[jsxgraph width="400px" height="400px"]]
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-1, 12, 12, -1], 
    axis: true, 
    showCopyright: false
});

var p1 = board.create('point', [{#x1#}, {#y1#}], {name: 'P', size: 3, fixed: true});
var p2 = board.create('point', [{#x2#}, {#y2#}], {name: 'Q', size: 3, fixed: true});
var vec = board.create('arrow', [p1, p2], {strokeColor: 'blue', strokeWidth: 3, withLabel: true});

// Component guidelines
board.create('segment', [[{#x1#}, {#y1#}], [{#x2#}, {#y1#}]], {dash: 2, strokeColor: 'gray'});
board.create('segment', [[{#x2#}, {#y1#}], [{#x2#}, {#y2#}]], {dash: 2, strokeColor: 'gray'});
[[/jsxgraph]]


<ol>
    <li>Determine the components of \(\vec{PQ}\): <br>
        \(\vec{PQ}\) =  [[input:ans1]] \(, \) [[input:ans2]] [[validation:ans1]][[feedback:prt1]] [[validation:ans2]][[feedback:prt2]]
    </li>
    <li>Calculate the magnitude of \(\vec{PQ}\): <br>
        \(|\vec{PQ}| = \) [[input:ans3]] [[validation:ans3]][[feedback:prt3]]
    </li>
</ol>

<ol start="3">
<li> <p>Given another vector \(b = ({@bx@}, {@by@})\)</b>, find the resultant vector \(R\) where \(R= \vec{PQ} + b\).</p>  
    <p>\(R = \) [[input:ans4]] \(, \) [[input:ans5]] [[validation:ans4]][[feedback:prt4]] [[validation:ans5]][[feedback:prt5]]</p></li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<b>Step 1: Determine the components of \(\vec{PQ}\)</b>
\(\begin{align}
\vec{PQ} &= (x_Q - x_P, y_Q - y_P) \\
&= ({@x2@} - {@x1@}, {@y2@} - {@y1@}) \\
&= {@ta1@}, {@ta2@}
\end{align}\)

<b>Step 2: Calculate the magnitude \(|\vec{PQ}|\)</b>
\(\begin{align}
|\vec{PQ}| &= \sqrt{{@ta1@}^2 + {@ta2@}^2} \\
&= \sqrt{{@mag_sq@}} \\
&= {@ta3@}
\end{align}\)

<b>Step 3: Add corresponding components for \(\mathbf{R}\)</b>
\(\begin{align}
\mathbf{R} &= \vec{PQ} + \mathbf{b} \\
&= ({@ta1@} + {@bx@}, {@ta2@} + {@by@}) \\
&= {@ta4@}, {@ta5@}
\end{align}\)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Define Points P and Q (for magnitude) */
x1: 1 + rand(3); y1: 1 + rand(3);
x2: x1 + 3 + rand(3); y2: y1 + 3 + rand(3);

/* 2. Vector PQ components and Magnitude */
ta1: x2 - x1;
ta2: y2 - y1;
mag_sq: (ta1^2) + (ta2^2);
ta3: sqrt(mag_sq);

/* 3. Define Vector b (for addition) */
bx: 1 + rand(4);
by: 1 + rand(4);

/* 4. Resultant Vector R = PQ + b */
ta4: ta1 + bx;
ta5: ta2 + by;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta1@}</p>
<p>{@ta2@}</p>
<p>{@ta3@}</p>
<p>{@ta4@}</p>
<p>{@ta5@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>406500667</deployedseed>
    <deployedseed>1731748359</deployedseed>
    <deployedseed>2010430875</deployedseed>
    <deployedseed>1475995015</deployedseed>
    <deployedseed>295544152</deployedseed>
    <deployedseed>1483241231</deployedseed>
    <deployedseed>291196386</deployedseed>
    <deployedseed>421589145</deployedseed>
    <deployedseed>270221619</deployedseed>
    <deployedseed>1309325481</deployedseed>
    <deployedseed>293149855</deployedseed>
    <deployedseed>1554825296</deployedseed>
    <deployedseed>218032458</deployedseed>
    <deployedseed>996108517</deployedseed>
    <deployedseed>664815723</deployedseed>
    <deployedseed>1870168803</deployedseed>
    <deployedseed>1790138088</deployedseed>
    <deployedseed>2008025198</deployedseed>
    <deployedseed>1674699833</deployedseed>
    <deployedseed>1955345785</deployedseed>
  </question>
</quiz>
