<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Analysis of ungrouped data for kabras sugar company</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The monthly salaries (in KES) of <strong>55 employees</strong> at <strong>Kabras Sugar Company</strong> are listed below:</p>
<div><p>{@raw_data@}</p></div>

<p>Analyze this ungrouped data to determine the measures of central tendency.</p>
<p><strong>a) Mean Salary:</strong><br>
Calculate the mean monthly salary.<br>
Mean = KES [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<p><strong>b) Median Salary:</strong><br>
Determine the median monthly salary.<br>
Median = KES [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

<p><strong>c) Mode Salary:</strong><br>
Identify the mode of the dataset.<br>
<em>(Note: If there is more than one mode, enter your answer as a set, e.g., <code>{30000, 35000}</code>).</em><br>
Mode = KES [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>a) Mean Salary</strong></p>
<p>The mean is the sum of all salaries divided by the total number of employees (\(N=55\)).</p>
\( \bar{x} = \frac{\sum x}{N} \)
\( \bar{x} = \frac{@sum_x@}{55} \)
\( \bar{x} \approx \mathbf{@decimalplaces(mean_val, 2)@} \text{ KES} \)

<br>
<p><strong>b) Median Salary</strong></p>
<p>To find the median, we first arrange the salaries in ascending order:</p>
<p>{@sorted_data@}</p>

<p>Since \( N=55 \) is an <strong>odd number</strong>, the median is exactly the middle value.</p>
<p>Position of Median: \( \frac{N+1}{2} = \frac{55+1}{2} = 28^{\text{th}} \) value.</p>

<p>Counting to the \( 28^{\text{th}} \) value in the sorted list, we find:</p>
\( \text{Median} = \mathbf{@median_val@} \text{ KES} \)

<br>
<p><strong>c) Mode Salary</strong></p>
<p>The mode is the salary that appears most frequently. We can verify this by organizing the data into a frequency table:</p>

<table border="1">
    <thead>
        <tr style="background-color: #f2f2f2;">
            <th>Salary (KES)</th>
            <th>Frequency</th>
        </tr>
    </thead>
    <tbody>
        {@freq_table_rows@}
    </tbody>
</table>

<p>The highest frequency is <strong>{@max_freq@}</strong>.</p>
<p>The salary with this frequency is: <strong>{@ta_mode@}</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. DATA GENERATION */
/* Pool of realistic monthly salaries in KES */
pool: [25000, 28000, 30000, 32000, 35000, 38000, 40000, 42000, 45000, 50000, 55000, 60000];

/* Select 55 salaries randomly from this pool */
/* We use random selection with replacement to ensure we have a Mode */
raw_data: makelist(random_permutation(pool)[1], k, 1, 55);

N: length(raw_data);

/* 2. MEAN CALCULATIONS */
sum_x: apply("+", raw_data);
mean_val: float(sum_x / N);

/* 3. MEDIAN CALCULATIONS */
sorted_data: sort(raw_data);

/* Logic for N=55 (Odd number) */
/* Position is (N+1)/2 = 28 */
med_pos: (N + 1) / 2;
median_val: sorted_data[med_pos];


/* 4. MODE CALCULATIONS */
unique_vals: unique(raw_data);
/* Count frequency of each salary */
freqs: map(lambda([x], length(sublist(raw_data, lambda([y], y=x)))), unique_vals);
max_freq: lmax(freqs);
/* Identify which salaries have the max frequency */
mode_indices: sublist_indices(freqs, lambda([x], x=max_freq));
modes: map(lambda([idx], unique_vals[idx]), mode_indices);

/* Format Mode answer (handles single or multiple modes) */
ta_mode: if length(modes)=1 then modes[1] else setify(modes);

/* 5. FEEDBACK TABLE GENERATION */
/* Build the HTML table rows dynamically */
u_vals: sort(unique_vals);
u_freqs: map(lambda([x], length(sublist(raw_data, lambda([y], y=x)))), u_vals);

freq_table_rows: "";
for i:1 thru length(u_vals) do (
   /* Highlight the Mode row in green */
   row_style: if u_freqs[i] = max_freq then "background-color: #d4edda; font-weight: bold;" else "",
   freq_table_rows: sconcat(
       freq_table_rows,
       "<tr style='", row_style, "'><td>", u_vals[i], "</td><td>", u_freqs[i], "</td></tr>"
   )
);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta_mode@}
{@median_val@}
{@mean_val@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">In this question, you are provided with \(55\) raw data points representing the monthly salaries (in KES) of employees at Kabras Sugar Company. You are required to analyze this ungrouped data to calculate the mean, median, and mode salaries.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>mean_val</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>median_val</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta_mode</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>mean_val</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>median_val</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta_mode</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>120616395</deployedseed>
    <deployedseed>752979998</deployedseed>
    <deployedseed>381788619</deployedseed>
    <deployedseed>1480294244</deployedseed>
    <deployedseed>837797560</deployedseed>
    <deployedseed>1025677972</deployedseed>
    <deployedseed>2137698053</deployedseed>
    <deployedseed>711154130</deployedseed>
    <deployedseed>939492572</deployedseed>
    <deployedseed>1931756128</deployedseed>
    <deployedseed>894350476</deployedseed>
    <deployedseed>1660413098</deployedseed>
    <deployedseed>572853496</deployedseed>
    <deployedseed>1774142952</deployedseed>
    <deployedseed>743864503</deployedseed>
    <deployedseed>151105167</deployedseed>
    <deployedseed>1649420688</deployedseed>
    <deployedseed>548664759</deployedseed>
    <deployedseed>1266316635</deployedseed>
    <deployedseed>1562633275</deployedseed>
    <deployedseed>745363087</deployedseed>
    <deployedseed>1135051677</deployedseed>
    <deployedseed>499449153</deployedseed>
    <deployedseed>1856931897</deployedseed>
    <deployedseed>442387196</deployedseed>
    <deployedseed>1817623465</deployedseed>
    <deployedseed>1809838695</deployedseed>
    <deployedseed>1894661263</deployedseed>
    <deployedseed>290886377</deployedseed>
    <deployedseed>1548246279</deployedseed>
    <deployedseed>550214144</deployedseed>
    <deployedseed>852309571</deployedseed>
    <deployedseed>384181321</deployedseed>
    <deployedseed>1968959735</deployedseed>
    <deployedseed>395206597</deployedseed>
    <deployedseed>1374187414</deployedseed>
    <deployedseed>1022923778</deployedseed>
    <deployedseed>1309401374</deployedseed>
    <deployedseed>382040280</deployedseed>
    <deployedseed>954148211</deployedseed>
    <deployedseed>810141684</deployedseed>
    <deployedseed>1447645588</deployedseed>
    <deployedseed>104877042</deployedseed>
    <deployedseed>1083007434</deployedseed>
    <deployedseed>676607486</deployedseed>
    <deployedseed>221205647</deployedseed>
    <deployedseed>181282938</deployedseed>
    <deployedseed>2122070352</deployedseed>
    <deployedseed>891517783</deployedseed>
    <deployedseed>541620038</deployedseed>
  </question>
</quiz>
