<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Grouping ungrouped data</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<h3>
Measures of Central Tendency (Real Life Data)
</h3>

<p>
Below are salaries of <b>{@N@}</b> employees in Bidco Africa company in Kenyan Shillings.
</p>

<p>
The recorded salaries are shown below:
</p>

<p style="text-align:center;">
\({@raw_data@}\)
</p>

<p>
Using the data above, answer the following questions.
</p>

<p>
<b>(a)</b> Calculate the <b>mean monthly salary</b>, to 2 decimal places.
</p>

<p>
Mean = [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
</p>

<p>
<b>(b)</b> Determine the <b>median monthly salary</b>.
</p>

<p>
Median = [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
</p>

<p>
<b>(c)</b> Determine the <b>mode of the salaries</b>.
</p>

<p>
Mode = [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution </h3>

<p>
Central tendency describes the value that represents the centre of a dataset.
The three most common measures are the <b>mean</b>, <b>median</b>, and <b>mode</b>.
</p>

<h4>1. Mean</h4>

<p>
The mean is calculated by dividing the sum of all values by the number of observations. That is 
\[
\text{Mean} = \frac{\sum x}{N}
\]
</p>
<p>
For this dataset:
</p>

<p>
Total number of employees: <b>{@N@}</b>
</p>

<p>
Correct Mean:
</p>

<p>
<b>{@mean_val@} KES</b>
</p>

<hr>

<h4>2. Median</h4>

<p>
The median is the middle value when the data is arranged in ascending order.
</p>

<p>
Since the dataset contains <b>{@N@}</b> values (an odd number), the median position is:
</p>

<p>
\[
\frac{N+1}{2} = \frac{{@N@}+1}{2}
\]
</p>

<p>
Correct Median:
</p>

<p>
<b>{@median_val@} KES</b>
</p>

<hr>

<h4>3. Mode</h4>

<p>
The mode is the value that appears most frequently in the dataset.
</p>

<p>
The table below shows the frequency of each salary.
The mode is highlighted.
</p>

<table border="1" style="border-collapse:collapse; text-align:center;">
<tr>
<th>Salary (KES)</th>
<th>Frequency</th>
</tr>

{@freq_table_rows@}

</table>

<p>
Therefore, the <b>Mode</b> is:
</p>

<p>
<b>{@ta_mode@} KES</b>
</p>

<hr>

<h4>Interpretation</h4>

<ul>
<li><b>Mean</b> gives the overall average salary.</li>
<li><b>Median</b> represents the middle salary when ordered.</li>
<li><b>Mode</b> shows the most common salary in the company.</li>
</ul>

<p>
These three measures help statisticians understand how income or other real-world data is distributed.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* ============================= */
/* 1. DATA GENERATION */
/* ============================= */

/* Pool of realistic monthly salaries in KES */
pool: [25000, 28000, 30000, 32000, 35000, 38000, 40000, 42000, 45000, 50000, 55000, 60000];

/* Select 55 salaries randomly from the pool */
raw_data: makelist(random_permutation(pool)[1], k, 1, 55);

N: length(raw_data);


/* ============================= */
/* 2. MEAN CALCULATIONS */
/* ============================= */

sum_x: apply("+", raw_data);
mean_val: float(sum_x / N);
mean_disp: round(mean_val*100)/100.0;
ta1:mean_val;

/* ============================= */
/* 3. MEDIAN CALCULATIONS */
/* ============================= */

sorted_data: sort(raw_data);

/* Position for odd number of observations */
med_pos: (N + 1) / 2;
median_val: sorted_data[med_pos];
ta2:median_val;

/* ============================= */
/* 4. MODE CALCULATIONS */
/* ============================= */

unique_vals: unique(raw_data);

/* Frequency of each salary */
freqs: map(lambda([x], length(sublist(raw_data, lambda([y], y=x)))), unique_vals);

max_freq: lmax(freqs);

/* Identify salaries with highest frequency */
mode_indices: sublist_indices(freqs, lambda([x], x=max_freq));

modes: map(lambda([idx], unique_vals[idx]), mode_indices);

/* Handle single or multiple modes */
ta_mode: if length(modes)=1 then modes[1] else setify(modes);
ta3:ta_mode;

/* ============================= */
/* 5. FEEDBACK TABLE GENERATION */
/* ============================= */

/* Sorted unique values for table */
u_vals: sort(unique_vals);

u_freqs: map(lambda([x], length(sublist(raw_data, lambda([y], y=x)))), u_vals);

/* Generate HTML table rows */
freq_table_rows: "";
for i:1 thru length(u_vals) do (
   row_style: if u_freqs[i] = max_freq then "background-color:#d4edda;font-weight:bold;" else "",
   freq_table_rows: sconcat(
       freq_table_rows,
       "<tr style='", row_style, "'><td>", u_vals[i], "</td><td>", u_freqs[i], "</td></tr>"
   )
);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}
{@ta2@}
{@ta3@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check your working: add every value, then divide by the number of values \(N\). Give the result to 2 decimal places.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Sort all the values in ascending order, then find the middle value. For an even count, average the two middle values.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The mode is the value that occurs most often. If two or more values tie for the most, enter them all as a set, e.g. <code>{a, b}</code>.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>606124195</deployedseed>
    <deployedseed>1843415159</deployedseed>
    <deployedseed>1158378682</deployedseed>
    <deployedseed>1802594339</deployedseed>
    <deployedseed>1536429111</deployedseed>
    <deployedseed>1235580854</deployedseed>
    <deployedseed>1419264566</deployedseed>
    <deployedseed>550739278</deployedseed>
    <deployedseed>545472545</deployedseed>
    <deployedseed>1494118480</deployedseed>
    <deployedseed>1231736812</deployedseed>
    <deployedseed>1507029187</deployedseed>
    <deployedseed>807863423</deployedseed>
    <deployedseed>494675504</deployedseed>
    <deployedseed>1658952651</deployedseed>
    <deployedseed>1864692028</deployedseed>
    <deployedseed>1418315020</deployedseed>
    <deployedseed>942508806</deployedseed>
    <deployedseed>1244445484</deployedseed>
    <deployedseed>1084940735</deployedseed>
    <deployedseed>1120748126</deployedseed>
    <deployedseed>1681901444</deployedseed>
    <deployedseed>1809159444</deployedseed>
    <deployedseed>1331025197</deployedseed>
    <deployedseed>521191217</deployedseed>
    <deployedseed>1019904095</deployedseed>
    <deployedseed>904311803</deployedseed>
    <deployedseed>1798180742</deployedseed>
    <deployedseed>1924171569</deployedseed>
    <deployedseed>1112735640</deployedseed>
    <deployedseed>343077848</deployedseed>
    <deployedseed>1134039016</deployedseed>
    <deployedseed>613655331</deployedseed>
    <deployedseed>481495136</deployedseed>
    <deployedseed>1385824477</deployedseed>
    <deployedseed>877373416</deployedseed>
    <deployedseed>1649099619</deployedseed>
    <deployedseed>686548701</deployedseed>
    <deployedseed>1465849383</deployedseed>
    <deployedseed>1285372873</deployedseed>
    <deployedseed>1957286409</deployedseed>
    <deployedseed>1906687515</deployedseed>
    <deployedseed>122958536</deployedseed>
    <deployedseed>136895074</deployedseed>
    <deployedseed>1047769166</deployedseed>
    <deployedseed>337999708</deployedseed>
    <deployedseed>308228905</deployedseed>
    <deployedseed>1791537660</deployedseed>
    <deployedseed>1598083241</deployedseed>
    <deployedseed>96636176</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Clearly wrong answers: each PRT fires its F branch.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1000</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+3000</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>-1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-1-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
