<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the total capacity of a rainwater harvesting tank</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A school is installing a new rainwater harvesting tank. The design consists of a cylindrical main body with a conical roof to allow debris to slide off. The dimensions of the tank are:</p>
<ul>
    <li>Radius of the tank = <b>{@r@} m</b></li>
    <li>Height of the cylindrical part = <b>{@h_cyl@} m</b></li>
    <li>Vertical height of the conical roof = <b>{@h_cone@} m</b></li>
</ul>

[[jsxgraph width="400px" height="450px"]]
var board = JXG.JSXGraph.initBoard(divid, {
    /* Bounding box accommodates the max possible heights (up to 20m total) */
    boundingbox: [-12, 25, 12, -6], 
    axis: false,
    showCopyright: false
});

var R = {#r#};
var H_cyl = {#h_cyl#};
var H_cone = {#h_cone#};
var p = 0.3; /* Perspective squash factor */

/* 1. Base of Cylinder */
/* Front solid curve */
board.create('curve', [function(t){return R*Math.cos(t)}, function(t){return (R*p)*Math.sin(t)}, Math.PI, 2*Math.PI], {strokeColor: 'black', strokeWidth: 2});
/* Back dashed curve */
board.create('curve', [function(t){return R*Math.cos(t)}, function(t){return (R*p)*Math.sin(t)}, 0, Math.PI], {strokeColor: 'black', strokeWidth: 1, dash: 2});

/* 2. Seam (Where cylinder meets cone) */
board.create('curve', [function(t){return R*Math.cos(t)}, function(t){return H_cyl + (R*p)*Math.sin(t)}, Math.PI, 2*Math.PI], {strokeColor: 'black', strokeWidth: 2});
board.create('curve', [function(t){return R*Math.cos(t)}, function(t){return H_cyl + (R*p)*Math.sin(t)}, 0, Math.PI], {strokeColor: 'black', strokeWidth: 1, dash: 2});

/* 3. Vertical Sides of Cylinder */
board.create('segment', [[-R, 0], [-R, H_cyl]], {strokeColor: 'black', strokeWidth: 2});
board.create('segment', [[R, 0], [R, H_cyl]], {strokeColor: 'black', strokeWidth: 2});

/* 4. Conical Roof */
var apex = [0, H_cyl + H_cone];
board.create('segment', [[-R, H_cyl], apex], {strokeColor: 'black', strokeWidth: 2});
board.create('segment', [[R, H_cyl], apex], {strokeColor: 'black', strokeWidth: 2});

/* Internal vertical dash for cone height */
board.create('segment', [[0, H_cyl], apex], {strokeColor: 'gray', strokeWidth: 1, dash: 2});

/* 5. Labels */
// Radius (Blue)
board.create('segment', [[0, 0], [R, 0]], {strokeColor: 'blue', strokeWidth: 2});
board.create('text', [R/2, -1.8, "{#r#} m"], {color: 'blue', fontSize: 14, fontWeight: 'bold'});

// Cylinder Height (Red)
board.create('segment', [[-R - 1.5, 0], [-R - 1.5, H_cyl]], {strokeColor: 'red', strokeWidth: 2});
board.create('text', [-R - 4.5, H_cyl/2, "{#h_cyl#} m"], {color: 'red', fontSize: 14, fontWeight: 'bold'});

// Cone Height (Green)
board.create('segment', [[-R - 1.5, H_cyl], [-R - 1.5, H_cyl + H_cone]], {strokeColor: 'green', strokeWidth: 2});
board.create('text', [-R - 4.5, H_cyl + H_cone/2, "{#h_cone#} m"], {color: 'green', fontSize: 14, fontWeight: 'bold'});

[[/jsxgraph]]

<p>Calculate the total capacity of the tank in <b>litres</b>. (Take $\pi = 3.142$)</p>
<p>Total Capacity = [[input:ans1]] litres [[validation:ans1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>To find the total capacity, we calculate the volume of the cylindrical body and the conical roof, then convert the total to litres.</p>

<p><b>Step 1: Volume of the Cylindrical Part (\(V_1\))</b></p>
\begin{align}
V_1 &= \pi r^2 h \\
&= \pi \times {@r@}^2 \times {@h_cyl@} \\
&= \pi \times {@r2@} \times {@h_cyl@} \\
&= {@V_cyl@} \text{ m}^3
\end{align}

<p><b>Step 2: Volume of the Conical Roof (\(V_2\))</b></p>
\begin{align}
V_2 &= \frac{1}{3} \pi r^2 h \\
&= \frac{1}{3} \times \pi \times {@r@}^2 \times {@h_cone@} \\
&= \frac{1}{3} \times \pi \times {@r2@} \times {@h_cone@} \\
&= {@V_cone@} \text{ m}^3
\end{align}

<p><b>Step 3: Total Capacity in Litres</b></p>
\begin{align}
\text{Total Volume} &= V_1 + V_2 \\
&= {@V_cyl@} + {@V_cone@} \\
&= {@total_m3@} \text{ m}^3
\end{align}

<p>Since \(1 \text{ m}^3 = 1,000 \text{ litres}\):</p>
\begin{align}
\text{Total Capacity} &= {@total_m3@} \times 1000 \\
&= \mathbf{{@ta@} \text{ litres}}
\end{align}]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Base Multiplier for Cone Proportions (3-4-5 ratio for clean slant height) */
multiplier: 1 + rand(3);     /* Generates 1, 2, or 3 */

/* 2. Dimensions */
r: 3 * multiplier;           /* Radius: 3, 6, or 9 m */
h_cone: 4 * multiplier;      /* Conical height: 4, 8, or 12 m */
h_cyl: 4 + rand(5);          /* Cylinder height: 4 to 8 m */

/* 3. Volume Calculations */
V_cyl: %pi * r^2 * h_cyl;
V_cone: (1/3) * %pi * r^2 * h_cone;

total_m3: V_cyl + V_cone;
total_litres: total_m3 * 1000;

/* Teacher Answer */
ta: total_litres;

/* 4. Display Variables for Feedback */
r2: r^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt1]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta@}</p>
<p>{@total_m3@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct volume</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Checks for unconverted units</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>total_m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have correctly calculated the volume of the tank but you failed to convert the units from m\(^3\) to litres. Please check the worked solution below.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>787045650</deployedseed>
    <deployedseed>318380194</deployedseed>
    <deployedseed>593582161</deployedseed>
    <deployedseed>604541413</deployedseed>
    <deployedseed>1234742911</deployedseed>
    <deployedseed>228355305</deployedseed>
    <deployedseed>1881849077</deployedseed>
    <deployedseed>1986500606</deployedseed>
    <deployedseed>149096845</deployedseed>
    <deployedseed>930484262</deployedseed>
    <deployedseed>621209505</deployedseed>
    <deployedseed>615661030</deployedseed>
  </question>
</quiz>
