<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the Depth of Water in a Cylindrical Container</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A spherical container has a diameter of \({@d_sph@}\,\text{cm}. The volume of water is \(\frac{4}{3}\) times the volume of the sphere.</p>

<p>The water is poured into a cylindrical container with a diameter of \({@d_cyl@}\,\text{cm}.\)</p>

<p>Use the volume formula for a sphere to determine the volume of water. Then find the depth of the water in the cylindrical container.</p>

<p>\(\text{The depth of the water}=\)[[input:ans1]]\(\,\text{cm}\)[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>1) Find the radius of the sphere</strong></p>

<p>The diameter of the sphere is \({@d_sph@}\,\text{cm},\) so its radius is</p>

<p>\(r_{\text{sph}}=\frac{{@d_sph@}}{2}\,\text{cm}.\)</p>

<p><strong>2) Calculate the volume of the sphere</strong></p>

<p>The volume of a sphere is \(V=\frac{4}{3}\pi r^3.\)</p>

<p>Therefore, the volume of the sphere is</p>

<p>\(V_{\text{sphere}}=\frac{4}{3}\pi\left(\frac{{@d_sph@}}{2}\right)^3={@V_sphere@}\,\text{cm}^3.\)</p>

<p><strong>3) Calculate the volume of water</strong></p>

<p>The volume of water is \(\frac{4}{3}\) times the volume of the sphere. Therefore,</p>

<p>\(V_{\text{water}}=\frac{4}{3}V_{\text{sphere}}={@V_water@}\,\text{cm}^3.\)</p>

<p><strong>4) Find the depth of the water in the cylinder</strong></p>

<p>The diameter of the cylinder is \({@d_cyl@}\,\text{cm},\) so its radius is</p>

<p>\(r_{\text{cyl}}=\frac{{@d_cyl@}}{2}\,\text{cm}.\)</p>

<p>The volume of a cylinder is \(V=\pi r^2h.\)</p>

<p>Since the volume of water is unchanged when it is poured into the cylinder,</p>

<p>\(V_{\text{water}}=\pi r_{\text{cyl}}^2h.\)</p>

<p>Hence,</p>

<p>\(h=\frac{V_{\text{water}}}{\pi r_{\text{cyl}}^2}={@h@}\,\text{cm}.\)</p>

<p>Therefore, the depth of the water is</p>

<p>\(\boxed{{@Depth_cylinder@}\,\text{cm}}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>15</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>
/* Randomize diameter of the cylinder */
d_cyl : rand_with_step(6, 20, 2);

/* Randomize diameter of the sphere (always greater than cylinder) */
d_sph : d_cyl + rand_with_step(10, 40, 2);

/* Calculate radii */
r_sph : d_sph / 2;
r_cyl : d_cyl / 2;

/* The volume of water is 4/3 times the volume of the sphere */
fraction : 4/3;

/* Use pi */
pi_val : 22/7;

/* Volume of the sphere */
V_sphere : (4/3) * pi_val * r_sph^3;

/* Volume of water */
V_water : fraction * V_sphere;

/* Depth of water in the cylinder */
h : V_water / (pi_val * r_cyl^2);

/* Convert to decimal and round */
depth : float(h);
Depth_cylinder : decimalplaces(depth, 2);
V_mistake1 : (3/4) * V_sphere;
h_mistake1 : V_mistake1 / (pi_val * r_cyl^2);
depth_mistake1 : decimalplaces(float(h_mistake1), 2);

</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The diameter of the sphere}={@d_sph@}\,\text{cm}\)</p>
<p>\(\text{The diameter of the cylinder}={@d_cyl@}\,\text{cm}\)</p>
<p>\(\text{The volume of the sphere}={@V_sphere@}\,\text{cm}^3\)</p>
<p>\(\text{The height of the cylinder}={@depth@}\,\text{cm}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>depth</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct height of the cylinder.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>depth</tans>
        <testoptions>0.000000000001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Using the incorrect formula for volume of a sphere</description>
        <answertest>NumRelative</answertest>
        <sans>ans1</sans>
        <tans>depth_mistake1</tans>
        <testoptions>0.00001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have used \(\frac{3}{4}\) of the volume of the sphere instead of \(\frac{4}{3}\).</p>

<p>The volume of the sphere is</p>

<p>\(V_{\text{sphere}}={@V_sphere@}\,\text{cm}^3.\)</p>

<p>The question states that the volume of water is \(\frac{4}{3}\) times the volume of the sphere. Therefore, the correct volume of water is</p>

<p>\(V_{\text{water}}=\frac{4}{3}V_{\text{sphere}}.\)</p>

<p>You have instead calculated</p>

<p>\(V_{\text{mistake1}}=\frac{3}{4}V_{\text{sphere}}={@V_mistake1@}\,\text{cm}^3.\)</p>

<p>Check how you have interpreted the fraction \(\frac{4}{3}\) in the question, and try again.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1552983850</deployedseed>
    <deployedseed>2065100228</deployedseed>
    <deployedseed>218362514</deployedseed>
    <deployedseed>1946267100</deployedseed>
    <deployedseed>963643092</deployedseed>
    <deployedseed>2112298858</deployedseed>
    <deployedseed>706929964</deployedseed>
    <deployedseed>2036964292</deployedseed>
    <deployedseed>1719564714</deployedseed>
    <deployedseed>1370883810</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>depth</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
