<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating the Volume of a Triangular-Based Pyramid</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The figure below shows a pyramid with a triangular base. The base of the triangle measures \({@b@}\,\text{m}\), the height of the triangle is \({@h_tri@}\,\text{m}\), and the vertical height of the pyramid is \({@H@}\,\text{m}\). Find the volume of the pyramid.</p>

[[jsxgraph height='400px' width='500px']]
    var board = JXG.JSXGraph.initBoard(divid, {
        boundingbox: [-2.5, 8.5, 10.5, -2.5],
        axis: false,
        showCopyright: false,
        keepaspectratio: true
    });

    /* Triangular base */
    var A = board.create('point', [0, 1], {
        name: '',
        visible: false
    });

    var B = board.create('point', [7, 1], {
        name: '',
        visible: false
    });

    var C = board.create('point', [2.5, -1], {
        name: '',
        visible: false
    });

    /* Apex */
    var D = board.create('point', [3.5, 7], {
        name: '',
        visible: false
    });

    /* Base triangle */
    board.create('polygon', [A, B, C], {
        fillOpacity: 0.15,
        borders: {
            strokeWidth: 2
        },
        vertices: {
            visible: false
        }
    });

    /* Pyramid edges */
    board.create('segment', [D, A], {strokeWidth: 2});
    board.create('segment', [D, B], {strokeWidth: 2});
    board.create('segment', [D, C], {strokeWidth: 2});

    /* Base edges */
    board.create('segment', [A, B], {strokeWidth: 2});
    board.create('segment', [B, C], {strokeWidth: 2});

    /* Hidden base edge */
    board.create('segment', [A, C], {
        strokeWidth: 2,
        dash: 2
    });

    /* Base length */
    board.create('text', [3.2, 1.35, '{#b#} m'], {
        fontSize: 16
    });

    /* Foot of triangle height on AB */
    var E = board.create('point', [2.5, 1], {
        name: '',
        visible: false
    });

    /* Triangle height */
    board.create('segment', [C, E], {
        strokeWidth: 2,
        dash: 2
    });

    board.create('text', [2.7, -0.05, '{#h_tri#} m'], {
        fontSize: 16
    });

    /* Foot of vertical pyramid height */
    var F = board.create('point', [3.5, 1], {
        name: '',
        visible: false
    });

    /* Vertical height */
    board.create('segment', [D, F], {
        strokeWidth: 2,
        dash: 2
    });

    board.create('text', [3.8, 4.0, '{#H#} m'], {
        fontSize: 16
    });

    board.unsuspendUpdate();
[[/jsxgraph]]
<p>\(\text{Volume}=\)[[input:ans1]]\(\,\text{m}^3\)[[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>Given:</strong></p>
<p>
Base of triangle, \(b = {@b@}\,\text{m}\)<br>
Height of triangle, \(h = {@h_tri@}\,\text{m}\)<br>
Vertical height of pyramid, \(H = {@H@}\,\text{m}\)
</p>

<ol type="i">
  <li>
    <p>Compute the area of the triangular base:</p>
    <p>
\(
      \begin{align}
      \text{Base area} &= \frac{1}{2} \times \text{base} \times \text{height of triangle} \\
      &= \frac{1}{2} \times {@b@} \times {@h_tri@} \\
      &= {@Base_area@}\,\text{m}^2
      \end{align}
\)
    </p>
  </li>
  <li>
    <p>Compute the volume of the pyramid using \(V = \frac{1}{3} \times \text{Base area} \times \text{vertical height}\):</p>
    <p>
\(
      \begin{align}
      V &= \frac{1}{3} \times {@Base_area@} \times {@H@} \\
      &= {@V@}\,\text{m}^3
      \end{align}
\)
    </p>
  </li>
</ol>

<p>
Therefore,<br>
The volume of the pyramid is \( {@Volume_pyramid@}\,\text{m}^3\) </p>]]></text>
    </generalfeedback>
    <defaultgrade>15</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Randomize base of the triangle (b) */
b : rand_with_step(4, 24, 2);  /* base between 4 cm and 14 cm in steps of 2 */

/* Randomize height of the triangle (h_tri) */
h_tri : rand_with_step(3, 20, 1);  /* triangle height between 3 cm and 10 cm */

/* Randomize vertical height of the pyramid (H) */
H : rand_with_step(5, 30, 1);   /* pyramid height between 5 cm and 20 cm */

/* Compute base area of the triangle */
Base_area : (1/2) * b * h_tri;

/* Compute volume of the pyramid */
V : (1/3) * Base_area * H;

/* Convert to float and round to 2 decimal places */
volume : float(V);
Volume_pyramid : decimalplaces(volume, 2);
/* Node 2 mistake: forget to multiply the base area by pyramid height */
V_mistake1: Base_area;
Volume_mistake1: decimalplaces(float(V_mistake2), 2);
</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(\text{The Base of the triangle}={@b@}\,\text{m}\)</p>
<p>\(\text{The Vertical height of the triangle}={@h_tri@}\,\text{m}\)</p>
<p>\(\text{The Vertical height of the Pyramid}={@H@}\,\text{m}\)</p>
<p>\(\text{The base area}={@Base_area@}\,\text{m}^2\)</p>
<p>\(\text{The Volume of the pyramid}={@V@}\,\text{m}^3\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>V</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct volume of the pyramid.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>V</tans>
        <testoptions>0.00000002</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Failing to use the vertical height of the pyramid</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>Volume_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have calculated the area of the triangular base, but you have not used the vertical height of the pyramid.</p>

<p>The area of the triangular base is</p>

<p>\(B=\frac{1}{2}b h_{\text{tri}}={@Base_area@}\,\text{cm}^2.\)</p>

<p>To calculate the volume of a pyramid, you need to use the base area together with the vertical height \(H\):</p>

<p>\(V=\frac{1}{3}BH.\)</p>

<p>Your answer \({@Volume_mistake2@}\,\text{cm}^3\) is the area of the triangular base, not the volume of the pyramid.</p>

<p>Try again by using the vertical height of the pyramid in the volume calculation.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>622173973</deployedseed>
    <deployedseed>1472648742</deployedseed>
    <deployedseed>1696544507</deployedseed>
    <deployedseed>303023302</deployedseed>
    <deployedseed>1960018969</deployedseed>
    <deployedseed>1435306534</deployedseed>
    <deployedseed>2141764339</deployedseed>
    <deployedseed>1386674134</deployedseed>
    <deployedseed>272415443</deployedseed>
    <deployedseed>1660317320</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>V</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
