<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Constructing frequency tables and calculating mean and median</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Mathematics test scores for \( {@n_students@} \) students in a class are listed below:</p>
<div>\({@raw_data@}\)</div>
<p><strong>a)</strong> Complete the frequency distribution table below by counting the number of scores that fall into each class interval.</p>
<table border="1">
    <thead>
        <tr style="background-color: #eee;">
            <th>Class Interval</th>
            <th>Frequency</th>
        </tr>
    </thead>
    <tbody>
        <tr><td>10 – 19</td><td>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</td></tr>
        <tr><td>20 – 29</td><td>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</td></tr>
        <tr><td>30 – 39</td><td>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</td></tr>
        <tr><td>40 – 49</td><td>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</td></tr>
        <tr><td>50 – 59</td><td>[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</td></tr>
        <tr><td>60 – 69</td><td>[[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</td></tr>
        <tr><td>70 – 79</td><td>[[input:ans7]] [[validation:ans7]] [[feedback:prt7]]</td></tr>
    </tbody>
</table>
<p><strong>b)</strong> Identify the Modal Class: [[input:ans8]] [[validation:ans8]] [[feedback:prt_mode]]</p>
<p><strong>c)</strong> Using the table you created, calculate the following estimates (round to 2 decimal places):</p>
<ul>
    <li>Mean Score: [[input:ans9]] [[validation:ans9]] [[feedback:prt_mean]]</li>
    <li>Median Score: [[input:ans10]] [[validation:ans10]] [[feedback:prt_median]]</li>
</ul>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><strong>a) Frequency Table Construction</strong></p>
<p>To fill the table, we count how many scores appear in each interval. It helps to sort the data first:</p>
<p style="font-family: monospace; font-size: 0.9em; color: #555;">{@raw_data_sorted@}</p>

<table border="1">
    <tr style="background-color: #eee;"><th>Class</th><th>Midpoint \( (x) \)</th><th>Frequency \( (f) \)</th><th>\( fx \)</th><th>Cum. Freq \( (CF) \)</th></tr>
    <tr><td>10-19</td><td>{@midpoints[1]@}</td><td>{@freqs[1]@}</td><td>{@fx_list[1]@}</td><td>{@cum_freqs[1]@}</td></tr>
    <tr><td>20-29</td><td>{@midpoints[2]@}</td><td>{@freqs[2]@}</td><td>{@fx_list[2]@}</td><td>{@cum_freqs[2]@}</td></tr>
    <tr><td>30-39</td><td>{@midpoints[3]@}</td><td>{@freqs[3]@}</td><td>{@fx_list[3]@}</td><td>{@cum_freqs[3]@}</td></tr>
    <tr><td>40-49</td><td>{@midpoints[4]@}</td><td>{@freqs[4]@}</td><td>{@fx_list[4]@}</td><td>{@cum_freqs[4]@}</td></tr>
    <tr><td>50-59</td><td>{@midpoints[5]@}</td><td>{@freqs[5]@}</td><td>{@fx_list[5]@}</td><td>{@cum_freqs[5]@}</td></tr>
    <tr><td>60-69</td><td>{@midpoints[6]@}</td><td>{@freqs[6]@}</td><td>{@fx_list[6]@}</td><td>{@cum_freqs[6]@}</td></tr>
    <tr><td>70-79</td><td>{@midpoints[7]@}</td><td>{@freqs[7]@}</td><td>{@fx_list[7]@}</td><td>{@cum_freqs[7]@}</td></tr>
    <tr style="font-weight:bold;"><td>Total</td><td></td><td>{@sum_f@}</td><td>{@sum_fx@}</td><td></td></tr>
</table>

<p><strong>b) Modal Class</strong></p>
<p>The highest frequency is \({@max_freq@}\), which corresponds to the interval <strong>{@mode_str@}</strong>.</p>

<p><strong>c) Estimated Mean and Median</strong></p>
<p><strong>Mean:</strong></p>
\( \bar{x} = \frac{\sum fx}{\sum f} = \frac{@sum_fx@}{@sum_f@} \approx \mathbf{@decimalplaces(mean_est, 2)@} \)

<p><strong>Median:</strong></p>
<p>Position = \( \frac{N}{2} = \frac{@sum_f@}{2} = {@med_pos@} \). Looking at the CF column, the median lies in the {@lower_bounds[med_idx]@} - {@upper_bounds[med_idx]@} class.</p>
\( Median = L + \left( \frac{\frac{N}{2} - CF_{prev}}{f_{med}} \right) \times h \)
\( Median = {@L_med@} + \left( \frac{{@med_pos@} - {@cf_prev@}}{@f_med@} \right) \times {@class_width@} \)
\( Median \approx \mathbf{@decimalplaces(median_est, 2)@} \)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. SETUP CLASSES AND RAW DATA */
/* Define class intervals: 10-19, 20-29, ... 70-79 */
lower_bounds: [10, 20, 30, 40, 50, 60, 70];
upper_bounds: [19, 29, 39, 49, 59, 69, 79];
class_width: 10;
num_classes: length(lower_bounds);

/* Generate 40 random test scores between 10 and 79 */
/* We sort them only for the 'worked solution' display variable, but keep raw for the question */
n_students: 40;
raw_data: makelist(rand(70)+10, k, 1, n_students);
raw_data_sorted: sort(raw_data);

/* 2. BIN THE DATA (Count Frequencies) */
/* Function to count how many numbers fall in a range [a, b] */
count_in_range(list, a, b) := length(sublist(list, lambda([x], x >= a and x <= b)));

/* Create the list of correct frequencies */
freqs: makelist(count_in_range(raw_data, lower_bounds[i], upper_bounds[i]), i, 1, num_classes);

/* 3. GROUPED STATISTICS CALCULATIONS */
/* Midpoints */
midpoints: map(lambda([l, u], (l + u)/2), lower_bounds, upper_bounds);

/* Mean Calculation (Grouped) */
fx_list: map("*", freqs, midpoints);
sum_fx: apply("+", fx_list);
sum_f: apply("+", freqs); /* Should equal n_students */
mean_est: float(sum_fx / sum_f);

/* Modal Class */
max_freq: lmax(freqs);
/* Find index of max freq */
mode_idx: sublist_indices(freqs, lambda([x], x=max_freq))[1];
/* Format string for dropdown answer */
mode_str: sconcat(lower_bounds[mode_idx], " - ", upper_bounds[mode_idx]);

/* Median Calculation (Grouped Interpolation) */
/* Cumulative Frequencies */
cum_freqs: [];
running: 0;
for f in freqs do (running: running + f, cum_freqs: endcons(running, cum_freqs));

/* Find Median Class */
med_pos: sum_f / 2;
med_idx: 1;
for i:1 thru num_classes do (
    if cum_freqs[i] >= med_pos then (med_idx: i, return())
);

/* Median Variables */
L_med: lower_bounds[med_idx]; /* Using lower bound of class (e.g. 30). Note: strictly speaking boundaries are x.5, but standard school math often uses integer lower bound. Adjust to L_med - 0.5 if your curriculum requires continuous boundaries. */
f_med: freqs[med_idx];
cf_prev: if med_idx = 1 then 0 else cum_freqs[med_idx - 1];
h: class_width;

median_est: float(L_med + ((med_pos - cf_prev) / f_med) * h);

/* 4. DISPLAY HELPERS */
/* Create a string of the raw data for display */
data_str: sconcat(raw_data);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@raw_data@}
{@mean_est@}
{@mode_str@}
{@median_est@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author"> Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">In this question, you are required to organize a raw dataset of student scores into a frequency distribution table using specified class intervals. Based on this table, you must then identify the modal class and calculate estimates for the mean and median scores.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>freqs[1]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans10</name>
      <type>algebraic</type>
      <tans>median_est</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>freqs[2]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>freqs[3]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>freqs[4]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>freqs[5]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>freqs[6]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans7</name>
      <type>algebraic</type>
      <tans>freqs[7]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans8</name>
      <type>algebraic</type>
      <tans>[mode_str]</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans9</name>
      <type>algebraic</type>
      <tans>mean_est</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt_mean</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans9</sans>
        <tans>mean_est</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_mean-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_mean-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_median</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans10</sans>
        <tans>median_est</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_median-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_median-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_mode</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans8</sans>
        <tans>[mode_str]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_mode-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_mode-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>freqs[1]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>freqs[2]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>freqs[3]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>freqs[4]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans5</sans>
        <tans>freqs[5]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans6</sans>
        <tans>freqs[6]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans7</sans>
        <tans>freqs[7]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2027977487</deployedseed>
    <deployedseed>878348757</deployedseed>
    <deployedseed>1440552277</deployedseed>
    <deployedseed>1361978873</deployedseed>
    <deployedseed>380080332</deployedseed>
    <deployedseed>380636029</deployedseed>
    <deployedseed>747807277</deployedseed>
    <deployedseed>1586289523</deployedseed>
    <deployedseed>690707360</deployedseed>
    <deployedseed>102149499</deployedseed>
    <deployedseed>1533165610</deployedseed>
    <deployedseed>1669086780</deployedseed>
    <deployedseed>1521549681</deployedseed>
    <deployedseed>1524024748</deployedseed>
    <deployedseed>1971312658</deployedseed>
    <deployedseed>1275496568</deployedseed>
    <deployedseed>973485873</deployedseed>
    <deployedseed>1523828927</deployedseed>
    <deployedseed>1466426290</deployedseed>
    <deployedseed>268184173</deployedseed>
    <deployedseed>1022511375</deployedseed>
    <deployedseed>695294049</deployedseed>
    <deployedseed>150919686</deployedseed>
    <deployedseed>810108978</deployedseed>
    <deployedseed>313236840</deployedseed>
    <deployedseed>396224084</deployedseed>
    <deployedseed>1070231102</deployedseed>
    <deployedseed>1143346597</deployedseed>
    <deployedseed>270507053</deployedseed>
    <deployedseed>2085917654</deployedseed>
    <deployedseed>284385709</deployedseed>
    <deployedseed>1147266877</deployedseed>
    <deployedseed>645613378</deployedseed>
    <deployedseed>772100547</deployedseed>
    <deployedseed>844898654</deployedseed>
    <deployedseed>1575930908</deployedseed>
    <deployedseed>203194353</deployedseed>
    <deployedseed>363148901</deployedseed>
    <deployedseed>141213600</deployedseed>
    <deployedseed>239281531</deployedseed>
    <deployedseed>1652486139</deployedseed>
    <deployedseed>296921496</deployedseed>
    <deployedseed>1311624036</deployedseed>
  </question>
</quiz>
