<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding magnitude of a given vector</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Find the magnitude of each of the following vectors. Give your answers as integers or exact surds (e.g. <code>sqrt(13)</code>).</p>
<ol class="HELM_parts">
    <li>
        <p>
            \( \mathbf{a} = {@vec_a@} \)
        </p>
        <p>
            Magnitude \( |\mathbf{a}| = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
        </p>
    </li>
    <li>
        <p>
            \( \mathbf{b} = {@vec_b@} \)
        </p>
        <p>
            Magnitude \( |\mathbf{b}| = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]
        </p>
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>
<p>The magnitude (or length) of a column vector \(\mathbf{v} = \begin{pmatrix} x \\ y \end{pmatrix}\) is given by the formula:</p>
\[ |\mathbf{v}| = \sqrt{x^2 + y^2} \]

<p><b>Part (a)</b></p>
<p>For \(\mathbf{a} = {@vec_a@}\):</p>
\[ |\mathbf{a}| = \sqrt{({@ax@})^2 + ({@ay@})^2} \]
\[ |\mathbf{a}| = \sqrt{{@ax^2@} + {@ay^2@}} \]
\[ |\mathbf{a}| = \sqrt{{@ax^2 + ay^2@}} = {@ta1@} \]

<p><b>Part (b)</b></p>
<p>For \(\mathbf{b} = {@vec_b@}\):</p>
\[ |\mathbf{b}| = \sqrt{({@bx@})^2 + ({@by@})^2} \]
\[ |\mathbf{b}| = \sqrt{{@bx^2@} + {@by^2@}} \]
\[ |\mathbf{b}| = \sqrt{{@bx^2 + by^2@}} \]
<p>
    If the number under the square root is not a perfect square, we leave it in exact surd form:
    \( \mathbf{@ta2@} \)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025073100</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Setup for Part (a): Pythagorean Triples (Integer Magnitude) */
triples: [[3,4], [4,3], [6,8], [8,6], [5,12], [12,5], [8,15], [15,8]];
sel: rand(triples);
s1: rand([-1, 1]);
s2: rand([-1, 1]);

ax: sel[1] * s1;
ay: sel[2] * s2;

/* Define vector a */
vec_a: matrix([ax],[ay]);
/* Answer for a */
ta1: sqrt(ax^2 + ay^2);

/* 2. Setup for Part (b): Random Vector */
bx: rand_with_step(-9, 9, 1);
by: rand_with_step(-9, 9, 1);
/* Avoid zero vector */
bx: if bx=0 and by=0 then 3 else bx;

/* Define vector b */
vec_b: matrix([bx],[by]);
/* Answer for b */
ta2: sqrt(bx^2 + by^2);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}
{@ta2@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Daniel Murunga</p>
<p id="concept"></p>
<p id="reviewer"></p>
<p id="reviewer"></p>
<p id="description">Calculates the magnitude of two 2D column vectors.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>222549107</deployedseed>
    <deployedseed>1065838551</deployedseed>
    <deployedseed>608563890</deployedseed>
    <deployedseed>1882287494</deployedseed>
    <deployedseed>893359562</deployedseed>
    <deployedseed>1406264140</deployedseed>
    <deployedseed>1922568966</deployedseed>
    <deployedseed>1182395856</deployedseed>
    <deployedseed>496337036</deployedseed>
    <deployedseed>721068846</deployedseed>
    <deployedseed>678125668</deployedseed>
    <deployedseed>663870359</deployedseed>
    <deployedseed>1082749579</deployedseed>
    <deployedseed>1990852393</deployedseed>
    <deployedseed>629539443</deployedseed>
    <deployedseed>371098598</deployedseed>
    <deployedseed>518912456</deployedseed>
    <deployedseed>365364246</deployedseed>
    <deployedseed>137170296</deployedseed>
    <deployedseed>984869845</deployedseed>
    <deployedseed>489549834</deployedseed>
    <deployedseed>1121797739</deployedseed>
    <deployedseed>1468591010</deployedseed>
    <deployedseed>963145766</deployedseed>
    <deployedseed>884634108</deployedseed>
    <deployedseed>223358754</deployedseed>
    <deployedseed>1378906282</deployedseed>
    <deployedseed>1723763475</deployedseed>
    <deployedseed>710816419</deployedseed>
    <deployedseed>222225738</deployedseed>
    <deployedseed>653876700</deployedseed>
    <deployedseed>1271758015</deployedseed>
    <deployedseed>967600305</deployedseed>
    <deployedseed>693935376</deployedseed>
    <deployedseed>878841359</deployedseed>
    <deployedseed>701883537</deployedseed>
    <deployedseed>663437654</deployedseed>
    <deployedseed>1585755388</deployedseed>
    <deployedseed>878053638</deployedseed>
    <deployedseed>549838533</deployedseed>
    <deployedseed>334011902</deployedseed>
    <deployedseed>558213286</deployedseed>
  </question>
</quiz>
