<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating slant height of the original cone and volume a frustum</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A solid right circular cone has a base radius of <b>{@R@} cm</b> and a vertical height of <b>{@H@} cm</b>. A smaller cone is cut off from the top to leave a frustum of height <b>{@h_frustum@} cm</b> and top radius <b>{@r@} cm</b> as shown below.</p>

[[jsxgraph width="400px" height="350px"]]
var board = JXG.JSXGraph.initBoard(divid, {
    /* Expanded boundingbox: [left, top, right, bottom] */
    boundingbox: [-7, 12, 7, -3], 
    axis: false,
    showCopyright: false
});

var R = {#Rdraw#};
var r = {#rdraw#};
var h = {#hdraw#};
var H = {#Hdraw#};

/* 1. Points */
var pBottomLeft = board.create('point', [-R, 0], {visible: false});
var pBottomRight = board.create('point', [R, 0], {visible: false});
var pTopLeft = board.create('point', [-r, h], {visible: false});
var pTopRight = board.create('point', [r, h], {visible: false});
var apex = board.create('point', [0, H], {visible: false});

/* 2. Frustum Sides */
board.create('segment', [pBottomLeft, pTopLeft], {strokeColor: 'black', strokeWidth: 2});
board.create('segment', [pBottomRight, pTopRight], {strokeColor: 'black', strokeWidth: 2});

/* 3. Top Ellipse (Solid) */
board.create('curve', [
    function(t){ return r * Math.cos(t); },
    function(t){ return h + (r * 0.25) * Math.sin(t); },
    0, 2 * Math.PI
], {strokeColor: 'black', strokeWidth: 2, fillColor: '#f0f0f0', fillOpacity: 0.3});

/* 4. Bottom Ellipse (Front solid, Back dashed) */
board.create('curve', [
    function(t){ return R * Math.cos(t); },
    function(t){ return (R * 0.25) * Math.sin(t); },
    Math.PI, 2 * Math.PI
], {strokeColor: 'black', strokeWidth: 2});
board.create('curve', [
    function(t){ return R * Math.cos(t); },
    function(t){ return (R * 0.25) * Math.sin(t); },
    0, Math.PI
], {strokeColor: 'black', strokeWidth: 1, dash: 2});

/* 5. "Ghost" Lines to Apex */
board.create('segment', [pTopLeft, apex], {dash: 2, strokeColor: 'gray'});
board.create('segment', [pTopRight, apex], {dash: 2, strokeColor: 'gray'});

/* 6. Indicators and Labels */
// Height (Red)
board.create('segment', [[0, 0], [0, H]], {dash: 1, strokeColor: 'red'});
board.create('text', [-2.5, H/2, "{#H#} cm"], {color: 'red', fontSize: 14, fontWeight: 'bold'});

// Radius (Blue)
board.create('segment', [[0, 0], [R, 0]], {strokeColor: 'blue', strokeWidth: 2});
board.create('text', [R/2, -1.2, "{#R#} cm"], {color: 'blue', fontSize: 14, fontWeight: 'bold'});

[[/jsxgraph]]
<p><b>(a)</b>  Calculate the slant height (L) of the original full cone.</p>
<p> Slant height = [[input:ans1]] [[validation:ans1]][[feedback:prt1]] </p>
<p><b>(b)</b> Calculate the volume of the remaining frustum. </p>
<p>Volume = [[input:ans2]] [[validation:ans2]][[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution </h3>
<p><b>Part (a): Slant Height of the Original Cone</b></p>
Using Pythagoras' Theorem:
\[ L^2 = R^2 + H^2 \]
\[ L^2 = {@R@}^2 + {@H@}^2 = {@R^2 + H^2@} \]
\[ L = \sqrt{{@R^2 + H^2@}}\]
<p><b>Part (b): Volume of the Frustum</b></p>
We use the formula: \( V = \frac{1}{3}\pi h(R^2 + Rr + r^2) \)
<p>
    Substituting our values:<br>
    \( R = {@R@} \), \( r = {@r@} \), \( h = {@h_frustum@} \)
</p>
<p>
\( \begin{align} V &= \frac{1}{3} \times \pi \times {@h_frustum@} \times ({@R^2@} + {@R*r@} + {@r^2@}) \\ V &= \frac{1}{3} \times \pi \times {@h_frustum@} \times {@R^2 + R*r + r^2@} \\ V &= {@ta@} \text{ cm}^3 \end{align}\)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Pick radii such that (R-r) is a simple small number */
R: 10 + rand(6);
r: 5 + rand(3);
diff: R - r;

/* 2. Pick h_frustum as a multiple of the difference to keep H a whole number */
h_factor: 2 + rand(3); 
h_frustum: h_factor * diff;

/* 3. Calculate full height H */
H: (R * h_frustum) / diff;

/* 4. NEW: Calculate Slant Height L using Pythagoras */
/* We use float() because this will likely be a decimal */
L: R^2 + H^2;

/* 5. Calculation variables for the General Feedback */
bracket_value: R^2 + R*r + r^2;
height_bracket_value: bracket_value * h_frustum;
V: (1/3) * %pi * height_bracket_value;
ta: V / 1000;

/* 6. JSXGraph Scaling */
Rdraw: R/4;
rdraw: r/4;
hdraw: h_frustum/4;
Hdraw: H/4;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@L@}</p>
<p>{@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>sqrt(L)</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct value of the slight height</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>sqrt(L)</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for the correct volume of the frustum</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1526049761</deployedseed>
    <deployedseed>839969387</deployedseed>
    <deployedseed>1897158016</deployedseed>
    <deployedseed>1133287341</deployedseed>
    <deployedseed>1591924957</deployedseed>
    <deployedseed>959817895</deployedseed>
    <deployedseed>1356865330</deployedseed>
    <deployedseed>1592618242</deployedseed>
    <deployedseed>2034409984</deployedseed>
    <deployedseed>1114553496</deployedseed>
    <deployedseed>36235439</deployedseed>
    <deployedseed>523332874</deployedseed>
    <deployedseed>1727166544</deployedseed>
    <deployedseed>1840273182</deployedseed>
    <deployedseed>2082496739</deployedseed>
    <deployedseed>1433336537</deployedseed>
    <deployedseed>388745564</deployedseed>
    <deployedseed>109915867</deployedseed>
    <deployedseed>1760173120</deployedseed>
    <deployedseed>323951327</deployedseed>
  </question>
</quiz>
