<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the centre and angle of rotation of a symmetrical layout for a city plaza</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A designer is creating a symmetrical layout for a city plaza. A decorative stone bench, represented by the segment <b>AB</b>, is rotated around a central fountain <b>O</b> to create a matching bench <b>A'B'</b>.</p>

<ul>
    <li><b>Bench AB:</b> A({@pA[1]@}, {@pA[2]@}) to B({@pB[1]@}, {@pB[2]@})</li>
    <li><b>Bench A'B':</b> A'({@pA_prime[1]@}, {@pA_prime[2]@}) to B'({@pB_prime[1]@}, {@pB_prime[2]@})</li>
</ul>

[[jsxgraph width="450px" height="450px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox: [-10, 10, 10, -10], axis: true, showCopyright: false});
var seg1 = board.create('segment', [{#pA#}, {#pB#}], {strokeColor: 'blue', strokeWidth: 4, name: 'AB', withLabel: true});
var seg2 = board.create('segment', [{#pA_prime#}, {#pB_prime#}], {strokeColor: 'red', strokeWidth: 4, name: "A'B'", withLabel: true});
[[/jsxgraph]]


<ol>
    <li>Use the perpendicular bisector method to find the coordinates of the fountain (Centre O). <br>
        O = [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]
    </li>
    <li>Determine the angle of rotation. <br>
        Angle = [[input:ans2]]° [[validation:ans2]] [[feedback:prt2]]
    </li>
</ol>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<b>Finding the Centre</b>
<p>We join <b>A to A'</b> and <b>B to B'</b> with dotted lines. The intersection of their perpendicular bisectors identifies the centre of rotation.</p>

[[jsxgraph width="350px" height="350px"]]
var b1 = JXG.JSXGraph.initBoard(divid, {boundingbox: [-10, 10, 10, -10], axis: true, showCopyright: false});
var A = b1.create('point', [{#pA[1]#}, {#pA[2]#}], {name: 'A', color: 'blue'});
var Ap = b1.create('point', [{#ax_p#}, {#ay_p#}], {name: "A'", color: 'red'});
var B = b1.create('point', [{#pB[1]#}, {#pB[2]#}], {name: 'B', color: 'blue'});
var Bp = b1.create('point', [{#bx_p#}, {#by_p#}], {name: "B'", color: 'red'});

b1.create('segment', [A, Ap], {dash: 2, strokeColor: 'gray'});
b1.create('segment', [B, Bp], {dash: 2, strokeColor: 'gray'});
b1.create('perpendicularBisector', [A, Ap], {strokeColor: 'purple'});
b1.create('perpendicularBisector', [B, Bp], {strokeColor: 'orange'});
b1.create('point', [{#cx#}, {#cy#}], {name: 'O', color: 'black', size: 4});
[[/jsxgraph]]

<b>Identifying the Angle</b>
<p>By connecting the original point <b>A</b> and its image <b>A'</b> to the centre <b>O</b>, we can measure the rotation angle \(\angle AOA'\).</p>
<p>The rotation for this specific layout is <b>{#ta2#}°</b>.</p>

[[jsxgraph width="350px" height="350px"]]
var b2 = JXG.JSXGraph.initBoard(divid, {boundingbox: [-10, 10, 10, -10], axis: true, showCopyright: false});

// Define key points
var O = b2.create('point', [{#cx#}, {#cy#}], {name: 'O', color: 'black', size: 5, fixed: true});
var A = b2.create('point', [{#ax_orig#}, {#ay_orig#}], {name: 'A', color: 'blue', size: 3, fixed: true});
var Ap = b2.create('point', [{#ax_p#}, {#ay_p#}], {name: "A'", color: 'red', size: 3, fixed: true});

// Draw the lines from centre to points to visualize the angle
b2.create('segment', [O, A], {dash: 1, strokeColor: 'gray', strokeWidth: 2});
b2.create('segment', [O, Ap], {dash: 1, strokeColor: 'gray', strokeWidth: 2});

// Visual Angle arc (ta2 is the angle)
b2.create('angle', [Ap, O, A], {
    radius: 1.5, 
    name: '{#abs(ta2)#}°', 
    withLabel: true,
    label: {offset: [15, 15]}
});
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* 1. Randomize Centre O */
cx: -2 + rand(5); 
cy: -2 + rand(5); 
ta1: [cx, cy];

/* 2. Randomize Angle */
angle_list: [90, -90, 180, 270];
ta2: random_permutation(angle_list)[1];

/* 3. Define the 'Bench' (Line Segment AB) */
/* Coordinates are relative to centre to ensure they stay on the grid */
ax_orig: cx + 2 + rand(2);
ay_orig: cy + 1 + rand(2);
bx_orig: ax_orig - 1;
by_orig: ay_orig - 2;

pA: [ax_orig, ay_orig];
pB: [bx_orig, by_orig];

/* 4. Rotation Logic - CRITICAL: Pre-calculate as pure numbers */
rad: (ta2 * %pi) / 180;

/* Calculate A' */
ax_p: round(float(cx + (pA[1]-cx)*cos(rad) - (pA[2]-cy)*sin(rad)));
ay_p: round(float(cy + (pA[1]-cx)*sin(rad) + (pA[2]-cy)*cos(rad)));
pA_prime: [ax_p, ay_p];

/* Calculate B' */
bx_p: round(float(cx + (pB[1]-cx)*cos(rad) - (pB[2]-cy)*sin(rad)));
by_p: round(float(cy + (pB[1]-cx)*sin(rad) + (pB[2]-cy)*cos(rad)));
pB_prime: [bx_p, by_p];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>{@ta1@}</p>
<p>{@ta2@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1779462980</deployedseed>
    <deployedseed>219232324</deployedseed>
    <deployedseed>1225511114</deployedseed>
    <deployedseed>2136155793</deployedseed>
    <deployedseed>302496037</deployedseed>
    <deployedseed>1796596506</deployedseed>
    <deployedseed>780149011</deployedseed>
    <deployedseed>805578650</deployedseed>
    <deployedseed>1257085255</deployedseed>
    <deployedseed>1540074191</deployedseed>
    <deployedseed>1377044641</deployedseed>
    <deployedseed>1159240542</deployedseed>
    <deployedseed>1517252299</deployedseed>
    <deployedseed>1022952175</deployedseed>
    <deployedseed>261094242</deployedseed>
    <deployedseed>1306203723</deployedseed>
    <deployedseed>35099174</deployedseed>
    <deployedseed>1028222410</deployedseed>
    <deployedseed>782302633</deployedseed>
    <deployedseed>2043585660</deployedseed>
    <deployedseed>592220845</deployedseed>
    <deployedseed>442338001</deployedseed>
    <deployedseed>323002357</deployedseed>
    <deployedseed>869964364</deployedseed>
    <deployedseed>1890222978</deployedseed>
    <deployedseed>580150959</deployedseed>
    <deployedseed>1377990656</deployedseed>
    <deployedseed>1009825505</deployedseed>
    <deployedseed>2125671906</deployedseed>
    <deployedseed>678238883</deployedseed>
    <deployedseed>627566283</deployedseed>
    <deployedseed>1342215923</deployedseed>
    <deployedseed>1400834361</deployedseed>
    <deployedseed>1870151034</deployedseed>
    <deployedseed>1280662745</deployedseed>
    <deployedseed>151264580</deployedseed>
    <deployedseed>1442978955</deployedseed>
    <deployedseed>1391228881</deployedseed>
    <deployedseed>2121991457</deployedseed>
    <deployedseed>1265219446</deployedseed>
  </question>
</quiz>
