<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Application of reciprocals - period of a pendulum</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The <strong>period</strong> (\(T\)) of a swinging pendulum is the reciprocal of its <strong>frequency</strong> (\(f\)). The relationship is given by the formula:</p>
\(T = \frac{1}{f} \)

<p>This means: if you know the frequency, you can find the period by taking <strong>1 divided by the frequency</strong>.</p><p>A digital sensor measures a pendulum's frequency as <strong>{@freq_disp@} swings per second (Hz)</strong>.</p>
<p>Use the <strong>table of reciprocals</strong> provided to calculate the pendulum's period in seconds.</p>
<p>Period (\(T\)) \(=\) [[input:ans1]] s [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>We are asked to find the period of a pendulum. The period <strong>T</strong> is the reciprocal of the frequency:</p>

\[ T = \frac{1}{{@freq_disp@}} \]

<p>We use a mathematical table of reciprocals to find this value:</p>

<ol>
  <li>Locate the row corresponding to <strong>{@float(part_a)@}</strong> in the column labelled \(x\).</li>
  <li>Move across to the main column headed <strong>{@part_b@}</strong>. The value at this intersection is <strong>{@table_val_main@}</strong>.</li>
  <li>Continue along the same row to the "SUBTRACT" (Mean Difference) column headed <strong>{@part_c@}</strong>. The value here corresponds to <strong>{@sub_num@}</strong>, representing 0.000{@sub_num@}.</li>
  <li>Subtract the mean difference from the main table value.  
      (We subtract because the reciprocal function \(y = 1/x\) decreases as \(x\) increases.)</li>
</ol>

<p>Calculation:</p>
\[ {@table_val_main@} - 0.000{@sub_num@} = {@ta@} \]

<p>Thus, the period of the pendulum is <strong>{@ta@} seconds</strong>.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* 1. Generate the Frequency components (1.000 to 9.999) */
d1 : rand(8) + 1;       /* Integer part: 1-9 (Avoid 0) */
d2 : rand(10);          /* 1st decimal: 0-9 */
d3 : rand(10);          /* 2nd decimal: 0-9 */
d4 : rand(9) + 1;       /* 3rd decimal (mean difference): 1-9 */

/* Construct the Frequency 'f' */
freq : d1 + d2/10 + d3/100 + d4/1000;
freq : float(freq);

/* 2. Simulate the Table Lookup Values */
/* 'part_a' is the Row (e.g., 2.3) */
part_a : d1 + d2/10;

/* 'part_b' is the Main Column (e.g., 4) */
part_b : d3;

/* 'part_c' is the Difference Column (e.g., 7) */
part_c : d4;

/* 3. Calculate the Table Values */
/* Calculate the value found in the main body (1 / d1.d2d3) rounded to 4dp */
val_lookup_base : d1 + d2/10 + d3/100;
table_val_main : decimalplaces(1/val_lookup_base, 4);

/* Calculate the exact final answer (1/f) */
ta : decimalplaces(1/freq, 4);

/* Calculate the 'Subtract' value (Mean Difference) */
/* We reverse-engineer the difference to ensure the feedback math works perfectly */
/* Difference = (Main Table Value) - (True Answer) */
diff_val : table_val_main - ta;

/* Convert this to the integer seen in the "Subtract" column (e.g., 0.0005 -> 5) */
sub_num : round(diff_val * 10000);

/* Safety: Ensure sub_num is positive (Reciprocals always decrease) */
sub_num : abs(sub_num);

/* Formatting for display */
/* Ensures 4.500 displays as 4.500 not 4.5 in some contexts if needed */
freq_disp : dispdp(freq, 3);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Sheila Cherotich</p>
 <p id="concept"></p> 
<p id="reviewer"></p> 
<p id="description">This question requires the learner to find the reciprocal of a given decimal number using a table.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>830071103</deployedseed>
    <deployedseed>1873782826</deployedseed>
    <deployedseed>1803225670</deployedseed>
    <deployedseed>236680742</deployedseed>
    <deployedseed>518282498</deployedseed>
    <deployedseed>752548061</deployedseed>
    <deployedseed>1476275683</deployedseed>
    <deployedseed>1050836297</deployedseed>
    <deployedseed>1901855063</deployedseed>
    <deployedseed>579234699</deployedseed>
    <deployedseed>28783169</deployedseed>
    <deployedseed>1120616938</deployedseed>
    <deployedseed>1969746631</deployedseed>
    <deployedseed>213984693</deployedseed>
    <deployedseed>2143482703</deployedseed>
    <deployedseed>1607864787</deployedseed>
    <deployedseed>2049455653</deployedseed>
    <deployedseed>2088501389</deployedseed>
    <deployedseed>2063483827</deployedseed>
    <deployedseed>1721866342</deployedseed>
    <deployedseed>856206045</deployedseed>
    <deployedseed>1294336658</deployedseed>
    <deployedseed>1104149188</deployedseed>
    <deployedseed>232415072</deployedseed>
    <deployedseed>792634296</deployedseed>
    <deployedseed>1484641015</deployedseed>
    <deployedseed>99812067</deployedseed>
    <deployedseed>671227243</deployedseed>
    <deployedseed>575746081</deployedseed>
    <deployedseed>1832310373</deployedseed>
    <deployedseed>906111223</deployedseed>
    <deployedseed>1254850404</deployedseed>
    <deployedseed>302630984</deployedseed>
    <deployedseed>1365757643</deployedseed>
    <deployedseed>1839632096</deployedseed>
    <deployedseed>1722376711</deployedseed>
    <deployedseed>1192557368</deployedseed>
    <deployedseed>814839088</deployedseed>
    <deployedseed>1530952977</deployedseed>
    <deployedseed>540168453</deployedseed>
    <deployedseed>1859422338</deployedseed>
    <deployedseed>933651225</deployedseed>
    <deployedseed>1379005282</deployedseed>
    <deployedseed>1648108777</deployedseed>
    <deployedseed>1850975519</deployedseed>
    <deployedseed>1744734768</deployedseed>
    <deployedseed>1843408851</deployedseed>
    <deployedseed>1049763915</deployedseed>
    <deployedseed>432079243</deployedseed>
    <deployedseed>508523524</deployedseed>
    <deployedseed>1902709901</deployedseed>
    <deployedseed>1571664558</deployedseed>
    <deployedseed>638064452</deployedseed>
    <deployedseed>226479867</deployedseed>
    <deployedseed>1521372871</deployedseed>
    <deployedseed>1632618651</deployedseed>
    <deployedseed>1127599662</deployedseed>
    <deployedseed>1614825483</deployedseed>
    <deployedseed>959212604</deployedseed>
    <deployedseed>1114715748</deployedseed>
    <deployedseed>1604434882</deployedseed>
    <deployedseed>352900733</deployedseed>
    <deployedseed>1075315992</deployedseed>
    <deployedseed>756816423</deployedseed>
    <deployedseed>565189991</deployedseed>
    <deployedseed>1518507998</deployedseed>
    <deployedseed>1325673657</deployedseed>
    <deployedseed>359738678</deployedseed>
    <deployedseed>381384752</deployedseed>
    <deployedseed>1689233774</deployedseed>
    <deployedseed>811250208</deployedseed>
    <deployedseed>297372699</deployedseed>
    <deployedseed>1874366312</deployedseed>
    <deployedseed>531002598</deployedseed>
    <deployedseed>361418810</deployedseed>
    <deployedseed>1313635934</deployedseed>
    <deployedseed>914617142</deployedseed>
    <deployedseed>2019052632</deployedseed>
    <deployedseed>1404960370</deployedseed>
    <deployedseed>1857133599</deployedseed>
    <deployedseed>741353251</deployedseed>
    <deployedseed>1455104109</deployedseed>
    <deployedseed>243119816</deployedseed>
    <deployedseed>664984458</deployedseed>
    <deployedseed>986637094</deployedseed>
    <deployedseed>757865879</deployedseed>
    <deployedseed>1093566903</deployedseed>
    <deployedseed>242825257</deployedseed>
    <deployedseed>695168685</deployedseed>
    <deployedseed>914890441</deployedseed>
    <deployedseed>1639711684</deployedseed>
    <deployedseed>1171824884</deployedseed>
    <deployedseed>1469881543</deployedseed>
    <deployedseed>1759708046</deployedseed>
    <deployedseed>1443236499</deployedseed>
    <deployedseed>1773674595</deployedseed>
    <deployedseed>64854711</deployedseed>
    <deployedseed>569275471</deployedseed>
    <deployedseed>2051228357</deployedseed>
    <deployedseed>1520727387</deployedseed>
  </question>
</quiz>
