<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a square</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A square has a side length of <b>{@s@} cm</b>. Find the area of the square.
</p>

[[jsxgraph]]
    var s = {#s#};
    
    var board = JXG.JSXGraph.initBoard(BOARDID, {
        boundingbox: [-2, s + 3, s + 3, -2],
        keepAspectRatio: true,
        showNavigation: false,
        showCopyright: false,
        pan: {enabled: false},
        zoom: {enabled: false}
    });
    
    // Vertices locked in place
    var p1 = board.create('point', [0, 0], {visible: false, fixed: true});
    var p2 = board.create('point', [s, 0], {visible: false, fixed: true});
    var p3 = board.create('point', [s, s], {visible: false, fixed: true});
    var p4 = board.create('point', [0, s], {visible: false, fixed: true});
    
    // Square Polygon locked in place
    var poly = board.create('polygon', [p1, p2, p3, p4], {
        fillColor: '#e8eaf6', fillOpacity: 0.8, fixed: true, highlight: false,
        borders: {strokeColor: '#3949ab', strokeWidth: 2.5, fixed: true, highlight: false}
    });
    
    // Right angle markers in all 4 corners (names explicitly cleared to remove the alpha)
    var rad = Math.min(s*0.15, 2);
    board.create('angle', [p2, p1, p4], {radius: rad, type: 'square', strokeColor: 'black', fixed: true, name: ''});
    board.create('angle', [p3, p2, p1], {radius: rad, type: 'square', strokeColor: 'black', fixed: true, name: ''});
    board.create('angle', [p4, p3, p2], {radius: rad, type: 'square', strokeColor: 'black', fixed: true, name: ''});
    board.create('angle', [p1, p4, p3], {radius: rad, type: 'square', strokeColor: 'black', fixed: true, name: ''});
    
    // Labels
    board.create('text', [-0.5, s/2, s + " cm"], {anchorX: 'right', anchorY: 'middle', fontSize: 14, fixed: true});
[[/jsxgraph]]

<p>
Area = [[input:ans1]] <b>\(\text{cm}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
The area of a square is found using the formula:
</p>

<p>
\[
A = s^2
\]
</p>

<p>
Here, the side length is:
</p>

<p>
\[
s = {@s@}\text{ cm}
\]
</p>

<p>
Step 1: Substitute the value of the side into the formula.
</p>

<p>
\[
A = {@s@}^2
\]
</p>

<p>
Step 2: Multiply the side by itself.
</p>

<p>
\[
A = {@s@} \times {@s@}
\]
</p>

<p>
\[
A = {@A@}
\]
</p>

<p>
Therefore, the area of the square is:
</p>

<p>
\[
\boxed{{@A@}\text{ cm}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Side length of the square */
s: rand(30) + 3;

/* Area */
A: s^2;

/* Common mistakes */
m1: 4*s;        /* perimeter instead of area */
m2: 2*s;        /* doubles the side */
m3: s+s;        /* repeated addition confusion */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>s: {@s@}, Area: {@A@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of squares and rectangles</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to calculate the area of a square by applying the formula A = s².
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>A</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>perimeter instead of area </description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You found the perimeter. For the area of a square, multiply the side by itself (\(s^2\)).</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>doubling the sides</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Area is not found by doubling the side. Use side &times; side.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>repeated addition confusion</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The area of a square is found by multiplying the side by itself, not by adding the sides together.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1608528124</deployedseed>
    <deployedseed>247696548</deployedseed>
    <deployedseed>1847533907</deployedseed>
    <deployedseed>1644553511</deployedseed>
    <deployedseed>765900417</deployedseed>
    <deployedseed>1186568348</deployedseed>
    <deployedseed>542942347</deployedseed>
    <deployedseed>1187946555</deployedseed>
    <deployedseed>975563647</deployedseed>
    <deployedseed>599361331</deployedseed>
    <deployedseed>1742718192</deployedseed>
    <deployedseed>2017546093</deployedseed>
    <deployedseed>464319046</deployedseed>
    <deployedseed>1889345483</deployedseed>
    <deployedseed>1075674304</deployedseed>
    <deployedseed>967859772</deployedseed>
    <deployedseed>1425515208</deployedseed>
    <deployedseed>628245635</deployedseed>
    <deployedseed>719913301</deployedseed>
    <deployedseed>1937375162</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
