<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a rectangular school garden</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A school is preparing a rectangular garden for planting <b>{@crop@}</b>. The garden measures <b>{@l_disp@} m</b> in length and <b>{@w_disp@} m</b> in width. Find the area of the garden.
</p>

<p>
Give your answer correct to <b>2 decimal places</b>.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{m}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
The garden is rectangular.
</p>

<p>
The formula for the area of a rectangle is:
</p>

<p>
\[
A = l \times w
\]
</p>

<p>
Step 1: Identify the dimensions.
</p>

<p>
\[
l = {@l_disp@}\text{ m}, \quad w = {@w_disp@}\text{ m}
\]
</p>

<p>
Step 2: Substitute into the formula.
</p>

<p>
\[
A = {@l_disp@} \times {@w_disp@}
\]
</p>

<p>
Step 3: Multiply.
</p>

<p>
\[
A \approx {@A_disp@}
\]
</p>

<p>
Therefore, the area of the garden is:
</p>

<p>
\[
\boxed{{@A_disp@}\text{ m}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Length and width as decimals */
l: float((rand(120)+90)/10);
w: float((rand(50)+30)/10);

/* Display formatting to prevent floating-point artifacts */
l_disp: dispdp(l, 1);
w_disp: dispdp(w, 1);

/* Area calculated natively as float */
A_float: float(l * w);

/* Rounded display answer (2 d.p.) */
A_disp: dispdp(A_float, 2);

/* Context */
crop: rand(["vegetables","flowers","maize","spinach","kales"]);

/* Mistakes evaluated as floats for PRT catching */
m1: float(2*(l+w));
m2: float(l+w);
m3: float(l^2);
m4: float(w^2);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>l: {@l_disp@}, w: {@w_disp@}, Area: {@A_disp@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of squares and rectangles</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to determine the area of a rectangular garden in a real-life context by applying the formula A = l × w.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A_disp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>NumDecPlaces</answertest>
        <sans>ans1</sans>
        <tans>A_float</tans>
        <testoptions>2</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>perimeter</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You calculated the perimeter instead of the area. To find the area of a rectangle, multiply the length by the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>adds dimensions</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You added the dimensions. The area is found by <strong>multiplying</strong> the length and the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>squares length</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You squared the length. Because this is a rectangle, you must multiply the length by the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>squares width</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m4</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You squared the width. Because this is a rectangle, you must multiply the length by the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2012224842</deployedseed>
    <deployedseed>612503299</deployedseed>
    <deployedseed>848392839</deployedseed>
    <deployedseed>1763148486</deployedseed>
    <deployedseed>1008813701</deployedseed>
    <deployedseed>1118095198</deployedseed>
    <deployedseed>460607550</deployedseed>
    <deployedseed>1308087481</deployedseed>
    <deployedseed>477296797</deployedseed>
    <deployedseed>1046566553</deployedseed>
    <deployedseed>312278076</deployedseed>
    <deployedseed>837493444</deployedseed>
    <deployedseed>1662798185</deployedseed>
    <deployedseed>1888536111</deployedseed>
    <deployedseed>1409045429</deployedseed>
    <deployedseed>475491556</deployedseed>
    <deployedseed>1019483805</deployedseed>
    <deployedseed>804317582</deployedseed>
    <deployedseed>1904997791</deployedseed>
    <deployedseed>1362081613</deployedseed>
    <deployedseed>503729524</deployedseed>
    <deployedseed>228445311</deployedseed>
    <deployedseed>2094770312</deployedseed>
    <deployedseed>658841905</deployedseed>
    <deployedseed>1243310607</deployedseed>
    <deployedseed>907964649</deployedseed>
    <deployedseed>1843743268</deployedseed>
    <deployedseed>1434892691</deployedseed>
    <deployedseed>46002749</deployedseed>
    <deployedseed>142401052</deployedseed>
    <deployedseed>131935123</deployedseed>
    <deployedseed>382831345</deployedseed>
    <deployedseed>86049053</deployedseed>
    <deployedseed>1578704938</deployedseed>
    <deployedseed>1878022713</deployedseed>
    <deployedseed>451896223</deployedseed>
    <deployedseed>1043575106</deployedseed>
    <deployedseed>1436621730</deployedseed>
    <deployedseed>1696062878</deployedseed>
    <deployedseed>338038974</deployedseed>
    <deployedseed>1700824425</deployedseed>
    <deployedseed>562573041</deployedseed>
    <deployedseed>705847211</deployedseed>
    <deployedseed>1866541925</deployedseed>
    <deployedseed>2025485610</deployedseed>
    <deployedseed>344371540</deployedseed>
    <deployedseed>1618293847</deployedseed>
    <deployedseed>440419060</deployedseed>
    <deployedseed>638905783</deployedseed>
    <deployedseed>1492313219</deployedseed>
    <deployedseed>28287909</deployedseed>
    <deployedseed>1770445797</deployedseed>
    <deployedseed>460090171</deployedseed>
    <deployedseed>605550372</deployedseed>
    <deployedseed>735366687</deployedseed>
    <deployedseed>1537282293</deployedseed>
    <deployedseed>663050676</deployedseed>
    <deployedseed>679097604</deployedseed>
    <deployedseed>419636206</deployedseed>
    <deployedseed>816839737</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A_disp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>( prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
