<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the area of a rectangular classroom floor</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
A classroom floor is rectangular. Its length is <b>{@l@} m</b> and its width is <b>{@w@} m</b>. Find the area of the floor.
</p>

<p>
Area = [[input:ans1]] <b>\(\text{m}^2\)</b> [[validation:ans1]] [[feedback:prt1]]
</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution</h3>

<p>
The area of a rectangle is found using the formula:
</p>

<p>
\[
A = l \times w
\]
</p>

<p>
Here,
</p>

<p>
\[
l = {@l@}\text{ m}, \quad w = {@w@}\text{ m}
\]
</p>

<p>
Step 1: Substitute the given values into the formula.
</p>

<p>
\[
A = {@l@} \times {@w@}
\]
</p>

<p>
Step 2: Multiply.
</p>

<p>
\[
A = {@A@}
\]
</p>

<p>
Therefore, the area of the classroom floor is:
</p>

<p>
\[
\boxed{{@A@}\text{ m}^2}
\]
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>/* Length and width */
l: rand(20) + 10;
w: rand(6) + 3;

/* Area */
A: l*w;

/* Common mistakes */
m1: 2*(l+w);    /* perimeter */
m2: l+w;        /* adds dimensions */
m3: l^2;        /* squares length */
m4: w^2;        /* squares width */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>l: {@l@}, w: {@w@}, Area: {@A@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="author">Michael Onyimbo</p>
<p id="concept">Area of squares and rectangles</p>
<p id="reviewer"></p>
<p id="reviewer"></p>

<p id="description">
The question requires the learner to calculate the area of a rectangle in a real-life context by applying the formula A = l × w.
</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>A</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>checks for the correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>A</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>perimeter</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You found the perimeter. For the area of a rectangle, multiply length by width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>adds dimensions</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Do not add the dimensions. Area is found by multiplying length and width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>squares length</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-4-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is a rectangle, not a square. You squared the length. Use length &times; width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-4-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>squares width</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>m4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-5-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is a rectangle, not a square. You squared the width. Use length &times; width.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-5-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>826014208</deployedseed>
    <deployedseed>34605616</deployedseed>
    <deployedseed>1928714838</deployedseed>
    <deployedseed>345767455</deployedseed>
    <deployedseed>917711036</deployedseed>
    <deployedseed>236526306</deployedseed>
    <deployedseed>164059010</deployedseed>
    <deployedseed>597288230</deployedseed>
    <deployedseed>1522772841</deployedseed>
    <deployedseed>1592625685</deployedseed>
    <deployedseed>400869514</deployedseed>
    <deployedseed>4781190</deployedseed>
    <deployedseed>317422933</deployedseed>
    <deployedseed>1519884517</deployedseed>
    <deployedseed>1697948877</deployedseed>
    <deployedseed>1374079791</deployedseed>
    <deployedseed>623482548</deployedseed>
    <deployedseed>1482310016</deployedseed>
    <deployedseed>736614171</deployedseed>
    <deployedseed>481037573</deployedseed>
    <deployedseed>941119460</deployedseed>
    <deployedseed>1115243276</deployedseed>
    <deployedseed>723565334</deployedseed>
    <deployedseed>1527361744</deployedseed>
    <deployedseed>454050483</deployedseed>
    <deployedseed>2049523733</deployedseed>
    <deployedseed>788738515</deployedseed>
    <deployedseed>1087944574</deployedseed>
    <deployedseed>1492954551</deployedseed>
    <deployedseed>430566628</deployedseed>
    <deployedseed>1366078114</deployedseed>
    <deployedseed>1703571741</deployedseed>
    <deployedseed>1878977049</deployedseed>
    <deployedseed>1068016871</deployedseed>
    <deployedseed>1132010128</deployedseed>
    <deployedseed>1586178936</deployedseed>
    <deployedseed>652392454</deployedseed>
    <deployedseed>1059305525</deployedseed>
    <deployedseed>453550191</deployedseed>
    <deployedseed>112496485</deployedseed>
    <deployedseed>1066041139</deployedseed>
    <deployedseed>1166574393</deployedseed>
    <deployedseed>394418983</deployedseed>
    <deployedseed>446193988</deployedseed>
    <deployedseed>1956912005</deployedseed>
    <deployedseed>1945588589</deployedseed>
    <deployedseed>1008900454</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>A</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
